---
title: "Бинарная логистическая регрессия"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(titanic)
library(dplyr)
data("titanic_train")
titanic<-titanic_train
titanic_clean <- titanic |> 
  select(Survived, Sex, Pclass, Age, Fare, SibSp, Parch, Embarked) |>
  mutate(
    Survived = factor(Survived),
    Sex      = factor(Sex),
    Pclass   = factor(Pclass),
    Embarked = factor(Embarked)
  ) |>
  mutate(
    Age = ifelse(is.na(Age), median(Age, na.rm = TRUE), Age)
  )
model_sex <- glm(Survived ~ Sex,
                 data   = titanic_clean,
                 family = binomial)
model_sex_pclass_age <- glm(Survived ~ Sex + Pclass + Age,
                            data   = titanic_clean,
                            family = binomial)
P_hat <- predict(model_sex_pclass_age,
                 type = "response")

pred_class <- ifelse(P_hat >= 0.5, 1, 0)

y_true <- as.numeric(titanic_clean$Survived) - 1
conf_mat <- table(Predicted = pred_class, Actual = y_true)

knitr::opts_chunk$set(echo=FALSE, warning=FALSE, message = FALSE)
```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Введение

Мы продолжаем погружаться в регрессионный анализ, и данный тьюториал посвящен логистической регрессии.

**Логистическая регрессия** относится к моделям бинарного выбора. Что это значит? Это значит, что зависимая переменная может принимать только два значения (например, 0/1, «да/нет», «сдал экзамен/не сдал экзамен»), а модель описывает, как вероятность одного из этих двух вариантов зависит от набора предикторов.


Примеры типичных задач:

- Участие в выборах: голосовал / не голосовал как функция возраста, образования, доверия к институтам, политических установок.
- Трудовая занятость: работает / не работает в зависимости от пола, образования, семейного статуса, наличия детей.
- Проживание в городе или селе, миграция: остался / уехал, как функция дохода, возраста, социального капитала.
- Принятие рискованного решения, безопасность: участвует / не участвует в проекте, принимает на себя какие-то риски / выбирает безопасное поведение и т.п.

На первый взгляд, если мы столкнулись с одним из вопросов, сформулированных выше, мы можем легко обойтись теми инструментами регрессионного анализа, которые нам уже знакомы: взять в качестве зависимой переменной бинарную переменную и провести множественный регрессионный анализ.

В примере с трудовой занятостью мы могли бы собрать данные про $n$ индивидов, пришедших на биржу труда, и в качестве зависимой переменной выбрать бинарную переменную, равную единице для тех, кто устроился на работу, и нулю для тех, кто не смог найти подходящее место

В такой модели мы можем считать, что вероятность наступления события (например, трудоустройства) линейно зависит от некоторого фактора $x$ (например, от заработной платы индивида). Тогда такую модель называют линейной моделью вероятности и записывают следующим образом:

$$
p_i=P(y_i=1)=\beta_1+\beta_2x_i,
$$
где $Y_i$ -- событие, которое равно 1, если оно наступило, и 0 -- если нет, $x_i$ -- фактор, влияющий на наступления данного события, и  $P(y_i=1)$ -- вероятность того, что событие наступит.

Проблематичность использования данной модели заключается в том, что предсказанные значения вероятности могут быть отрицательными или превышать единицу. Например, если мы пытаемся спрогнозировать вероятность сдачи зачета от количества потраченного на подготовку времени, то может оказаться, что вероятность при 2-х часах будет отрицательной, а при 100 часах, превысит единицу.Вот [здесь](https://books.econ.msu.ru/Introduction-to-Econometrics/chap10/10.1/) этот пример рассмотрен подробнее. 

Вызывает сомнения и предположение о том, что вероятность успеха зависит от объясняющей переменной линейно. Вряд ли увеличение времени подготовки с 20 до 30 часов должно приводить к такому же изменению вероятности сдачи зачета, как и увеличение с 1000 до 1010 часов.

Поэтому в прикладных исследованиях вместо линейной модели вероятности обычно используют одну из двух альтернатив: логит-модель или пробит-модель.

Логит‑модель — это другая привычная русская формулировка логистической регрессии.


В логит‑модели мы описываем не вероятность, а **логит вероятности** события через линейную комбинацию признаков:

$$
\text{logit}(p) = \log\left(\frac{p}{1 - p}\right) = \beta_0 + \beta_1 x_1 + \dots + \beta_k x_k.
$$

- $p = P(Y = 1 \mid X)$ — вероятность «успеха», события, выбора варианта 1.  
- Левая часть -— логарифм **шансов** события (odds) -- отношения вероятности наступления события к вероятности того, что это событие не наступит $(1-p)$.  
- Правая часть —- обычная линейная регрессия по предикторам $x_1, \dots, x_k$.

Из этого выражения  можно получить саму вероятность (формулу логистической функции):

$$
p = \frac{1}{1 + \exp\big(-(\beta_0 + \beta_1 x_1 + \dots + \beta_k x_k)\big)}.
$$

Таким образом, логит‑преобразование переводит вероятность события $p$ в логарифм шансов $\log\big(p/(1-p)\big)$, который может принимать любые значения от $-\infty$ до $+\infty$. Это позволяет строить линейную модель на шкале логита, а затем через логистическую функцию автоматически возвращать результат в интервал вероятностей $[0;1]$ и интерпретировать коэффициенты как изменения логита и отношений шансов.

Например, пусть вероятность события $p = 0{.}8$. Тогда шансы (odds) равны $0{.}8/(1-0{.}8) = 4$: «шансы 4 к 1 в пользу события». Логит — это логарифм этих шансов: $\log(4) \approx 1{.}39$. Если же $p = 0{.}2$, то шансы $0{.}2/0{.}8 = 0{.}25$, а логит $\log(0{.}25) \approx -1{.}39$ — симметричен предыдущему. 

Именно на этой логит‑шкале логистическая регрессия задает прямую $\text{logit}(p) = \beta_0 + \beta_1 x_1 + \dots$, а уже затем эта прямая превращается в S‑образную кривую вероятностей.

**S‑образная** кривая, в свою очередь, при больших отрицательных значениях дает значения, близкие к 0, при больших положительных — близкие к 1, а в середине возрастает наиболее быстро.

Формула логистической функции:
$$
p(x) = \frac{1}{1 + \exp(-(\beta_0 + \beta_1 x))}
$$

Простейший график при $\beta_0 = 0$, $\beta_1 = 1$:

$$
p(x) = \frac{1}{1 + e^{-x}}
$$

```{r}
logistic <- function(x) {
  1 / (1 + exp(-x))
}

# диапазон x
x <- seq(-10, 10, length.out = 400)

# значения p(x)
p <- logistic(x)

# график
plot(
  x, p, type = "l",
  xlab = "x",
  ylab = "p(x)",
  main = "Логистическая S-образная кривая"
)
abline(h = 0.5, v = 0, lty = 2, col = "grey")
```

**Пробит‑модель** — это альтернативная логит‑модели регрессия для бинарного выбора, в которой вместо логистической функции используется стандартная нормальная функция распределения. 

В логите мы моделируем логарифм шансов события, а в пробите моделируем **z‑балл** латентной непрерывной переменной: вероятность события задается как $\Phi(\beta_0 + \beta_1 x_1 + \dots)$, где $\Phi$ — стандартная нормальная кумулятивная функция. На практике обе модели решают одну и ту же задачу (моделирование $P(Y=1)$ по предикторам) и часто дают очень похожие вероятности, но в пробит‑модели коэффициенты интерпретируются через нормальное распределение, а не через отношения шансов; выбор между логитом и пробитом обычно определяется традицией в конкретной области или удобством интерпретации.

Изучение пробит-модели, к сожалению, выходит за рамки нашего курса, и с ней желающие могут познакомиться самостоятельно. Например, [вот здесь](https://books.econ.msu.ru/Introduction-to-Econometrics/chap10/10.4/).

## Предварительные вопросы

```{r regression-quiz, echo=FALSE}
quiz(
  caption = "Тест: Логистическая регрессия",

  question(
    "1. Для чего используется логистическая регрессия?",
    answer("Для прогнозирования непрерывной зависимой переменной (например, доход)", correct = FALSE,
           message = "Это описание линейной регрессии."),
    answer("Для моделирования вероятности бинарного исхода (0/1)", correct = TRUE,
           message = "Верно: логистическая регрессия моделирует $P(Y = 1)$."),
    answer("Только для кластеризации наблюдений", correct = FALSE,
           message = "Кластеризация — это метод обучения без учителя."),
    answer("Для оценки корреляции между двумя количественными переменными", correct = FALSE),
    allow_retry = TRUE
  ),
 question_radio(
    "2. В логистической регрессии зависимая переменная является...",
    answer("Непрерывной количественной"),
    answer("Порядковой"),
    answer("Бинарной (категориальной с двумя исходами)", correct = TRUE),
    answer("Счётной"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),
 
  question(
    "3. Какой вид связи между предикторами и логитом вероятности предполагает базовая логистическая регрессия?",
    answer("Линейная связь между предикторами и самой вероятностью", correct = FALSE),
    answer("Линейная связь между предикторами и логитом (log-odds)", correct = TRUE),
    answer("Нелинейная связь любого общего вида без ограничений", correct = FALSE),
    answer("Связь не задаётся явно", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "4. Как выглядит логистическая функция для вероятности $P(Y = 1 | X)$?",
    answer("$P = \\beta_0 + \\beta_1X$", correct = FALSE),
    answer("$P = exp(\\beta_0 + \\beta_1X)$", correct = FALSE),
    answer("$P = 1 / (1 + exp(-(\\beta_0 + \\beta_0X)))$", correct = TRUE),
    answer("$P = (\\beta_0 + \\beta_1X)^2$", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "5. Какой R-код корректно оценивает логистическую регрессию с бинарным исходом y и предиктором x?",
    answer("glm(y ~ x, family = gaussian)", correct = FALSE,
           message = "Здесь указано нормальное семейство, это не логистическая регрессия."),
    answer("glm(y ~ x, family = binomial)", correct = TRUE,
           message = "Да, логистическая регрессия через glm с binomial."),
    answer("lm(y ~ x)", correct = FALSE,
           message = "lm() строит линейную, не логистическую регрессию."),
    answer("glm(y ~ x, family = poisson)", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "6. Как интерпретируется коэффициент $\\beta_1$ в простой логистической регрессии с одним предиктором $X$?",
    answer("Изменение средней $Y$ при увеличении $X$ на 1", correct = FALSE),
    answer("Изменение логита (log-odds) события $Y=1$ при увеличении $X$ на 1", correct = TRUE,
           message = "Верно: $\\beta_1$ — приращение логита на 1 единицу $X$."),
    answer("Изменение стандартного отклонения $Y$ при увеличении $X$ на 1", correct = FALSE),
    answer("Коэффициент корреляции между $X$ и $Y$", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
  "7. Как получить отношение шансов (odds ratio) для коэффициента предиктора X из модели model <- glm(y ~ x, family = binomial)?",
  answer("coef(model)[\"x\"]", correct = FALSE,
         message = "Это логарифм отношения шансов, а не само отношение шансов."),
  answer("exp(coef(model)[\"x\"])", correct = TRUE,
         message = "Да, экспонента коэффициента дает отношение шансов."),
  answer("log(coef(model)[\"x\"])", correct = FALSE),
  answer("1 / coef(model)[\"x\"]", correct = FALSE),
  allow_retry = TRUE
),

  question(
    "8. Как трактовать отношение шансов (odds ratio) = 2 для предиктора X?",
    answer("При увеличении X на 1 вероятность события удваивается при любых значениях X", correct = FALSE,
           message = "Удваиваются шансы, а не вероятность в абсолютном выражении."),
    answer("При увеличении X на 1 шансы события Y=1 умножаются на 2", correct = TRUE,
           message = "Да, это стандартная интерпретация OR."),
    answer("При увеличении X на 1 логит уменьшается в 2 раза", correct = FALSE),
    answer("Это означает отсутствие эффекта X", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "9. Какой параметр в `predict()` следует указать, чтобы получить предсказанные вероятности для логистической регрессии в R?",
    answer("type = \"link\"", correct = FALSE,
           message = "\"link\" возвращает логиты, не вероятности."),
    answer("type = \"response\"", correct = TRUE,
           message = "Да, \"response\" даёт вероятности P(Y=1)."),
    answer("type = \"prob\"", correct = FALSE),
    answer("type = \"class\"", correct = FALSE),
    allow_retry = TRUE
  ),
question(
  "10. Как обычно формируется бинарный прогноз (0/1) на основе предсказанной вероятности P(Y=1)?",
  answer("Всегда присваивать 1, так как модель логистическая", correct = FALSE),
  answer("Использовать порог, например 0.5: если P >= 0.5, прогноз = 1, иначе 0", correct = TRUE,
         message = "Наиболее типичный подход — фиксировать порог, часто 0.5."),
  answer("Использовать случайные числа вместо порога", correct = FALSE),
  answer("Всегда присваивать 0, чтобы снизить количество ошибок", correct = FALSE),
  allow_retry = TRUE
),
   question(
    "11. Что представляет собой матрица ошибок (confusion matrix) в задаче бинарной классификации?",
    answer("Матрицу корреляций между предикторами", correct = FALSE),
    answer("Таблицу истинных и предсказанных классов (TP, FP, TN, FN)", correct = TRUE,
           message = "Да, это матрица ошибок классификатора."),
    answer("Матрицу ковариаций параметров модели", correct = FALSE),
    answer("Таблицу значений логарифма правдоподобия", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "12. Какой из критериев часто используется для сравнения логистических моделей в R?",
    answer("Среднеквадратическая ошибка (MSE)", correct = FALSE),
    answer("AIC (информационный критерий Акаике)", correct = TRUE,
           message = "Да, AIC широко используют для сравнения моделей."),
    answer("$R^2$ из линейной регрессии без модификаций", correct = FALSE),
    answer("Коэффициент вариации", correct = FALSE),
    allow_retry = TRUE
  ),

  question(
    "13. Какое из нижеперечисленных утверждений ближе всего к предположениям логистической регрессии?",
    answer("Нормальность распределения зависимой переменной", correct = FALSE),
    answer("Линейность по логиту, независимость наблюдений, отсутствие сильной мультиколлинеарности", correct = TRUE,
           message = "Это ключевые предположения: линейность по логиту, независимость и разумная мультиколлинеарность."),
    answer("Требуется гомоскедастичность ошибок как в МНК", correct = FALSE),
    answer("Никаких предположений не требуется", correct = FALSE),
    allow_retry = TRUE
  ),
 question(
  "14. Как можно обобщить логистическую регрессию на случай зависимой переменной с более чем двумя категориями (мультиклассовая классификация)?",
  answer("Использовать набор бинарных моделей «один против всех» (one-vs-rest) или мультиномиальную логистическую регрессию", correct = TRUE),
  answer("Применять линейную регрессию отдельно для каждого класса", correct = FALSE),
  answer("Использовать только деревья решений", correct = FALSE),
  answer("Никак, логистическая регрессия применима только к бинарным исходам", correct = FALSE),
  allow_retry = TRUE
),
question(
  "15. Какой подход уместен при сильном дисбалансе классов в задаче бинарной логистической регрессии?",
  answer("Оценивать только точность (accuracy), этого достаточно", correct = FALSE),
  answer("Использовать дополнительные метрики (precision, recall, F1), изменять порог вероятности, балансировать выборку (oversampling/undersampling)", correct = TRUE),
  answer("Всегда удалять часть наблюдений из большего класса случайным образом без анализа", correct = FALSE),
  answer("Ничего не менять, дисбаланс классов логистической регрессии не мешает", correct = FALSE),
  allow_retry = TRUE
)
)
```

## Практические задания

Для выполнения упражнений по данной теме мы возьмем один из классических наборов данных, посвященных легендарной истории корабля «Титаник».


### Исторический контекст

14–15 апреля 1912 года произошло крушение пассажирского лайнера «Титаник», считавшегося на тот момент вершиной технического прогресса и «практически непотопляемым» судном. В катастрофе погибло более полутора тысяч человек, и уже тогда были сформированы подробные документальные записи: списки пассажиров, классы билетов, информация о посадке, данных о спасшихся и погибших. Позже эти архивы начали систематизировать и оцифровывать, что привело к появлению машиночитаемых таблиц с данными о пассажирах.

![Рис. 1. Кадр из фильма «Титаник» (1997)](images\titanic1.jpg){height=200px} ![Рис. 2. Титаник перед отплытием](images\titanic2.jpg){height=200px}

Ключевой момент популяризации датасета Titanic в аналитике и машинном обучении связан с запуском в 2012 году конкурса «Titanic: Machine Learning from Disaster» на платформе Kaggle. Организаторы предложили участникам построить модель, предсказывающую, выжил ли пассажир, на основе нескольких демографических и социальных характеристик. С этого момента Titanic стал де‑факто «классическим» учебным набором данных для введения в анализ данных, статистическое моделирование и бинарную классификацию.


Датасет Titanic представляет собой таблицу, где каждая строка — это пассажир, а столбцы — его характеристики.

- **Целевая переменная**  
  - `Survived` — выжил (1) / не выжил (0).

- **Социально‑демографические и контекстные признаки**  
  - `Pclass` — класс билета (1, 2, 3), часто трактуется как индикатор социального статуса.  
  - `Sex` — пол пассажира.  
  - `Age` — возраст.  
  - `SibSp` — количество братьев/сестер и супругов на борту.  
  - `Parch` — количество родителей и детей на борту.  
  - `Fare` — стоимость билета.  
  - `Embarked` — порт посадки (C, Q, S).  
  
Набор содержит 891 наблюдения. Датасет интересен тем, что объединяет числовые и категориальные переменные, имеет пропуски (например, в `Age` или `Cabin`) и хорошо подходит для демонстрации полного цикла работы: от очистки и кодирования признаков до построения логистической регрессии и оценки качества классификации.


### Упражнение 1 — Просмотр данных

Как всегда мы начинаем с просмотра данных. Загрузите набор `titanic`
и посмотрите на значения.

```{r ex1, exercise = TRUE, eval=TRUE}
titanic
```

```{r ex1-hint-1, eval = FALSE}
titanic
```


```{r ex1-solution, include = FALSE}
titanic
```

```{r ex1-code-check}
grade_code()
```


### Упражнение 2 — Базовая подготовка данных Titanic

Вам необходимо:

1. Создать объект `titanic_clean` на основе данных `titanic`.  
2. Оставить только переменные: `Survived, Sex, Pclass, Age, Fare, SibSp, Parch, Embarked`.  
3. Преобразовать переменные `Survived, Sex, Pclass, Embarked` в фактор.  
4. Заменить пропуски в `Age` на медиану возраста.

```{r ex2, exercise = TRUE, eval = TRUE}
titanic_clean <- titanic |> 
  select() |>
  mutate(
    Survived = ,
    Sex      = ,
    Pclass   = ,
    Embarked = 
  ) |>
  mutate(
    Age = ifelse(is.na(Age), , Age)
  )

str(titanic_clean)
```

```{r ex2-hint-1, eval = FALSE}
titanic_clean <- titanic |> 
  select(Survived, Sex, Pclass, Age, Fare, SibSp, Parch, Embarked) |>
  mutate(
    Survived = factor(Survived),
    Sex      = factor(Sex),
    Pclass   = factor(Pclass),
    Embarked = factor(Embarked)
  ) |>
  mutate(
    Age = ifelse(is.na(Age), median(Age, na.rm = TRUE), Age)
  )

str(titanic_clean)
```

```{r ex2-solution, include = FALSE}
titanic_clean <- titanic |> 
  select(Survived, Sex, Pclass, Age, Fare, SibSp, Parch, Embarked) |>
  mutate(
    Survived = factor(Survived),
    Sex      = factor(Sex),
    Pclass   = factor(Pclass),
    Embarked = factor(Embarked)
  ) |>
  mutate(
    Age = ifelse(is.na(Age), median(Age, na.rm = TRUE), Age)
  )

str(titanic_clean)
```

```{r ex2-code-check}
grade_code()
```

```{r ex2-q}
question_text(
  "Какие переменные из выбранных, по Вашему мнению, могут сильнее всего влиять на вероятность выживания?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Попробуйте ещё раз уточнить свой ответ",
  incorrect = NULL
)
```


### Упражнение 3 — Таблица выживаемости по полу

Построим таблицу выживаемости по полу пассажиров и оценим долю выживших в каждой группе.

```{r echo=FALSE}
library(dplyr)
```

Вам необходимо:

1. Построить таблицу частот `Survived × Sex`.  
2. Посчитать долю выживших для мужчин и женщин (по столбцам).

```{r ex3, exercise = TRUE, eval=TRUE}
tab <- table( , )

prop_tab <- prop.table(tab, margin = )

tab
prop_tab
```

```{r ex3-hint-1, eval = FALSE}
tab <- table(titanic_clean$Survived, titanic_clean$Sex)

prop_tab <- prop.table(tab, margin = 2)

tab
prop_tab
```

```{r ex3-solution, include = FALSE}
tab <- table(titanic_clean$Survived, titanic_clean$Sex)

prop_tab <- prop.table(tab, margin = 2)

tab
prop_tab
```


```{r ex3-code-check}
grade_code()
```

```{r ex3-q}
question_text(
  "Что можно сказать о разнице в долях выживших между мужчинами и женщинами?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Попробуйте описать, у кого шансы выжить были выше.",
  incorrect = NULL
)
```


### Упражнение 4 — Базовая логит‑модель с полом

Оценим логистическую регрессию, в которой вероятность выживания зависит только от пола.

Вам необходимо:

1. Оценить модель `Survived ~ Sex` с помощью `glm()` и `family = binomial`.  
2. Сохранить модель в объект `model_sex`.  
3. Просмотреть результат `summary(model_sex)`.

```{r ex4, exercise = TRUE, eval=TRUE}
model_sex <- glm( ,
                 data   = ,
                 family = )

summary(model_sex)
```

```{r ex4-hint-1, eval = FALSE}
model_sex <- glm(Survived ~ Sex,
                 data   = titanic_clean,
                 family = binomial)

summary(model_sex)
```

```{r ex4-solution, include = FALSE}
model_sex <- glm(Survived ~ Sex,
                 data   = titanic_clean,
                 family = binomial)

summary(model_sex)
```


```{r ex4-code-check}
grade_code()
```

```{r ex4-q}
question_text(
  "Каков знак коэффициента при Sexmale и что он означает для шансов выживания мужчин по сравнению с женщинами?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Укажите знак и опишите, увеличивает он или уменьшает шансы выживания.",
  incorrect = NULL
)
```


### Упражнение 5 — Отношение шансов для мужчин

Посчитаем отношение шансов выживания мужчин по сравнению с женщинами на основе модели `model_sex`.

Вам необходимо:

1. Извлечь коэффициент при `Sexmale` из `model_sex`.  
2. Преобразовать его с помощью `exp()` в отношение шансов.  
3. Вывести полученное значение.

```{r ex5, exercise = TRUE, eval=TRUE}
coef(model_sex)

or_male <- exp(coef(model_sex)[ ])

or_male
```

```{r ex5-hint-1, eval = FALSE}
coef(model_sex)

or_male <- exp(coef(model_sex)["Sexmale"])

or_male
```

```{r ex5-solution, include = FALSE}
coef(model_sex)

or_male <- exp(coef(model_sex)["Sexmale"])

or_male
```


```{r ex5-code-check}
grade_code()
```

```{r ex5-q}
question_text(
  "Как можно интерпретировать найденное отношение шансов для мужчин относительно женщин?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Опишите, во сколько раз шансы выживания мужчин больше или меньше, чем у женщин.",
  incorrect = NULL
)
```


### Упражнение 6 — Добавляем класс и возраст

Теперь расширим модель, добавив класс билета и возраст пассажира.

Вам необходимо:

1. Оценить модель `Survived ~ Sex + Pclass + Age`.  
2. Сохранить модель в объект `model_sex_pclass_age`.  
3. Посмотреть `summary(model_sex_pclass_age)`.

```{r ex6, exercise = TRUE, eval=TRUE}
model_sex_pclass_age <- glm( ,
                            data   = ,
                            family = )

summary(model_sex_pclass_age)
```

```{r ex6-hint-1, eval = FALSE}
model_sex_pclass_age <- glm(Survived ~ Sex + Pclass + Age,
                            data   = titanic_clean,
                            family = binomial)

summary(model_sex_pclass_age)
```

```{r ex6-solution, include = FALSE}
model_sex_pclass_age <- glm(Survived ~ Sex + Pclass + Age,
                            data   = titanic_clean,
                            family = binomial)

summary(model_sex_pclass_age)
```


```{r ex6-code-check}
grade_code()
```

```{r ex6-q}
question_text(
  "Как меняется эффект пола на выживание после добавления Pclass и Age (по сравнению с моделью только с Sex)?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Сравните знак и величину коэффициента Sexmale в двух моделях.",
  incorrect = NULL
)
```


### Упражнение 7 — Сравнение моделей по AIC

Сравним качество простой и расширенной моделей с помощью AIC.

Вам необходимо:

1. Вычислить AIC для `model_sex` и `model_sex_pclass_age`.  
2. Сравнить значения AIC и определить, какая модель лучше по этому критерию.

```{r ex7, exercise = TRUE, eval=TRUE}
AIC()

AIC()
```

```{r ex7-hint-1, eval = FALSE}
AIC(model_sex)

AIC(model_sex_pclass_age)
```

```{r ex7-solution, include = FALSE}
AIC(model_sex)

AIC(model_sex_pclass_age)
```


```{r ex7-code-check}
grade_code()
```

```{r ex7-q}
question_text(
  "Какая из двух моделей предпочтительнее по AIC и почему?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Вспомните, как интерпретируется AIC: больше или меньше — лучше?",
  incorrect = NULL
)
```


### Упражнение 8 — Предсказанные вероятности и бинарный прогноз

Построим предсказанные вероятности выживания и бинарный прогноз при пороге 0.5.

Вам необходимо:

1. Получить предсказанные вероятности `P_hat` из `model_sex_pclass_age` с `type = "response"`.  
2. Построить бинарный прогноз `pred_class` (1, если `P_hat >= 0.5`, иначе 0).  
3. Преобразовать `Survived` в `y_true` (0/1) и посчитать долю верно классифицированных наблюдений.

```{r ex8, exercise = TRUE, eval=TRUE}
P_hat <- predict( ,
                 type = )

pred_class <- ifelse( , 1, 0)

y_true <- as.numeric(titanic_clean$Survived) - 

mean(pred_class == y_true)
```

```{r ex8-hint-1, eval = FALSE}
P_hat <- predict(model_sex_pclass_age,
                 type = "response")

pred_class <- ifelse(P_hat >= 0.5, 1, 0)

y_true <- as.numeric(titanic_clean$Survived) - 1

mean(pred_class == y_true)
```

```{r ex8-solution, include = FALSE}
P_hat <- predict(model_sex_pclass_age,
                 type = "response")

pred_class <- ifelse(P_hat >= 0.5, 1, 0)

y_true <- as.numeric(titanic_clean$Survived) - 1

mean(pred_class == y_true)
```


```{r ex8-code-check}
grade_code()
```

```{r ex8-q}
question_text(
  "Как Вы интерпретируете полученную долю правильно классифицированных наблюдений?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Опишите, насколько хорошо модель различает выживших и невыживших при пороге 0.5.",
  incorrect = NULL
)
```


### Упражнение 9 — Матрица ошибок (confusion matrix)

Построим матрицу ошибок для бинарного прогноза и фактических значений.

Вам необходимо:

1. Использовать уже полученные `pred_class` и `y_true`.  
2. Построить таблицу `table(Predicted = pred_class, Actual = y_true)`.

```{r ex9, exercise = TRUE, eval=TRUE}
conf_mat <- table( , )

conf_mat
```

```{r ex9-hint-1, eval = FALSE}
conf_mat <- table(Predicted = pred_class, Actual = y_true)

conf_mat
```

```{r ex9-solution, include = FALSE}
conf_mat <- table(Predicted = pred_class, Actual = y_true)

conf_mat
```


```{r ex9-code-check}
grade_code()
```

```{r ex9-q}
question_text(
  "Как Вы интерпретируете элементы матрицы ошибок (TP, FP, TN, FN) в контексте выживания на Титанике?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Опишите хотя бы по одному примеру: что означает FP или FN в этой задаче.",
  incorrect = NULL
)
```


### Упражнение 10 — Семейный размер как предиктор (без предсказаний)

В этом упражнении сделаем акцент на **интерпретации** нового предиктора, а не на предсказании.

Вам необходимо:

1. Создать новый объект `titanic_clean2`, добавив переменную `FamilySize = SibSp + Parch + 1`.  
2. Оценить модель `Survived ~ Sex + Pclass + Age + FamilySize`.  
3. Сохранить модель в объект `model_family`.  
4. Просмотреть `summary(model_family)` и сравнить AIC с `model_sex_pclass_age`.

```{r ex10, exercise = TRUE, eval=TRUE}
titanic_clean2 <- titanic_clean |> 
  mutate(FamilySize = )

model_family <- glm( ,
                    data   = ,
                    family = )

AIC(model_sex_pclass_age)
AIC(model_family)

summary(model_family)
```

```{r ex10-hint-1, eval = FALSE}
titanic_clean2 <- titanic_clean |> 
  mutate(FamilySize = SibSp + Parch + 1)

model_family <- glm(Survived ~ Sex + Pclass + Age + FamilySize,
                    data   = titanic_clean2,
                    family = binomial)

AIC(model_sex_pclass_age)
AIC(model_family)

summary(model_family)
```

```{r ex10-solution, include = FALSE}
titanic_clean2 <- titanic_clean |> 
  mutate(FamilySize = SibSp + Parch + 1)

model_family <- glm(Survived ~ Sex + Pclass + Age + FamilySize,
                    data   = titanic_clean2,
                    family = binomial)

AIC(model_sex_pclass_age)
AIC(model_family)

summary(model_family)
```

```{r ex10-code-check}
grade_code()
```


```{r ex10-q}
question_text(
  "Как можно интерпретировать знак и величину коэффициента при FamilySize в модели model_family?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Опишите, как изменение размера семьи на 1 человека влияет на логит (или шансы) выживания.",
  incorrect = NULL
)
```



```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
