---
title: "Основы статистического вывода и тестирование гипотез исследования"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(ggstatsplot)
knitr::opts_chunk$set(echo=FALSE)
library(datarium)
library(rstatix)
data("jobsatisfaction")
js<-jobsatisfaction
```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Предварительные тестовые задания

### Вопрос 1

```{r q1}
question_radio("Какая формулировка нулевой гипотезы наиболее уместна для двухвыборочного t-теста в R при сравнении средних значений роста у мужчин и женщин?",
  answer("Распределение роста мужчин является нормальным"),
  answer("Средний рост мужчин больше среднего роста женщин"),
  answer("Средний рост мужчин равен среднему росту женщин", correct = TRUE),
  answer("Средний рост мужчин отличается от определенного порогового значения"),
  incorrect = "Нужно учитывать не только характер распределения, но и все условия",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Вопрос 2

```{r q2}
question_radio("В R вы получили результат одновыборочного t-теста: $t = 2.5$, $df = 19$, $p-value = 0.021$ при уровне значимости $\\alpha= 0.05$. Какой наиболее корректный вывод?",
  answer("Среднее значение не отличается от гипотетического, и нулевая гипотеза принимается"),
  answer("Среднее значение статистически значимо отличается от гипотетического, и нулевая гипотеза отвергается", correct = TRUE),
  answer("Среднее значение сильно отличается от гипотетического, так как t-статистика положительна"),
  answer("Статистических оснований для вывода нет, так как p-value слишком мало"),
  incorrect = "Обратите внимание на p-value",
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Вопрос 3

```{r q3}
question_radio("Какое утверждение о p-значении в контексте вывода $p-value = 0.08$ в R при $\\alpha= 0.05$ наиболее корректно?",
  answer("При истинной нулевой гипотезе вероятность получить столь же экстремальные данные составляет 8%", correct = TRUE),
  answer("С вероятностью 8% мы сделаем ошибку первого рода"),
  answer("Вероятность того, что альтернативная гипотеза верна, равна 92%"),
  answer("Вероятность того, что нулевая гипотеза верна, равна 8%"),
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Вопрос 4

```{r q4}
question_radio("Какое утверждение о 95% доверительном интервале для среднего, построенном в R с помощью t.test, является наиболее корректным?",
  answer("Если многократно повторить выборочное исследование и строить интервалы, то примерно 95% таких интервалов накроют истинное среднее", correct = TRUE),
  answer("Истинное среднее с точностью 95% равно середине интервала"),
  answer("Ширина интервала зависит только от размера выборки и не зависит от разброса данных"),
  answer("Вероятность того, что истинное среднее попадет в конкретный построенный интервал, равна 95%"),
  allow_retry = TRUE,
  random_answer_order = TRUE
)
```

### Вопрос 5

```{r q5}
question_radio("Как центральная предельная теорема (ЦПТ) обосновывает использование нормального приближения для средней выборки в R при достаточно большом размере выборки?",
  answer("Распределение исходных данных становится нормальным при любом размере выборки"),
  answer("ЦПТ утверждает, что все статистики имеют нормальное распределение при большом объёме выборки"),
  answer("Распределение выборочной средней стремится к нормальному независимо от исходного распределения при увеличении размера выборки
", correct = TRUE),
  answer("Выборочная дисперсия всегда имеет точное нормальное распределение"),
  allow_retry = TRUE,
  random_answer_order = TRUE,
correct="Центральная предельная теорема утверждает, что при достаточно большом размере выборки $n$ распределение стандартизированной выборочной средней $\\bar{X}$ стремится к нормальному распределению с параметрами 
$N(\\mu,\\sigma^2/n)$, даже если исходные данные имеют ненормальное распределение, при выполнении некоторых условий (например, конечная дисперсия). Это объясняет, почему в R множество процедур (например, построение доверительных интервалов для среднего через `t.test` или нормальное приближение в `prop.test`) становятся более точными при большом $n$. На практике при $n\\geq30$ часто считают, что нормальное приближение уже достаточно хорошее, хотя точные критерии зависят от формы исходного распределения. ЦПТ позволяет использовать известные свойства нормального распределения для построения интервалов и проверки гипотез относительно средних, опираясь на асимптотическое поведение $\\bar{X}$")
```

### Вопрос 6

```{r q6, warning=FALSE}
question_radio("Вы моделируете распределение выборочной средней в R с помощью симуляции: многократно берете выборки n = 50 из одной и той же популяции и для каждой считаете среднее. Какой основной вывод позволяет сделать такая симуляция в контексте ЦПТ?",
  answer("Она позволяет вычислить точное p-значение для любого статистического теста", 
         message = "Cимуляция визуализирует распределение, а не p-value"),
  answer("Она наглядно показывает, что распределение выборочной средней близко к нормальному", correct = TRUE, 
         message = "**Правильно!** Такая симуляция строит гистограмму многих $\\bar{X}$ (n=50) → видно приближение к нормальному (ЦПТ).

**R-пример:**
```r (eval=T)
sample_means <- replicate(1000, mean(rexp(50)))  
hist(sample_means, freq=FALSE)
curve(dnorm(x, 1, 1/sqrt(50)), col='red', add=TRUE, lwd=2)
На гистограмме видна красная нормальная кривая, хотя данные сгенерированы из экспоненциального распределения!"),
answer("Она показывает, что выборочная дисперсия всегда постоянна",
message = "Дисперсия $\\bar{X} = \\sigma^2/n$, уменьшается с $n$"),
answer("Она доказывает, что исходное распределение данных всегда нормальное",
message = "Предельная теорема работает для любых распределений с конечной дисперсией"),
random_answer_order = TRUE,
allow_retry = TRUE
)
```

### Вопрос 7

```{r q7}
question_radio("В R вы использовали функцию t.test(x, mu = 0) и получили широкие доверительные интервалы. Какое объяснение этому наиболее обосновано?",
  answer("p-значение близко к нулю, поэтому интервал автоматически расширяется", 
         message = "Нет. Малое p-value говорит о значимости, но ширина CI зависит от разброса/n"),
  answer("Широкий интервал означает, что данные обязательно не нормально распределены", 
         message = "Нет. Нормальность влияет слабо. Главное — разброс данных и размер выборки"),
  answer("Ширина доверительного интервала не зависит от данных, а задаётся только уровнем значимости", 
         message = "Ошибка. Формула: CI = mean ± t × (sd/√n). Зависит от sd и n"),
  answer("Большой разброс данных или маленький размер выборки увеличивают стандартную ошибку, расширяя доверительный интервал", correct = TRUE, 
         message = "Правильно. Ширина $CI = 2 \\times t \\times SE$, где $SE = sd(x)/\\sqrt{n}$

CI расширяется при:
- Большом sd(x) — большой разброс  
- Маленьком n — $\\sqrt{n}$ в знаменателе
- Малых df — t больше

R-пример:
```r
x1 <- rnorm(10, 0, 5)    # n=10, sd=5 → Широкий CI
x2 <- rnorm(100, 0, 1)   # n=100, sd=1 → Узкий CI
t.test(x1, mu=0)
```"),
  allow_retry = TRUE
)
```

### Вопрос 8

```{r q8}
question_radio("Допустим, Вам нужно с помощью симуляции оценить в R распределение статистики t для выборки малого размера из нормальной популяции. Какой подход наиболее уместен?",
  answer("Многократно генерировать выборки из нормального распределения, вычислять t-статистику и строить гистограмму полученных значений", correct = TRUE, 
         message = "Правильно. Сделаем симуляцию: сгенерируем $n=20$ из $N(\\mu,\\sigma)$, затем посчитаем $t =(\\bar{X}-\\mu)/(s\\sqrt{n})$ и сделаем гистограмму через `hist()`.
Пример кода:
```r
set.seed(123); n<-20; sims<-5000; mu<-0
t_stats <- replicate(sims, {
  x <- rnorm(n, mu, 1)
  x_bar <- mean(x)
  s <- sd(x)
  t.stat <- (x_bar - mu)/(s/sqrt(n))
  t.stat
})
hist(t_stats, breaks=30, freq=FALSE, main='t-статистика (симуляция)')
curve(dt(x, df=n-1), add=TRUE, col='red', lwd=2)
```"),
  answer("Один раз построить выборку и использовать только теоретическое распределение Стьюдента", 
         message = "Нет. Одна выборка не покажет распределение t. Нужно много симуляций"),
  answer("Генерировать данные из равномерного распределения и считать их t-статистикой", 
         message = "Нет. t-статистика требует нормальных данных для корректного t-распределения"),
  answer("Использовать только функцию pt для вычисления теоретических вероятностей, не привлекая симуляции", 
         message = "Нет. pt дает теоретические значения. Симуляция визуализирует реальное распределение"),
  allow_retry = TRUE
)
```

### Вопрос 9

```{r q9, warning=FALSE}
question_radio("В R вы проводите множественные проверки гипотез (например, много t-тестов) и получаете множество p-значений. Почему интерпретация отдельных p-значений без поправки может быть проблематичной?",
  answer("Вероятность хотя бы одной ложной положительной ошибки возрастает с числом тестов", correct = TRUE, 
         message = "Правильно. При $\\alpha=0.05$ каждый тест дает 5% шанс ложного срабатывания. При 20 тестах P(хотя бы 1 ложная) = $1-(0.95)^20 = 64%!$

R-решение - поправка:
```r
pvals <- c(0.01, 0.04, 0.06, 0.001)  # 4 теста
p.adjust(pvals, method='BH')  # Benjamini-Hochberg
Без поправки: много ложных 'открытий'!"),
answer("t-тест автоматически корректирует p-значения при большом количестве тестов",
message = "Нет. t.test НЕ корректирует множественные сравнения автоматически"),
answer("p-значения становятся больше 1 при множественном тестировании",
message = "Нет. p всегда от 0 до 1. Поправка делает их БОЛЬШЕ (консервативнее)"),
answer("Центральная предельная теорема перестаёт действовать при множественных проверках",
message = "Нет. ЦПТ касается распределения средних, а не множественного тестирования"),
allow_retry = TRUE
)

```

### Вопрос 10

```{r q10}
question_radio("Вы проводите в R симуляционный эксперимент для оценки точности t-теста при ненормальном распределении данных. Какой вывод будет наиболее информативен для практики?",
  answer("t-тест может сохранять приемлемый уровень ошибки первого рода при умеренных отклонениях от нормальности, особенно при больших выборках", correct = TRUE, 
         message = "Правильно. t-тест робастен! При n>30 даже при скошенных данных уровень ошибки I рода остается близок к $\\alpha=0.05$.

R-симуляция показывает:
```r
set.seed(123); n<-30; sims<-1000; alpha<-0.05; rejects<-0
for(i in 1:sims){
  x <- rexp(n,1)  # Экспоненциальное (не нормальное) распределение!
  if(t.test(x)$p.value < alpha) rejects <- rejects + 1
}
rejects/sims  # ~0.05 при n=30!
```"),
  answer("При ненормальном распределении нельзя использовать никакие методы проверки гипотез", 
         message = "Нет. t-тест робастен к умеренным отклонениям, особенно при n>30"),
  answer("Симуляции в R не подходят для оценки свойств статистических тестов", 
         message = "Нет. Симуляции - золотой стандарт для проверки робастности!"),
  answer("t-тест всегда даёт абсолютно точные результаты вне зависимости от распределения данных", 
         message = "Нет. При сильных отклонениях лучше непараметрические wilcox.test()"),
  allow_retry = TRUE
)
```

### Вопрос 11

```{r q11, warning=FALSE}
question_radio("Какой вывод можно сделать, если в R однофакторный ANOVA показал значимое различие между группами, а последующий тест Тьюки (TukeyHSD) выявил, что различается только пара групп A и C?",
  answer("Средние групп A и C статистически значимо различаются, тогда как остальные пары не показали значимых различий", correct = TRUE, 
         message = "Правильно. ANOVA: 'Хотя бы одна пара групп различается'. TukeyHSD: 'Конкретно A vs C'.

R-пример:
```r
anova_model <- aov(y ~ group, data=df)
summary(anova_model)  # p < 0.05
TukeyHSD(anova_model) # Только A-C: p.adj < 0.05
ANOVA правильный - он чувствителен к любой паре различий!"),
answer("ANOVA был проведён неверно, раз не все пары групп оказались значимыми",
message = "Нет. H0 ANOVA: все mu равны. Один значимый Tukey достаточно!"),
answer("Все группы имеют одинаковые средние значения",
message = "Нет. ANOVA p < 0.05 уже отвергает равенство всех средних"),
answer("Только группа B отличается от остальных",
message = "Нет. TukeyHSD показал A-C, а не B vs остальные"),
allow_retry = TRUE
)
```

### Вопрос 12

```{r q12, warning=FALSE}
question_radio("Какой исследовательский вопрос лучше всего подходит для однофакторного ANOVA в R?",
  answer("Сравнить средние значения количественной переменной между тремя и более независимыми группами", correct = TRUE, 
         message = "Правильно. Однофакторный ANOVA: сравнение $\\mu_1 = \\mu_2 = \\mu_3...$ для трех и более групп.

R-синтаксис:
```r
aov(y ~ factor, data = df)  # y: количественная, factor: 3+ уровней
summary(aov_model)
Если p < 0.05 → TukeyHSD(aov_model) для парных сравнений"),
answer("Оценить, связаны ли две количественные переменные, например рост и вес",
message = "Нет. Для 2 количественных → корреляция cor.test() или регрессия lm()"),
answer("Оценить, связано ли распределение категориальной переменной с другой категориальной переменной",
message = "Нет. Для 2 категориальных → chisq.test() или fisher.test()"),
answer("Проверить, отличается ли медиана показателя между двумя зависимыми измерениями у тех же людей",
message = "Нет. Для парных данных → t.test(paired=TRUE) или wilcox.test(paired=TRUE)"),
allow_retry = TRUE
)
```
### Вопрос 13
```{r q13, warning=FALSE}
question_radio("Какое из следующих условий является ключевой предпосылкой для применения критерия Манна–Уитни в R?",
  answer("Распределение в выборках может быть ненормальным, но шкала должна быть не ниже порядковой", correct = TRUE, 
         message = "Правильно. Манна-Уитни - непараметрический тест для 2 независимых выборок порядковых/количественных данных.

R-синтаксис:
```r
wilcox.test(group1, group2)  # По умолчанию Манна-Уитни
# или явно:
wilcox.test(x ~ group, data = df)

2 независимые выборки

Порядковые или количественные данные

НЕ требует нормальности!"),
answer("Данные должны быть категориальными с небольшими ожидаемыми частотами",
message = "Нет. Это для хи-квадрат или fisher.test()"),
answer("Необходимо минимум три независимые группы",
message = "Нет. Для 3+ групп - Kruskal-Wallis (kruskal.test())"),
answer("Наблюдения в выборках обязательно должны быть парными (зависимыми)",
message = "Нет. Для парных - wilcox.test(paired=TRUE)"),
allow_retry = TRUE
)
```
### Вопрос 14
```{r q14, warning=FALSE}
question_radio("В каком случае критерий Краскела–Уоллиса будет более уместен, чем однофакторный ANOVA в R?",
  answer("Когда есть три и более независимых группы, распределения далеки от нормальных и шкала как минимум порядковая", correct = TRUE, 
         message = "Правильно. Kruskal-Wallis - непараметрический аналог ANOVA для 3+ групп без нормальности.

R-синтаксис:
```r
kruskal.test(y ~ group, data = df)  # 3+ группы
# или:
kruskal.test(list(groupA, groupB, groupC))
Если p < 0.05 → Dunn.test() или парные wilcox.test() для пост-хок анализа```"),
answer("Когда есть только одна выборка и нужно проверить её на соответствие нормальному распределению",
message = "Нет. Для одной выборки - shapiro.test()"),
answer("Когда нужно сравнить две зависимые выборки с нормальным распределением",
message = "Нет. Для парных нормальных - t.test(paired=TRUE)"),
answer("Когда сравниваются пропорции успешных исходов между несколькими категориальными группами",
message = "Нет. Для пропорций - prop.test() или chisq.test()"),
allow_retry = TRUE
)
```
### Вопрос 15
```{r, q15, warning=FALSE}
question_radio("Как верно интерпретировать значимое p-значение, полученное в результате применения критерия хи-квадрат в R?",
  answer("Между двумя категориальными переменными обнаружена статистически значимая связь", correct = TRUE, 
         message = "Правильно. chisq.test(): H0 - переменные независимы. p < 0.05 → отвергаем, есть связь.

R-пример:
```r
tab <- matrix(c(20,30,40,10), 2,2)
chisq.test(tab)  # p < 0.05 → зависимость!
Значимое p-value = статистически значимая связь между категориями!```"),
answer("Распределение одной количественной переменной является нормальным",
message = "Нет. Хи-квадрат для категориальных данных, не проверяет нормальность"),
answer("Две количественные переменные имеют линейную корреляцию",
message = "Нет. Для количественных - cor.test(). Хи-квадрат только для категориальных переменных"),
answer("Средние значения количественной переменной во всех группах одинаковы",
message = "Нет. Это ANOVA. Хи-квадрат работает с частотами, не со средними"),
allow_retry = TRUE
)
```
### Вопрос 16
```{r q16, warning=FALSE}
question_radio("Какой тип выборок предполагает использование парного (зависимого) t-теста или непараметрического аналога в R?",
  answer("Повторные измерения одной и той же группы участников во времени или при разных условиях", correct = TRUE, 
         message = "Правильно. Парный t-тест для связанных выборок: те же люди 'до/после' какого-то воздействия, например, программы профилактики

R-синтаксис:
```r
t.test(before, after, paired=TRUE)           # Векторы
t.test(value ~ time, data=df, paired=TRUE)   # Формула
wilcox.test(before, after, paired=TRUE)      # Непараметрический
Пример: вес пациентов до/после диеты```"),
answer("Сравнение частот категориальных ответов между разными группами",
message = "Нет. Это chisq.test() для таблиц сопряженности"),
answer("Две независимые группы, сформированные случайным образом",
message = "Нет. Для независимых - t.test() без paired=TRUE"),
answer("Несколько независимых групп (три и более)",
message = "Нет. Для 3+ групп - aov() или kruskal.test()"),
allow_retry = TRUE
)

```
### Вопрос 17
```{r q17, warning=FALSE}
question_radio("Какой код в R корректно задает однофакторный ANOVA для проверки влияния фактора group на переменную y?",
  answer("aov(y ~ group, data = df)", correct = TRUE, 
         message = "Правильно. Синтаксис ANOVA: зависимая ~ независимая

Полный код:
```r
model <- aov(y ~ group, data = df)
summary(model)  # Таблица ANOVA
y - количественная (зависимая), group - фактор с 3+ уровнями```"),
answer("aov(group ~ y, data = df)",
message = "Неправильно. Формула перевёрнута: зависимая должна быть слева"),
answer("aov(y + group, data = df)",
message = "Неправильно. + используется для нескольких зависимых переменных"),
answer("anova(y ~ group, data = df)",
message = "Близко, но anova() для lm/aov моделей. Сначала нужен aov()"),
allow_retry = TRUE
)
```
### Вопрос 18
```{r q18, warning=FALSE}
question_radio("Какой аргумент функции wilcox.test() в R нужно установить, чтобы сравнивать зависимые (парные) выборки, например до и после вмешательства?",
  answer("paired = TRUE", correct = TRUE, 
         message = "Правильно. wilcox.test(paired=TRUE) - парный непараметрический тест (аналог парного t-теста).

Пример:
```r
before <- c(25, 28, 30, 27, 29)
after <- c(23, 26, 28, 25, 27)
wilcox.test(before, after, paired = TRUE)

# Или формулой:
wilcox.test(value ~ time, paired = TRUE, data = df)
```"),
  answer("alternative = 'paired'", 
         message = "Нет. alternative = 'two.sided', 'greater', 'less'"),
  answer("var.equal = TRUE", 
         message = "Нет. var.equal есть только в t.test()"),
  answer("paired = FALSE", 
         message = "Нет. paired = FALSE (по умолчанию) - для независимых выборок"),
  allow_retry = TRUE
)
```
### Вопрос 19
```{r q19, warning=FALSE}
question_radio("Какая формулировка нулевой гипотезы подходит для критерия Краскела–Уоллиса при сравнении трёх независимых групп?",
  answer("Распределения (включая медианы) во всех группах одинаковы", correct = TRUE, 
         message = "Правильно. Kruskal-Wallis: H0 - все группы имеют одинаковые распределения (медианы).

R-синтаксис:
```r
kruskal.test(y ~ group, data = df)  # 3+ группы
# H0: f1(x) = f2(x) = f3(x)
p < 0.05 → хотя бы одно распределение отличается```"),
answer("Медианы по крайней мере двух групп различаются",
message = "Нет. Это альтернативная гипотеза H1"),
answer("Средние всех групп равны нулю",
message = "Нет. Это для ANOVA или t.test(mu=0). Kruskal работает с медианами"),
answer("Две выборки являются зависимыми",
message = "Нет. Kruskal-Wallis только для независимых групп"),
allow_retry = TRUE
)
```
### Вопрос 20
```{r q20, warning=FALSE}
question_radio("При применении критерия хи-квадрат в R к таблице сопряжённости $2\\times3$ какое требование к ожидаемым частотам считается важным?",
  answer("Ожидаемые частоты в каждой ячейке должны быть не меньше 5", correct = TRUE, 
         message = "Правильно. Правило 5: ожидаемые частоты >=5 во всех ячейках для корректного хи-квадрат.

R-проверка:
```r
chisq_result <- chisq.test(tab_2x3)
print(chisq_result$expected)  # Ожидаемые частоты
# Если <5 → fisher.test(tab) или объединить категории
При нарушении: точный тест Фишера или объединение ячеек```"),
answer("Медианы по строкам и столбцам должны совпадать",
message = "Нет. Хи-квадрат работает с частотами, не с медианами"),
answer("Все наблюдаемые частоты должны быть равны",
message = "Нет. Ожидаемые вычисляются автоматически из marginal totals"),
answer("Сумма всех частот в таблице должна быть равна нулю",
message = "Нет. Сумма >0, обычно n_total наблюдений"),
allow_retry = TRUE
)
```

## Практические задания


### Упражнение 1

Согласно данным комплексного наблюдения условий жизни россиян, проведенного Росстатом, время ожидания общественного транспорта достигло рекорда за девять лет и составило в 2024 году 15 минут и 48 секунд в обе стороны (округлим до 16 минут). 

Вы провели аналогичное социологическое исследование в городе Барнауле, результаты которого содержатся в `data`.

С помощью одновыборочного t-критерия Стьюдента выясните, отличаются ли Ваши данные от общероссийских.

![Виталий Тимкив. Трамвайный вагон в Краснодаре](https://cdnn21.img.ria.ru/images/07e8/03/1e/1936839939_0:323:3065:2047_1280x0_80_0_0_cb260f324c387369eb4053bd7c46c2ba.jpg.webp){width=300px}  ![Барнаульский трамвай](https://www.ap22.ru/netcat_files/multifile/2546/189624/G0G6pBeW5p4_1_.jpg){width=300px}

```{r t-one-way-1-setup, include=FALSE}
set.seed(1234)
data<-rnorm(n=1000, mean = 16.4, sd=7.1)
```

```{r t-one-way-1, exercise = TRUE}

```

```{r t-one-way-1-hint-1, eval = FALSE}
t.test(data, mu = ...)
```


```{r t-one-way-1-solution, include = FALSE}
t.test(data, mu = 16)
```

```{r t-one-way-1-code-check}
grade_code()
```

```{r t-one-way-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 2

В наборе данных `js` содержатся результаты оценки удовлетворенности работой (переменная `score`). Проведите анализ средних значений удовлетворенности в зависимости от пола (переменная `gender`).

![](images\js.jpg)

```{r js-setup, include=FALSE}
library(datarium)
data("jobsatisfaction")
js<-jobsatisfaction
```

```{r js, exercise = TRUE, eval=TRUE}

```

```{r js-test, include = FALSE}
t.test(score~gender, data = js)
```

```{r js-hint-1, eval = FALSE}
t.test(score~..., data = ...)
```


```{r js-solution, include = FALSE}
t.test(score~gender, data = js)
```

```{r js-code-check}
grade_code()
```

```{r js-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 3

Используя возможности библиотеки `ggstatsplot` и функции `ggbetweenstats` создайте график, показывающий различия в средних значениях удовлетворенности работой у мужчин и женщин. 

У Вас должно получиться следующее:

```{r echo=FALSE, eval=TRUE}
ggbetweenstats(
  js,
  x    = gender,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах мужчин и женщин"
)
```

```{r js-2-setup, include=FALSE}
library(datarium)
data("jobsatisfaction")
js<-jobsatisfaction
```

```{r js-2, exercise = TRUE, eval=TRUE}

```

```{r js-2-test, include = FALSE}
ggbetweenstats(
  js,
  x    = gender,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах мужчин и женщин"
)
```

```{r js-2-hint-1, eval = FALSE}
ggbetweenstats(
  js,
  x    = ...,
  y    = ...,
  type = "parametric",
  xlab = "..."
)
```


```{r js-2-solution, include = FALSE}
ggbetweenstats(
  js,
  x    = gender,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах мужчин и женщин"
)
```

```{r js-2-code-check}
grade_code()
```


### Упражнение 4

В наборе данных `civil_society` содержатся результаты исследования о развитии гражданского общества в Алтайском крае, проведенного в 2025-2026 году, в том числе 10-балльные шкалы, оценивающие институциональные условия, необходимые для развития гражданского общества.

Проведите сравнительный анализ средних оценок показателя «НКО оказывают качественные услуги населению» (`NPO`) у жителей, имеющих и не имеющих опыта взаимодействия с НКО  (`NKO_link`).


```{r data, echo=FALSE}
civil_society<-haven::read_sav("data/civil_society.sav")
civil_society<-haven::as_factor(civil_society)
```

```{r mu-1, exercise = TRUE, eval=TRUE}

```

```{r mu-1-test, include = FALSE}
wilcox.test(NPO~NKO_link, data = civil_society)
```

```{r mu-1-hint-1, eval = FALSE}
wilcox.test(NPO~..., data = ...)
```


```{r mu-1-solution, include = FALSE}
wilcox.test(NPO~NKO_link, data = civil_society)
```

```{r mu-1-code-check}
grade_code()
```

```{r mu-1-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 5

Сделайте визуализацию различий с помощью функции `ggbetweenstats`. Чтобы вывести результаты критерия Манна-Уитни, установите `type="nonparametric"`.

У Вас должно получиться следующее:

```{r echo=FALSE}
ggbetweenstats(civil_society, NKO_link, NPO, xlab = "Оценка НКО", type="nonparametric")
```


```{r mu-vis, exercise = TRUE, eval=TRUE}

```

```{r mu-vis-test, include = FALSE}
ggbetweenstats(civil_society, NKO_link, NPO, xlab = "Оценка качества услуг НКО", type="nonparametric")
```

```{r mu-vis-hint-1, eval = FALSE}
ggbetweenstats(civil_society, NKO_link, NPO, xlab = "Оценка качества услуг НКО", type="nonparametric")
```


```{r mu-vis-solution, include = FALSE}
ggbetweenstats(civil_society, NKO_link, NPO, xlab = "Оценка качества услуг НКО", type="nonparametric")
```

```{r mu-vis-code-check}
grade_code()
```

### Упражнение 6

Продолжаем работать с данными об удовлетворенности работой `js`. В переменной `education_level` содержится информация об образовании. В ней три уровня:

- school (средняя школа)
- colledge (колледж)
- university (университет)

С помощью однофакторного дисперсионного анализа проведите анализ различий в среднем уровне удовлетворенности в зависимости от образования участника исследования. Сохраните результаты в `res` и выведите их с помощью `summary()`.

```{r js-aov-setup, include=FALSE}
library(datarium)
data("jobsatisfaction")
js<-jobsatisfaction
```

```{r js-aov, exercise = TRUE, eval=TRUE}

```

```{r js-aov-test, include = FALSE}
res<-aov(score~education_level, data = js)
summary(res)
```

```{r js-aov-hint-1, eval = FALSE}
res<-aov(score~..., data = ...)
summary(...)
```


```{r js-aov-solution, include = FALSE}
res<-aov(score~education_level, data = js)
summary(res)
```

```{r js-aov-code-check}
grade_code()
```

```{r js-aov-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 7

Между какими именно группами сотрудников по образованию различия значимы?. Проведите множественные сравнения методом `tukey_hsd`. 

[Узнать больше](https://rpkgs.datanovia.com/rstatix/reference/tukey_hsd.html) о функции `tukey_hsd`.

```{r js-pwc, exercise = TRUE, eval=TRUE}

```

```{r js-pwc-test, include = FALSE}
tukey_hsd(js, score ~ education_level)
```

```{r js-pwc-hint-1, eval = FALSE}
tukey_hsd(js, score ~ education_level
```

```{r js-pwc-solution, include = FALSE}
tukey_hsd(js, score ~ education_level)
```

```{r js-pwc-code-check}
grade_code()
```

```{r js-pwc-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 8

Используя возможности библиотеки `ggstatsplot` и функции `ggbetweenstats` создайте график, показывающий различия в средних значениях удовлетворенности работой у сотрудников с разным образовательным уровнем. 

У Вас должно получиться следующее:

```{r echo=FALSE, eval=TRUE}
ggbetweenstats(
  js,
  x    = education_level,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах с разным образованием"
)
```

```{r aov-viz-setup, include=FALSE}
library(datarium)
data("jobsatisfaction")
js<-jobsatisfaction
```

```{r aov-viz, exercise = TRUE, eval=TRUE}

```

```{r aov-viz-test, include = FALSE}
ggbetweenstats(
  js,
  x    = education_level,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах с разным образованием"
)
```

```{r aov-viz-hint-1, eval = FALSE}
ggbetweenstats(
  js,
  x    = ...,
  y    = ...,
  type = "parametric",
  xlab = "..."
)
```


```{r aov-viz-solution, include = FALSE}
ggbetweenstats(
  js,
  x    = education_level,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах с разным образованием"
)
```

```{r aov-viz-code-check}
grade_code()
```


### Упражнение 9

На основе данных `civil_society` проведите анализ средних оценок социальной ответственности бизнеса (переменная `business`) в группах жителей, отличающихся местом проживания (переменная `place`). Нужно провести анализ по трем группам:

- Барнаул
- Другие города (малые и средние)
- Сельские районы

Для общей оценки сделайте тест Краскела-Уоллиса `kruskal.test()`, для парных сравнений - тест `pairwise.wilcox.test()` с поправкой Бенджамини-Хохберга

`p.adjust(method = "BH")` — процедура Бенджамини-Хохберга для контроля FDR (False Discovery Rate) - ложноположительных результатов  при множественных сравнениях.

BH менее строгая альтернатива Бонферрони, сохраняет статистическую мощность.

Алгоритм работы

- Сортировка p-значений: $p_{(1)} ≤ p_{(2)} ≤ ... ≤ p_{(m)}$
- Коррекция: $\hat{p}i = \min\left(1, \frac{m \cdot p{(i)}}{i}\right)$
- Кумулятивный минимум для монотонности
[Большая обзорная статья про разные виды поправок](https://pmc.ncbi.nlm.nih.gov/articles/PMC5506159/)

```{r ks, exercise = TRUE, eval=TRUE}

```

```{r ks-test, include = FALSE}
kruskal.test(business~place, data = civil_society)
pairwise.wilcox.test(civil_society$business,civil_society$place,p.adjust.method = "BH")

```

```{r ks-hint-1, eval = FALSE}
kruskal.test(business~..., data = ...)
pairwise.wilcox.test(...$business, ...$place, p.adjust.method = "BH")
```


```{r ks-solution, include = FALSE}
kruskal.test(business~place, data = civil_society)
pairwise.wilcox.test(civil_society$business,civil_society$place,p.adjust.method = "BH")
```

```{r ks-code-check}
grade_code()
```

```{r ks-essay}
question_text(
    "Что Вы можете сказать по поводу результатов - дайте небольшой комментарий:",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```


### Упражнение 10

С помощью функции `ggbetweenstats` создайте график, визуализирующие результаты анализа из упражнения 9.  
У Вас должно получиться следующее:

```{r echo=FALSE, eval=TRUE}
ggbetweenstats(
  civil_society,
  x    = place,
  y    = business,
  type = "non-parametric",
  xlab = "Оценка социальной ответственности бизнеса в поселениях разного типа"
)
```


```{r ks-viz, exercise = TRUE, eval=TRUE}

```

```{r ks-viz-test, include = FALSE}
ggbetweenstats(
  js,
  x    = education_level,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах с разным образованием"
)
```

```{r ks-viz-hint-1, eval = FALSE}
ggbetweenstats(
  js,
  x    = ...,
  y    = ...,
  type = "parametric",
  xlab = "..."
)
```


```{r ks-viz-solution, include = FALSE}
ggbetweenstats(
  js,
  x    = education_level,
  y    = score,
  type = "parametric",
  xlab = "Сравнение уровня удовлетворенности работой в группах с разным образованием"
)
```

```{r ks-viz-code-check}
grade_code()
```

```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
