---
title: "Моделирование структурными уравнениями (SEM)"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
bibliography: references.bib
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(lavaan)
library(lavaanPlot)

knitr::opts_chunk$set(echo = FALSE)

tutorial_options(exercise.reveal_solution = FALSE)

set.seed(2024)

pop_model <- '
  trust     =~ 0.75*trust1 + 0.70*trust2 + 0.65*trust3
  wellbeing =~ 0.80*wb1    + 0.72*wb2    + 0.68*wb3
  civic     =~ 0.70*civic1 + 0.66*civic2 + 0.73*civic3
  wellbeing ~  0.40*trust  + 0.30*civic
  civic     ~  0.35*trust
  trust     ~~ 1*trust
  wellbeing ~~ 0.5*wellbeing
  civic     ~~ 0.5*civic
'

social_data <- lavaan::simulateData(
  pop_model,
  sample.nobs = 500,
  seed        = 2024
)

social_data$sex <- factor(
  sample(c("мужчины", "женщины"), 500, replace = TRUE, prob = c(0.48, 0.52)),
  levels = c("мужчины", "женщины")
)

# --- CFA ---
cfa_model <- '
  trust     =~ trust1 + trust2 + trust3
  wellbeing =~ wb1    + wb2    + wb3
  civic     =~ civic1 + civic2 + civic3
'
fit_cfa <- cfa(cfa_model, data = social_data)

fit_indices <- fitmeasures(fit_cfa,
  c("chisq", "df", "pvalue", "cfi", "rmsea", "srmr"))

# --- Модели инвариантности ---
fit_config <- cfa(cfa_model, data = social_data, group = "sex")
fit_metric <- cfa(cfa_model, data = social_data, group = "sex",
                  group.equal = "loadings")
fit_scalar <- cfa(cfa_model, data = social_data, group = "sex",
                  group.equal = c("loadings", "intercepts"))

lrt_result <- lavTestLRT(fit_config, fit_metric, fit_scalar)

# --- Структурная модель ---
sem_model <- '
  trust     =~ trust1 + trust2 + trust3
  wellbeing =~ wb1    + wb2    + wb3
  civic     =~ civic1 + civic2 + civic3
  wellbeing ~  trust + civic
  civic     ~  trust
'
fit_sem <- sem(sem_model, data = social_data)

# параметры структурных путей
struct_params <- parameterEstimates(fit_sem, standardized = TRUE) |>
  dplyr::filter(op == "~")
```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Предварительные замечания

Моделирование структурными уравнениями (**SEM**, от англ. *Structural Equation Modeling*) — семейство статистических методов, которое объединяет **анализ путей (path analysis)** и **анализ латентных переменных** в единой модели, что дает возможность протестировать множество исследовательских гипотез одновременно.

В отличие от регрессии, SEM позволяет:

- работать с **латентными конструктами** (переменными, которые нельзя измерить напрямую: доверие, благополучие, социальный капитал и др.);
- оценивать **ошибку измерения** явным образом;
- проверять как **измерительную**, так и **структурную** части модели одновременно;
- тестировать **сложные причинно-следственные цепочки**, включая медиацию и модерацию.

SEM широко применяется в:

- социологии (изучение неравенства, мобильности, межнациональных отношений, ценностных ориентаций, социальной напряженности и конфликтов, маркетинговые исследования);
- психологии (например, при изучении взаимовлияния между личностными свойствами, психического благополучия, стресса, личностной зрелости и других интересных тем);
- политической науке (гражданское участие, институциональное доверие);
- экономике (измерение качества жизни, человеческого капитала, социально-экономического развития регионов и др.).

### Компоненты модели

Модель SEM состоит из двух частей.

**Модель измерения (measurement model)** -- задает, как наблюдаемые переменные (\(y_i\)) соотносятся с латентными факторами (\(\eta\)):

\[
y_i = \lambda_i \eta + \varepsilon_i
\]

где \(\lambda_i\) —- факторная нагрузка, \(\varepsilon_i\) — ошибка измерения. Эта часть модели тестируется с помощью методов конфирматорного факторного анализа. Например, в нашем исследовании есть задачи по оценке участия граждан в реализации инициатив, направленных на городское развитие. Мы провели операционализацию и определили индикаторы такого участия, измерили их, проведя социологическое исследование. Задача данного этапа будет заключаться в обосновании, действительно ли индикаторы измеряют нечто общее, то, что можно назвать гражданским участием.


**Структурная модель (structural model)** описывает связи между латентными конструктами и наблюдаемыми переменными (которые не нужно измерять дополнительно):

\[
\eta = B\eta + \Gamma\xi + \zeta
\]

где \(B\) — матрица путей между эндогенными конструктами, \(\Gamma\) — матрица путей от экзогенных конструктов, \(\zeta\) — структурные остатки. Например, мы можем посмотреть, как на гражданское участие влияют разные факторы, такие как доход или образование, а также ввести в модель еще один латентный фактор (допустим, институциональное доверие, которое мы оценим на основе батареи шкал доверия к разным социальным институтам), и выявить, как влияет доверие на уровень участия [@kline2016].

::: {#fig-sem}

![Основные компоненты модели SEM](images\sem.png){width=500px}
[@ucla-rsem]
:::


### Оценка модели и индексы подгонки

В SEM важно различать два уровня: **как** оцениваются параметры (оцениватель, тип матрицы) и **как** оценивается качество модели в целом (индексы подгонки).


- Для приблизительно **непрерывных, нормально распределенных** показателей (классический случай) по умолчанию используется метод **ML** (maximum likelihood), который минимизирует расхождение между модельной и эмпирической матрицами ковариаций.  
- При **ненормальности** часто применяют робастные варианты ML:  
 При обсуждении выбора оценивателя в SEM обычно отдельно рассматривают случай ненормальных данных и анализ переменных порядкового типа (например, шкалы Лайкерта). Ниже приведены краткие характеристики основных методов в форме, подходящей для учебного пособия.

### Оцениватели при отсутствии нормального распределения

- **MLR (Maximum Likelihood Robust)**  -- представляет собой модификацию максимального правдоподобия, в которой стандартные ошибки и статистика \( \chi^2 \) корректируются на отклонения распределения от нормального. Такой подход позволяет получать более надежные выводы при умеренной ненормальности, сохраняя при этом интерпретацию параметров, аналогичную классическому ML.  

- **MLM / MLMV (Satorra–Bentler и родственные варианты)** -- используют корректировку статистики согласия \( \chi^2 \) с учетом избыточной или недостаточной дисперсии при ненормальности. В варианте MLMV дополнительный акцент делается на корректировку ковариационной матрицы, что улучшает свойства тестов согласия и доверительных интервалов в условиях сильной ненормальности. Эти методы часто рекомендуются, когда данные заметно отклоняются от многомерной нормальности, но сохраняется количественный характер переменных.

- **WLSMV (Weighted Least Squares Mean and Variance adjusted)**  -- применяется при моделировании порядковых, бинарных и смешанных переменных, объявленных как `ordered`. Он основан на методе взвешенных наименьших квадратов, но включает робастные стандартные ошибки и модифицированную статистику \( \chi^2 \), скорректированную по среднему и дисперсии. На практике WLSMV рассматривается как один из наиболее устойчивых и рекомендуемых вариантов для рейтинговых шкал (например, Лайкерта), особенно при умеренных и больших объемах выборки.  

- **DWLS (Diagonal Weighted Least Squares)**  -- является упрощенной версией WLS, в которой используется только диагональ весовой матрицы, то есть отдельные веса для каждой переменной без учета всех парных ковариаций в весах. Это снижает вычислительную нагрузку и делает метод применимым для моделей с большим числом индикаторов, однако часть информации о структуре ошибок в весах игнорируется. DWLS часто используется как практический компромисс между точностью и вычислительной эффективностью при анализе моделей с большим количеством порядковых показателей.

#### Индексы подгонки

После оценки модели исследователь проверяет, насколько хорошо она отражает эмпирические данные — для этого используются различные индексы подгонки:

- **Абсолютные индексы**:  
  - \(\chi^2\) — тестирует точное совпадение модели с данными; при больших выборках почти всегда значим, поэтому его используют скорее как ориентир.  
  - **SRMR** (Standardized Root Mean Square Residual) показывает среднюю величину несоответствия между корреляциями, полученными из данных, и корреляциями, предсказанными моделью.

- **Индексы относительной подгонки (инкрементные)**:  
  - **CFI, TLI** — сравнивают модель с «нулевой» (где переменные независимы).  
- **Индексы приближенной подгонки**:  
  - **RMSEA** — оценивает степень несовпадения модели с данными в пересчете на единицу степени свободы; для него часто приводят доверительный интервал.  

Типичные ориентиры для интерпретации коэффициентов:

| Индекс      | Приемлемая подгонка           | Хорошая подгонка              |
|-------------|-------------------------------|-------------------------------|
| \(\chi^2/df\) | < 5.0                        | < 2.0                         |
| CFI, TLI    | > 0.90                        | > 0.95                        |
| RMSEA       | < 0.08                        | < 0.06                        |
| SRMR        | < 0.08                        | < 0.05                        |

Важно отметить, что при использовании **WLSMV/DWLS** интерпретация этих индексов в целом сохраняется, но **χ² и RMSEA основываются на приближенных формулах**, а значения могут систематически отличаться от ML‑оцененных моделей; сравнивать модели с разными типами оценивателей по $\chi^2$ и RMSEA некорректно.

> На практике часто смотрят не только на отдельные значения CFI, RMSEA и SRMR, но и на **изменение** этих индексов при усложнении/ограничении модели (например, в тестах инвариантности: \(\Delta CFI\), \(\Delta RMSEA\)), а также комбинируют оценку подгонки с содержательной интерпретируемостью параметров и простотой модели.

### SEM в R
В литературе по структурным уравнениям в R обычно выделяют две группы инструментов: основные пакеты для спецификации и оценки SEM‑моделей и дополнительные пакеты, ориентированные на визуализацию, отчетность и вспомогательные процедуры.

**Основные пакеты**:

- **lavaan** — пакет для создания моделей структурных уравнений на основе синтаксиса, создающего формулу модели. Ориентирован на классический SEM в прикладных исследованиях.  
- **sem** — реализация SEM в матричной форме, отражающая традиционный ковариационно‑структурный подход к заданию моделей.  
- **OpenMx** — гибкий фреймворк для сложных и нестандартных структурных моделей с матричным описанием и высоким уровнем контроля.  
- **semPlot** — инструмент для графической визуализации путевых диаграмм и структуры SEM‑моделей.  
- **semTools** — набор вспомогательных процедур для оценки надежности, инвариантности и дополнительной диагностики моделей в SEM.  
- **sjPlot / report** — пакеты для автоматизированного табличного и текстового представления результатов моделирования, пригодные для оформления SEM‑анализов.  
- **lavaanPlot** — специализированный инструмент для построения наглядных диаграмм моделей lavaan с возможностью настройки подписей и оформления.  
- **lavaanExtra** — набор дополнительных функций, упрощающих выполнение типичных операций с моделями lavaan и подготовку результатов к представлению.

### Библиотека lavaan

Пакет lavaan (Latent Variable Analysis) — это один из ключевых инструментов для выполнения моделей структурных уравнений (SEM) в R. Он позволяет исследователю задавать и оценивать модели с латентными переменными — такими как факторные модели, модели путей, а также их комбинации. В основе lavaan лежит синтаксис, похожий на формулировку уравнений, что делает его удобным для описания как простых CFA (Confirmatory Factor Analysis), так и сложных структур с несколькими уровнями и связями.

Пакет lavaan был разработан бельгийским статистиком и психометриком Ивом Росселем (Yves Rosseel) как современная, открытая и воспроизводимая реализация моделей структурных уравнений. Его развитие началось в конце 2000‑х годов, а первые стабильные версии стали широко использоваться в исследовательской среде примерно с начала 2010‑х [@rosseel]. О популярности данной библиотеки можно судить по количеству цитирований: на статью «[lavaan: An R package for structural equation modeling](https://scholar.google.ru/citations?view_op=view_citation&hl=ru&user=0R_YqcMAAAAJ&citation_for_view=0R_YqcMAAAAJ:2osOgNQ5qMEC)» сослались более 34 тыс. ученых со всего мира.


Lavaan имеет удобные функции для проверки подгонки модели, он автоматически вычисляет распространенные индексы, такие как CFI, TLI, RMSEA и SRMR, позволяет строить доверительные интервалы, тестировать ограничения параметров и сравнивать модели. Пакет интегрируется с инструментами визуализации и отчетности (например, semPlot, lavaanPlot, semTools, lavaanExtra), что делает его основой для обучения и прикладного анализа в социальных, психологических и образовательных исследованиях.

### Синтаксис lavaan

В пакете `lavaan` модели задаются в виде строки с операторами:

- `=~` — определение латентной переменной («измеряется через»);
- `~` — регрессионный путь («зависит от»);
- `~~` — ковариация или дисперсия;
- `~1` — интерцепт переменной.

Типичный пример кода:

```r
# модель
cfa_model <- '
  trust =~ trust1 + trust2 + trust3    # trust измеряется trust1, trust2, trust3
  wellbeing =~ wb1 + wb2 + wb3
'
fit <- cfa(cfa_model, data = social_data)
summary(fit, fit.measures = TRUE, standardized = TRUE)
```

## Данные: социальное благополучие молодежи

В этом тьюториале используются **смоделированные данные** по образцу реальных опросов о социальном благополучии молодежи (N = 500 респондентов).

Данные содержат **9 индикаторов** трех латентных конструктов и переменную пола:

| Переменная | Конструкт | Содержание                                    |
|------------|-----------|-----------------------------------------------|
| `trust1`   | Доверие   | Доверие правительству                        |
| `trust2`   | Доверие   | Доверие судебной системе                     |
| `trust3`   | Доверие   | Доверие полиции                              |
| `wb1`      | Благополучие | Удовлетворенность жизнью                 |
| `wb2`      | Благополучие | Ощущение счастья                         |
| `wb3`      | Благополучие | Позитивный аффект                        |
| `civic1`   | Гражд. активность | Участие в голосовании              |
| `civic2`   | Гражд. активность | Участие в общественных организациях|
| `civic3`   | Гражд. активность | Волонтерство                       |
| `sex`      | —         | Пол (мужчины / женщины)                      |

Теоретические гипотезы:

1. Каждый конструкт хорошо измеряется своими индикаторами (CFA).
2. Структура измерения одинакова для мужчин и женщин (инвариантность).
3. Доверие институтам положительно предсказывает гражданскую активность и субъективное благополучие (структурные пути).


## Упражнения: Модель измерения (CFA)

### Упражнение 1 — Знакомство с данными

Изучите структуру набора данных `social_data`.

- Выведите первые 6 строк с помощью `head()`.
- Проверьте структуру данных с помощью `str()`.

```{r ex1, exercise = TRUE}
# ваш код
```

```{r ex1-hint}
# head(social_data)
# str(social_data)
```

```{r ex1-solution, include=FALSE}
head(social_data)
str(social_data)
```

```{r ex1-check}
grade_code()
```

```{r ex1-q}
question_radio(
  "Сколько наблюдаемых переменных (индикаторов) содержат данные (не считая переменную sex)?",
  answer("6"),
  answer("7"),
  answer("9", correct = TRUE),
  answer("10"),
  allow_retry = TRUE,
  incorrect = "Посчитайте столбцы trust1–trust3, wb1–wb3, civic1–civic3."
)
```

### Упражнение 2 — Создание модели CFA

Задайте модель конфирматорного факторного анализа.

- Создайте строку `cfa_model`, которая определяет три фактора:
  - `trust` — измеряется через `trust1`, `trust2`, `trust3`
  - `wellbeing` — измеряется через `wb1`, `wb2`, `wb3`
  - `civic` — измеряется через`civic1`, `civic2`, `civic3`
- Оцените модель функцией `cfa()` на данных `social_data`; сохраните в `fit_cfa`.

```{r ex2, exercise = TRUE}
cfa_model <- '
  # ваш синтаксис
'
fit_cfa <- ___(___, data = ___)
```

```{r ex2-hint, eval=FALSE}
# cfa_model <- '
#   trust     =~ trust1 + trust2 + trust3
#   wellbeing =~ wb1 + wb2 + wb3
#   civic     =~ civic1 + civic2 + civic3
# '
# fit_cfa <- cfa(cfa_model, data = social_data)
```

```{r ex2-solution, include=FALSE}
cfa_model <- '
  trust     =~ trust1 + trust2 + trust3
  wellbeing =~ wb1    + wb2    + wb3
  civic     =~ civic1 + civic2 + civic3
'
fit_cfa <- cfa(cfa_model, data = social_data)
```

```{r ex2-check}
grade_this({
  if (!exists("fit_cfa")) fail("Создайте объект fit_cfa.")
  if (!inherits(fit_cfa, "lavaan")) fail("fit_cfa должен быть объектом lavaan.")
  pass("CFA модель успешно задана и оценена.")
})
```

### Упражнение 3 — Итоги CFA

Выведите итоговый отчет по модели с помощью `summary()`.

- Используйте аргументы `fit.measures = TRUE` и `standardized = TRUE`.

```{r ex3, exercise = TRUE}
# ваш код
```

```{r ex3-hint, eval=FALSE}
# summary(fit_cfa, fit.measures = TRUE, standardized = TRUE)
```

```{r ex3-solution, include=FALSE}
summary(fit_cfa, fit.measures = TRUE, standardized = TRUE)
```

```{r ex3-check}
grade_code()
```

```{r ex3-q1}
question(
  "Что означает столбец `Std.all` в выводе `summary()` при `standardized = TRUE`?",
  answer("Стандартизованная факторная нагрузка (все переменные приведены к единичной дисперсии).", correct = TRUE),
  answer("Стандартная ошибка оценки."),
  answer("Статистика критерия хи-квадрат."),
  answer("Значение p для каждого параметра."),
  allow_retry = TRUE,
  incorrect = "Столбец Std.all содержит полностью стандартизованные нагрузки, сопоставимые между индикаторами."
)
```

```{r ex3-q2}
question(
  "Какова доля остаточной (необъясненной) дисперсии индикатора civic3 в полностью стандартизованной модели (см. строку `.civic3` в колонке `Std.all`)?",
  answer("Около 0.30."),
  answer("Около 0.55."),
  answer("Около 0.70.", correct = TRUE),
  answer("Около 0.90."),
  allow_retry = TRUE,
  incorrect = "В строке `.civic3` раздела Variances значение `Std.all = 0.702`, что соответствует приблизительно 70% остаточной дисперсии индикатора."
)
```
### Упражнение 4 — Индексы подгонки

Извлеките ключевые индексы подгонки модели с помощью `fitmeasures()`.

- Запросите: `"chisq"`, `"df"`, `"pvalue"`, `"cfi"`, `"rmsea"`, `"srmr"`.
- Сохраните результат в `fit_indices`.

```{r ex4, exercise = TRUE}
fit_indices <- fitmeasures(___, c("chisq", "df", "pvalue", "cfi", "rmsea", "srmr"))
fit_indices
```

```{r ex4-hint, eval=FALSE}
# fit_indices <- fitmeasures(fit_cfa, c("chisq", "df", "pvalue", "cfi", "rmsea", "srmr"))
# fit_indices
```

```{r ex4-solution, include=FALSE}
fit_indices <- fitmeasures(fit_cfa, c("chisq", "df", "pvalue", "cfi", "rmsea", "srmr"))
fit_indices
```

```{r ex4-check}
grade_this({
  if (!exists("fit_indices")) fail("Создайте объект fit_indices.")
  if (!all(c("cfi", "rmsea", "srmr") %in% names(fit_indices)))
    fail("Убедитесь, что вы запросили cfi, rmsea и srmr.")
  pass("Индексы подгонки извлечены.")
})
```

```{r ex4-q}
question_radio(
  "Как можно интерпретировать приведенные значения индексов с точки зрения качества модели?",
  answer("Модель имеет плохую подгонку к данным."),
  answer("Модель имеет удовлетворительную подгонку, но требует доработки."),
  answer("Модель демонстрирует отличную подгонку к данным.", correct = TRUE),
  answer("По этим индексам нельзя судить о подгонке модели."),
  allow_retry = TRUE,
  incorrect = "Высокий CFI  RMSEA близкий к 0 и SRMR < 0.05 обычно свидетельствуют об очень хорошей подгонке модели."
)
```

### Упражнение 5 — Визуализация модели CFA

Постройте диаграмму пути для модели CFA с помощью `semPaths()`.

- Используйте `model = fit_cfa` для отображения результатов конфирматорного анализа.
- Добавьте `stand = TRUE`, `stars = "latent"`.

```{r ex5, exercise = TRUE, }
lavaanPlot(
  model = ___,
  stand = ___,
  coeffs= ___,
  stars = ___
)
```

```{r ex5-hint, eval=FALSE}
# lavaanPlot(model = fit_cfa, stand = TRUE, coefs = TRUE, stars = "latent")
```

```{r ex5-solution, include=FALSE}
lavaanPlot(model = fit_cfa, stand = TRUE, coefs = TRUE, stars = "latent")
```

```{r ex5-check}
grade_code()
```

```{r ex5-q}
question_radio(
  "Что отображают числа на стрелках в диаграмме пути при `stand = TRUE`?",
  answer("Стандартизованные факторные нагрузки (корреляции между фактором и индикатором).", correct = TRUE),
  answer("p-значения для каждого пути."),
  answer("Коэффициенты детерминации $R^2$."),
  answer("Дисперсии ошибок измерения."),
  allow_retry = TRUE
)
```



## Упражнения: Проверка инвариантности измерения

Проверка инвариантности измерения позволяет убедиться, что один и тот же конструкт одинаково понимается и измеряется в разных группах (например, у мужчин и женщин).

Процедура включает последовательное тестирование нескольких моделей:

1. **Конфигуральная** — одинаковая структура (те же факторы, те же индикаторы), но все параметры свободны в каждой группе.
2. **Метрическая** — дополнительно ограничиваются **факторные нагрузки** (проверяется гипотеза об их равенстве в группах).
3. **Скалярная** — дополнительно ограничиваются **интерцепты** индикаторов.

Если скалярная модель подходит не хуже метрической (\(\Delta CFI < 0.010\), \(\Delta RMSEA < 0.015\)), то сравнение средних значений латентных конструктов корректно.

### Упражнение 6 — Конфигуральная модель

Оцените конфигуральную модель инвариантности: аргумент `group = "sex"` без дополнительных ограничений.

```{r ex6, exercise = TRUE}
fit_config <- cfa(
  ___,
  data  = ___,
  group = ___
)
summary(fit_config, fit.measures = TRUE)
```

```{r ex6-hint, eval=FALSE}
# fit_config <- cfa(cfa_model, data = social_data, group = "sex")
# summary(fit_config, fit.measures = TRUE)
```

```{r ex6-solution, include=FALSE}
fit_config <- cfa(cfa_model, data = social_data, group = "sex")
summary(fit_config, fit.measures = TRUE)
```

```{r ex6-check}
grade_code()
```

```{r ex-config-q}
question(
  "Как можно интерпретировать результаты конфигуральной модели инвариантности (Chi-square = 46.13, df = 48, p = 0.55; CFI = 1.00; RMSEA = 0.000; SRMR = 0.034) для мужчин и женщин?",
  answer("Базовая факторная структура одинакова в обеих группах, и модель хорошо согласуется с данными.", correct = TRUE),
  answer("Факторная структура существенно различается между группами, модель плохо согласуется с данными."),
  answer("Модель приемлема только для мужчин, но не для женщин."),
  answer("По этим показателям нельзя судить о конфигуральной инвариантности."),
  allow_retry = TRUE,
  incorrect = "Незначимый χ², высокие значения CFI/TLI и низкие RMSEA и SRMR указывают на хорошую подгонку конфигуральной модели в обеих группах."
)
```

### Упражнение 7 — Метрическая и скалярная модели

- Оцените **метрическую** модель: `group.equal = "loadings"`, результат — `fit_metric`.
- Оцените **скалярную** модель: `group.equal = c("loadings", "intercepts")`, результат — `fit_scalar`.

```{r ex7, exercise = TRUE}
fit_metric <- cfa(
  cfa_model, data = social_data, group = "sex",
  group.equal = ___
)

fit_scalar <- cfa(
  cfa_model, data = social_data, group = "sex",
  group.equal = ___
)
```

```{r ex7-hint, eval=FALSE}
# fit_metric <- cfa(cfa_model, data = social_data, group = "sex",
#                   group.equal = "loadings")
# fit_scalar <- cfa(cfa_model, data = social_data, group = "sex",
#                   group.equal = c("loadings", "intercepts"))
```

```{r ex7-solution, include=FALSE}
fit_metric <- cfa(cfa_model, data = social_data, group = "sex",
                  group.equal = "loadings")
fit_scalar <- cfa(cfa_model, data = social_data, group = "sex",
                  group.equal = c("loadings", "intercepts"))
```

```{r ex7-check}
grade_this({
  if (!exists("fit_metric") || !exists("fit_scalar"))
    fail("Создайте оба объекта: fit_metric и fit_scalar.")
  if (!inherits(fit_metric, "lavaan") || !inherits(fit_scalar, "lavaan"))
    fail("Оба объекта должны быть моделями lavaan.")
  pass("Метрическая и скалярная модели оценены.")
})
```

### Упражнение 8 — Сравнение моделей инвариантности

Используйте `lavTestLRT()` для сравнения трех моделей инвариантности.

```{r ex8, exercise = TRUE}
lrt_result <- lavTestLRT(___, ___, ___)
lrt_result
```

```{r ex8-hint, eval=FALSE}
# lrt_result <- lavTestLRT(fit_config, fit_metric, fit_scalar)
# lrt_result
```

```{r ex8-solution, include=FALSE}
lrt_result <- lavTestLRT(fit_config, fit_metric, fit_scalar)
lrt_result
```

```{r ex8-check}
grade_this({
  if (!exists("lrt_result")) fail("Создайте объект lrt_result.")
  if (!inherits(lrt_result, "data.frame") && !inherits(lrt_result, "anova"))
    fail("Используйте функцию lavTestLRT() для сравнения моделей.")
  pass("Тест сравнения моделей выполнен.")
})
```

```{r ex8-q}
question(
  "«Если при переходе от метрической к скалярной модели Pr(>Chisq) < 0.05 и RMSEA = 0.083, что это означает?»",
  answer("Скалярная инвариантность не поддерживается: интерцепты различаются между группами.", correct = TRUE),
  answer("Скалярная инвариантность полностью подтверждена."),
  answer("Необходимо добавить еще больше ограничений."),
  answer("Данные имеют выбросы."),
  allow_retry = TRUE,
  incorrect = "Значимый хи-квадрат и увеличение RMSEA говорят о том, что интерцепты неодинаковы, т.е. полная скалярная инвариантность нарушена."
)
```

Действительно, получившиеся результаты указывают на то, что в нашей модели достингута только метрическая инвариантность.

- При **метрической, но не скалярной** инвариантности сравнивать можно **связи** (регрессионные коэффициенты, ковариации факторов), но сравнение латентных средних считается методологически проблемным, потому что часть различий может быть связана именно с сдвигами интерцептов индикаторов. [Источник для чтения](https://dlab.berkeley.edu/news/testing-measurement-invariance-using-lavaan-r)
- Если удается построить модель **частичной скалярной** инвариантности (часть интерцептов равны, часть свободны), сравнение сред по соответствующим факторам допускается, но интерпретация должна опираться только на «инвариантный» набор индикаторов. [pmc.ncbi.nlm.nih](https://pmc.ncbi.nlm.nih.gov/articles/PMC10801901/)

## Упражнения: Структурная модель

### Упражнение 9 — Задание структурной модели SEM

Задайте полную SEM-модель: измерительная часть (CFA) + структурные пути.

- `wellbeing` зависит от `trust` и `civic`;
- `civic` зависит от `trust`.

```{r ex9, exercise = TRUE}
sem_model <- '
  # Модель измерения
  trust     =~ trust1 + trust2 + trust3
  wellbeing =~ wb1    + wb2    + wb3
  civic     =~ civic1 + civic2 + civic3
  # Структурные пути (ваш код)
  ___
  ___
'
fit_sem <- sem(___, data = ___)
summary(fit_sem, fit.measures = TRUE, standardized = TRUE)
```

```{r ex9-hint, eval=FALSE}
# sem_model <- '
#   trust     =~ trust1 + trust2 + trust3
#   wellbeing =~ wb1 + wb2 + wb3
#   civic     =~ civic1 + civic2 + civic3
#   wellbeing ~ trust + civic
#   civic     ~ trust
# '
# fit_sem <- sem(sem_model, data = social_data)
# summary(fit_sem, fit.measures = TRUE, standardized = TRUE)
```

```{r ex9-solution, include=FALSE}
sem_model <- '
  trust     =~ trust1 + trust2 + trust3
  wellbeing =~ wb1    + wb2    + wb3
  civic     =~ civic1 + civic2 + civic3
  wellbeing ~  trust + civic
  civic     ~  trust
'
fit_sem <- sem(sem_model, data = social_data)
summary(fit_sem, fit.measures = TRUE, standardized = TRUE)
```

```{r ex9-check}
grade_this({
  if (!exists("fit_sem")) fail("Создайте объект fit_sem.")
  if (!inherits(fit_sem, "lavaan")) fail("fit_sem должен быть объектом lavaan.")
  pt <- lavaan::parameterTable(fit_sem)
  struct_paths <- pt[pt$op == "~", ]
  if (nrow(struct_paths) < 2) fail("Модель должна содержать хотя бы 2 структурных пути (~).")
  pass("Структурная модель SEM задана и оценена.")
})
```

### Упражнение 10 — Интерпретация путей SEM

Извлеките стандартизованные оценки параметров с помощью `parameterEstimates()`.

```{r ex10, exercise = TRUE}
struct_params <- parameterEstimates(
  fit_sem,
  standardized = ___,
  rsquare      = ___
) |>
  dplyr::filter(op == "~")

struct_params
```

```{r ex10-hint, eval=FALSE}
# struct_params <- parameterEstimates(
#   fit_sem,
#   standardized = TRUE,
#   rsquare      = TRUE
# ) |>
#   dplyr::filter(op == "~")
# struct_params
```

```{r ex10-solution, include=FALSE}
struct_params <- parameterEstimates(
  fit_sem,
  standardized = TRUE,
  rsquare      = TRUE
) |>
  dplyr::filter(op == "~")
struct_params
```

```{r ex10-check}
grade_this({
  if (!exists("struct_params")) fail("Создайте объект struct_params.")
  if (!is.data.frame(struct_params)) fail("struct_params должен быть датафреймом.")
  if (!"std.all" %in% names(struct_params))
    fail("Убедитесь, что standardized = TRUE.")
  if (nrow(struct_params) < 2) fail("Должны быть отфильтрованы структурные пути (~).")
  pass("Структурные пути извлечены и стандартизованы.")
})
```

```{r ex10-q}
question_radio(
  "Что означает стандартизованный коэффициент пути (std.all) 0.56 от trust к wellbeing?",
  answer("При увеличении доверия на 1 SD благополучие увеличивается на 0.56 SD, при контроле прочих переменных.", correct = TRUE),
  answer("40% дисперсии благополучия объясняется доверием."),
  answer("Доверие и благополучие коррелируют на уровне 0.40."),
  answer("p-значение пути равно 0.40."),
  allow_retry = TRUE,
  incorrect = "Стандартизованный бета-коэффициент показывает изменение зависимой переменной (в SD) при изменении независимой на 1 SD."
)
```

## Тестовые вопросы

### Вопрос 1

Что такое латентная переменная в SEM?

```{r q1}
question(
  "",
  answer("Переменная, которая не наблюдается напрямую, а выводится из набора измеримых индикаторов.", correct = TRUE),
  answer("Любая переменная, содержащая пропуски."),
  answer("Переменная, добавленная исследователем искусственно."),
  answer("Переменная, стандартизованная до нулевого среднего."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "Латентная переменная (конструкт) — ненаблюдаемая, измеряется косвенно через набор индикаторов."
)
```

### Вопрос 2

В чем принципиальное отличие CFA от EFA?

```{r q2}
question(
  "",
  answer("В CFA исследователь заранее задает, какие индикаторы загружаются на какие факторы; в EFA структура выявляется из данных.", correct = TRUE),
  answer("CFA применяется только к количественным данным; EFA — к категориальным."),
  answer("EFA требует больше наблюдений, чем CFA."),
  answer("В CFA нельзя использовать более двух факторов."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "CFA — подтверждающий подход с заданной структурой; EFA — разведочный, ищущий структуру в данных."
)
```

### Вопрос 3

Что означает оператор `=~` в синтаксисе lavaan?

```{r q3}
question(
  "",
  answer("Определение латентной переменной через индикаторы («измеряется через»).", correct = TRUE),
  answer("Регрессионный путь между двумя переменными."),
  answer("Ковариация между двумя переменными."),
  answer("Равенство двух параметров."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "`=~` — оператор определения фактора. Например, `trust =~ trust1 + trust2`."
)
```

### Вопрос 4

Что означает оператор `~` в синтаксисе lavaan?

```{r q4}
question(
  "",
  answer("Регрессионный путь: левая переменная зависит от правой.", correct = TRUE),
  answer("Определение латентной переменной."),
  answer("Ковариацию между двумя переменными."),
  answer("Интерцепт переменной."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "`~` — регрессия. `wellbeing ~ trust` означает: благополучие зависит от доверия."
)
```

### Вопрос 5

Какой индекс подгонки RMSEA считается хорошим?

```{r q5}
question(
  "",
  answer("RMSEA > 0.10"),
  answer("RMSEA < 0.06", correct = TRUE),
  answer("RMSEA = 1.0"),
  answer("RMSEA < 0.50"),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "RMSEA < 0.06 — хорошая подгонка; RMSEA < 0.08 — приемлемая; > 0.10 — плохая."
)
```

### Вопрос 6

Что тестирует **конфигуральная** модель инвариантности?

```{r q6}
question(
  "",
  answer("Что одинаковый набор факторов и паттерн нагрузок присутствует в обеих группах (нагрузки не ограничены).", correct = TRUE),
  answer("Что факторные нагрузки одинаковы в обеих группах."),
  answer("Что интерцепты одинаковы в обеих группах."),
  answer("Что дисперсии факторов одинаковы."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "Конфигуральная инвариантность — базовый уровень: проверяется лишь одинаковость конфигурации (структуры) модели."
)
```

### Вопрос 7

Что позволяет делать **скалярная** инвариантность, но не позволяет метрическая?

```{r q7}
question(
  "",
  answer("Сравнивать средние значения латентных конструктов между группами.", correct = TRUE),
  answer("Сравнивать факторные нагрузки между группами."),
  answer("Запускать SEM с ковариатами."),
  answer("Исключать плохо подходящие индикаторы."),
  allow_retry = TRUE,
  incorrect = "Для корректного сравнения средних латентных факторов нужна скалярная инвариантность (равенство и нагрузок, и интерцептов)."
)
```

### Вопрос 8

Какой тест используется в lavaan для сравнения вложенных моделей инвариантности?

```{r q8}
question(
  "",
  answer("lavTestLRT() — тест отношения правдоподобия для вложенных моделей.", correct = TRUE),
  answer("lavTestScore() — тест модификационных индексов."),
  answer("anova() — только для lm-моделей."),
  answer("chisq.test() — тест хи-квадрат для таблиц сопряженности."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "lavTestLRT() проверяет значимость ухудшения подгонки при добавлении ограничений."
)
```

### Вопрос 9

Какой индекс подходит как дополнительный критерий при сравнении моделей инвариантности?

```{r q9}
question(
  "",
  answer("$\\Delta CFI$: снижение на 0.010 и более свидетельствует об ухудшении подгонки.", correct = TRUE),
  answer("$R^2$ общей модели."),
  answer("Коэффициент Кронбаха $\\alpha$."),
  answer("Инерция первой оси в MCA."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "$\\Delta CFI$ ≥ −0.010 при добавлении ограничений — сигнал возможного нарушения инвариантности."
)
```

### Вопрос 10

Что такое «идентификация модели» в SEM?

```{r q10}
question(
  "",
  answer("Условие, при котором количество известных величин достаточно для уникального оценивания всех свободных параметров.", correct = TRUE),
  answer("Способность модели предсказывать новые наблюдения."),
  answer("Совпадение факторной структуры с теоретической."),
  answer("Нормализация данных перед оцениванием."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "Идентифицированная модель имеет df ≥ 0 и каждый параметр может быть оценен единственным образом."
)
```

### Вопрос 11

Как в lavaan по умолчанию фиксируется масштаб латентной переменной при CFA?

```{r q11}
question(
  "",
  answer("Нагрузка первого индикатора фиксируется равной 1 (marker indicator).", correct = TRUE),
  answer("Дисперсия фактора фиксируется равной 0."),
  answer("Среднее фактора фиксируется равным 100."),
  answer("Все нагрузки суммируются до 1."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "По умолчанию lavaan фиксирует нагрузку первого индикатора = 1, задавая шкалу латентной переменной."
)
```

### Вопрос 12

Что такое «модификационные индексы» (modification indices) в SEM?

```{r q12}
question(
  "",
  answer("Ожидаемое снижение $\\chi^2$ при освобождении каждого зафиксированного параметра; помогают улучшить подгонку.", correct = TRUE),
  answer("Оценки факторных нагрузок из EFA."),
  answer("Коэффициенты стандартизации переменных."),
  answer("Метрика качества измерения Cronbach $\\chi^2$."),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  incorrect = "modindices(fit) показывает, насколько снизился бы χ² при освобождении каждого ограниченного параметра."
)
```

### Вопрос 13

Что такое «медиация» в контексте SEM?

```{r q13}
question(
  "",
  answer("Косвенный эффект: переменная A влияет на C через посредника B (A → B → C).", correct = TRUE),
  answer("Прямой путь от экзогенной переменной к эндогенной."),
  answer("Ограничение нагрузок двух индикаторов равными значениями."),
  answer("Метод оценивания модели вместо ML."),
  allow_retry = TRUE,
  incorrect = "Медиация — когда влияние X на Y частично или полностью реализуется через посредник M."
)
```

### Вопрос 14

Как вычислить косвенный эффект в lavaan?

```{r q14}
question(
  "",
  answer("Через метки параметров и произведение путей: `a := b * c` с бутстрапом для CI.", correct = TRUE),
  answer("Только с помощью пакета mediation."),
  answer("Прямым вычитанием прямого эффекта из общего."),
  answer("Функцией indirect.effect() из lavaan."),
  allow_retry = TRUE,
  incorrect = "В lavaan косвенный эффект задается через `:=` с произведением меченых путей и оценивается бутстрапом."
)
```

### Вопрос 15

Что означает R² для латентной переменной (эндогенного конструкта) в SEM?

```{r q15}
question(
  "",
  answer("Доля дисперсии конструкта, объясняемая предикторами в структурной модели.", correct = TRUE),
  answer("Надежность измерения фактора (альфа Кронбаха)."),
  answer("Процент инерции, объясняемой данным конструктом."),
  answer("Корреляция между двумя латентными факторами."),
  allow_retry = TRUE,
  incorrect = "R² эндогенного конструкта — это квадрат множественной корреляции: сколько его дисперсии объясняют входящие пути."
)
```

## Источники {.unnumbered}

::: {#refs}
::: 

```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
