---
title: "Кластерный анализ"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)

knitr::opts_chunk$set(echo=FALSE)


tutorial_options(
  exercise.reveal_solution = FALSE
)

library(tclust)
library(factoextra)
library(class)
set.seed(123)
data("wholesale")
wh_num <- wholesale[,c("Fresh","Milk","Grocery","Frozen","Detergent","Delicatessen")]
wh_scaled <- scale(wh_num)
km_wh4 <- kmeans(wh_scaled, centers = 4, nstart = 25)
channel <- factor(wholesale$Channel, labels = c("Horeca","Retail"))
dist_wh <- dist(wh_scaled, method = "euclidean")
hc_wh <- hclust(dist_wh, method = "ward.D2")
hc_clusters4 <- cutree(hc_wh, k = 4)
n <- nrow(wh_scaled)
idx_train <- sample(1:n, size = round(0.7 * n))
train_x <- wh_scaled[idx_train, ]
train_y <- channel[idx_train]
test_x  <- wh_scaled[-idx_train, ]
test_y  <- channel[-idx_train]
n <- nrow(wh_scaled)
pred_knn5 <- knn(train = train_x, test = test_x, cl = train_y, k = 5)
tab_knn5 <- table(pred_knn5, test_y)
acc_knn5 <- sum(diag(tab_knn5)) / sum(tab_knn5)



```

```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Предварительные замечания

Кластерный анализ — это группа методов, позволяющих **разделить объекты на однородные группы (кластеры)** на основе их сходства по набору переменных.
В этом тьюториале рассматриваются три классических подхода: k-means, иерархический кластерный анализ и метод k-ближайших соседей (k-nearest neighbors, kNN), применяемых для решения практических задач.

### Задачи кластеризации в социальных исследованиях

В прикладных социальных исследованиях кластерный анализ часто используется для содержательной сегментации респондентов, когда исследователя интересуют не столько отдельные переменные, сколько устойчивые сочетания установок и практик, образующие качественно различающиеся типы. 

- Так, в случае изучения стилей жизни и потребительского поведения исследователь может исходить из ряда вопросов о ценностях, досуге, потребительских стратегиях и поведения в сфере медиа. Применяя методы кластеризации, можно выделить группы респондентов с различными профилями ответов. Тогда кластеры интерпретируются как «традиционалисты», ориентирующиеся на семейные ценности и умеренное потребление, «карьеристы», для которых центральными оказываются профессиональный успех и индивидуальные достижения, «ориентированные на семью» респонденты с высокой значимостью родительства и домашнего досуга и другие типы, задаваемые эмпирически. Важно, что в этом подходе типология возникает не «по замыслу автора», а как результат структурных закономерностей в данных, после чего исследователь уже соотносит полученные кластеры с теоретическими моделями и контекстом исследования.

- Аналогичная логика может быть применена для анализа территориальной дифференциации. Имея набор индикаторов социально‑экономического благополучия, инфраструктурной обеспеченности, уровня гражданской активности и доверия к институтам, исследователь может рассматривать не отдельные показатели, а их совместные конфигурации на уровне регионов, муниципалитетов или районов. Кластерный анализ в этом случае позволяет выделить группы территорий, схожие по профилю: условно «развитые центры» с высоким уровнем услуг и активным третьим сектором, «стабильная периферия» с умеренными значениями по большинству индикаторов и «уязвимые территории» с дефицитами как в инфраструктуре, так и в гражданском участии. Полученные кластеры могут использоваться для таргетирования государственной или муниципальной политики, сравнения динамики групп во времени, а также для более точной постановки исследовательских вопросов о факторах территориального неравенства.

- Наконец, отдельное направление применения кластерного анализа связано с поиском типологий организаций – будь то некоммерческие организации, школы, медицинские учреждения или другие учреждения социальной сферы В этом случае кластерный анализ может быть проведен по совокупности переменных: ресурсная база, структура персонала и добровольцев, набор услуг, целевые группы, модели взаимодействия с государством и бизнесом. Кластеризация позволяет выделять, например, «ресурсно обеспеченные профессиональные НКО» с устойчивым финансированием и развитой инфраструктурой, «локальные инициативные группы» с ограниченными ресурсами, но высоким уровнем вовлеченности волонтеров, или «гибридные организации», совмещающие черты классической НКО и социального предпринимательства. Такая типология не только упорядочивает эмпирическое разнообразие, но и служит основой для дальнейшего анализа: сравнения устойчивости различных типов, их роли в локальных сетях, реакций на институциональные изменения и т.д.

Таким образом, имея в своем арсенале методы кластерного анализа, у исследователя появляется возможность значительно расширить рамки своего проекта, и разработать аналитические решения, которые могут помочь сделать управленческие решения (в бизнесе, государственном управлении и некоммерческом секторе) более адресными, а, значит, более эффективными.



## Тестовые вопросы

### Вопрос 1

Какова основная цель методов кластерного анализа?

```{r q1}
question(
  "",
  answer("Предсказать значения количественной зависимой переменной."),
  answer("Разделить объекты на группы по сходству признаков.", correct = TRUE),
  answer("Проверить статистические гипотезы о средних значениях."),
  answer("Оценить причинно-следственные связи."),
  incorrect = "Кластеризация не предсказывает исход, а группирует объекты по их сходству.",
  allow_retry = TRUE
)
```

### Вопрос 2

Что необходимо задать заранее в классическом алгоритме k-means?

```{r q2}
question(
  "",
  answer("Метку класса для каждого объекта."),
  answer("Число кластеров k.", correct = TRUE),
  answer("Только метрику расстояния."),
  answer("Порядок объектов в таблице."),
  incorrect = "k-means требует заранее заданного числа кластеров k.",
  allow_retry = TRUE
)
```

### Вопрос 3

Как обычно интерпретируется центр кластера в k-means?

```{r q3}
question(
  "",
  answer("Как медиана всех переменных в выборке."),
  answer("Как точка, наиболее удаленная от других кластеров."),
  answer("Как среднее значение признаков объектов, вошедших в кластер.", correct = TRUE),
  answer("Как случайно выбранный объект."),
  incorrect = "Центр (центроид) — это вектор средних значений признаков по объектам кластера.",
  allow_retry = TRUE
)
```

### Вопрос 4

Какой недостаток k-means делает важной предварительную стандартизацию переменных?

```{r q4}
question(
  "",
  answer("Нечувствительность к масштабу."),
  answer("Чувствительность к масштабу переменных.", correct = TRUE),
  answer("Неспособность работать с большими выборками."),
  answer("Отсутствие случайности в алгоритме."),
  incorrect = "Переменные с большим масштабом доминируют расстояния, поэтому их часто приводят к общему масштабу.",
  allow_retry = TRUE
)
```

### Вопрос 5

Что представляет собой дендрограмма в иерархическом кластерном анализе?

```{r q5}
question(
  "",
  answer("График остатков регрессии."),
  answer("Дерево, показывающее последовательность объединения (или разделения) кластеров.", correct = TRUE),
  answer("Корреляционная матрица переменных."),
  answer("Диаграмма рассеяния для двух переменных."),
  incorrect = "Дендрограмма визуализирует структуру вложенных кластеров и высоту их объединения.",
  allow_retry = TRUE
)
```

### Вопрос 6

К какому типу относится метод k-nearest neighbors (kNN)?

```{r q6}
question(
  "",
  answer("Метод несупервизируемого обучения."),
  answer("Метод супервизируемого обучения для классификации.", correct = TRUE),
  answer("Метод факторного анализа."),
  answer("Метод оценки плотности."),
  incorrect = "kNN использует обучающую выборку с метками классов для классификации новых объектов.",
  allow_retry = TRUE
)
```

### Вопрос 7

Как определяется класс нового наблюдения в kNN?

```{r q7}
question(
  "",
  answer("По медиане всех признаков."),
  answer("По классу наиболее удаленного соседа."),
  answer("По большинству классов среди k ближайших соседей.", correct = TRUE),
  answer("Случайным образом."),
  incorrect = "kNN ищет k ближайших наблюдений и присваивает класс, наиболее часто встречающийся среди них.",
  allow_retry = TRUE
)
```

### Вопрос 8

Какой из подходов НЕ относится к иерархической кластеризации?

```{r q8}
question(
  "",
  answer("Агломеративный (bottom-up)."),
  answer("Дивисивный (top-down)."),
  answer("Метод k-ближайших соседей.", correct = TRUE),
  answer("Метод одиночной/полной связи."),
  incorrect = "kNN — это метод классификации, а не иерархической кластеризации.",
  allow_retry = TRUE
)
```

### Вопрос 9

Почему важно интерпретировать кластеры с использованием исходных переменных?

```{r q9}
question(
  "",
  answer("Чтобы проверить нормальность остатков."),
  answer("Чтобы понять содержательное значение кластеров как групп объектов с похожими профилями.", correct = TRUE),
  answer("Чтобы вычислить p-значения."),
  answer("Чтобы избежать мультиколлинеарности."),
  incorrect = "Кластеры сами по себе — результат алгоритма; интерпретация требует анализа профилей переменных внутри кластеров.",
  allow_retry = TRUE
)
```

### Вопрос 10

Какой вывод можно сделать, если в дендрограмме большие «скачки» расстояний происходят при объединении последних нескольких кластеров?

```{r q10}
question(
  "",
  answer("Лучше оставить число кластеров примерно до этого скачка.", correct = TRUE),
  answer("Нужно всегда использовать максимальное число кластеров."),
  answer("Данные не подходят для кластеризации."),
  answer("Нужно уменьшить размер выборки."),
  incorrect = "Резкий рост расстояний при объединении кластеров часто используется как эвристика для выбора числа кластеров.",
  allow_retry = TRUE
)
```


### Вопрос 11

Каково ключевое отличие k-means от kNN?

```{r q11}
question(
  "",
  answer("k-means — метод кластеризации без учителя, тогда как kNN — супервизируемый метод классификации.", correct = TRUE),
  answer("Оба метода — без учителя"),
  answer("Оба метода применяются только для визуализации различий между группами."),
  answer("Оба метода используются только для временных рядов."),
  incorrect = "k-means не использует метки классов, а kNN опирается на обучающую выборку с известными классами.",
  allow_retry = TRUE
)
```

### Вопрос 12

Какой шаг НЕЛЬЗЯ пропускать при работе с расстояниями между переменными разного масштаба?

```{r q12}
question(
  "",
  answer("Предварительная стандартизация или нормализация переменных.", correct = TRUE),
  answer("Проверка автокорреляции."),
  answer("Выбор уровня значимости."),
  answer("Вычисление медиан."),
  incorrect = "Без стандартизации переменные с большим масштабом доминируют расстояния и влияют на результаты кластеризации и kNN.",
  allow_retry = TRUE
)
```


### Вопрос 13

Какая из следующих формулировок лучше всего описывает различие между мерой расстояния и мерой сходства?

```{r q13}
question(
  "",
  answer("Мера расстояния достигает максимума для идентичных объектов, мера сходства — минимума."),
  answer("Мера расстояния возрастает по мере различия объектов, мера сходства уменьшается при увеличении различий.", correct = TRUE),
  answer("Мера сходства всегда равна 1 минус мера расстояния."),
  answer("Мера расстояния используется только для количественных переменных, а мера сходства — только для качественных."),
  incorrect = "Подумайте о том, как интерпретируются большие и малые значения этих мер.",
  allow_retry = TRUE
)
```

### Вопрос 14

Чем евклидово расстояние принципиально отличается от манхэттенского (city-block) при кластеризации?

```{r q14}
question(
  "",
  answer("Евклидово расстояние учитывает только одну переменную, манхэттенское — несколько."),
  answer("Евклидово расстояние вычисляется как корень из суммы квадратов разностей, манхэттенское — как сумма модулей разностей по координатам.", correct = TRUE),
  answer("Манхэттенское расстояние всегда больше евклидова."),
  answer("Отличий нет, это полностью эквивалентные метрики."),
  incorrect = "Вспомните геометрическую интерпретацию обеих метрик.",
  allow_retry = TRUE
)
```

### Вопрос 15

Как влияет сильная корреляция между признаками на кластеризацию с использованием евклидова расстояния?

```{r q15}
question(
  "",
  answer("Никак, корреляция переменных не влияет на результаты кластеризации."),
  answer("Кластеризация может отражать в первую очередь структуру корреляций, а не содержательные различия между объектами.", correct = TRUE),
  answer("Евклидово расстояние автоматически устраняет влияние корреляции."),
  answer("Кластеризация становится невозможной из‑за сингулярности матрицы расстояний."),
  incorrect = "Подумайте, что происходит, когда несколько переменных фактически измеряют один и тот же латентный признак.",
  allow_retry = TRUE
)
```

### Вопрос 16

Чем отличаются методы `single linkage` и `complete linkage` в иерархической кластеризации?

```{r q16}
question(
  "",
  answer("Single linkage использует минимальное расстояние между элементами кластеров, complete linkage — максимальное.", correct = TRUE),
  answer("Оба метода используют только центры кластеров."),
  answer("Single linkage строит только один кластер, complete linkage — несколько."),
  answer("Отличий нет, это разные названия одного и того же алгоритма."),
  incorrect = "Обратите внимание, какие пары точек учитываются при объединении кластеров.",
  allow_retry = TRUE
)
```

### Вопрос 17

В чем ключевая идея метода Уорда (Ward) в агломеративной кластеризации?

```{r q17}
question(
  "",
  answer("Минимизировать суммарное расстояние между всеми парами объектов в выборке."),
  answer("На каждом шаге объединять те кластеры, которые дают наименьшее увеличение суммы квадратов внутри кластера.", correct = TRUE),
  answer("Объединять кластеры по случайному принципу."),
  answer("Строить кластеры фиксированного размера."),
  incorrect = "Метод Уорда на каждом шаге минимизирует рост внутрикластерной вариации.",
  allow_retry = TRUE
)
```

### Вопрос 18

Что измеряет силуэтная мера (silhouette) для отдельного объекта?

```{r q18}
question(
  "",
  answer("Разность между максимальным и минимальным значением признака."),
  answer("Насколько объект ближе к своему кластеру, чем к соседнему кластеру.", correct = TRUE),
  answer("Количество соседей в радиусе фиксированного расстояния."),
  answer("Вероятность принадлежности объекту к каждому кластеру."),
  incorrect = "Силуэт сравнивает средние расстояния до собственного и ближайшего «чужого» кластера.",
  allow_retry = TRUE
)
```

### Вопрос 19

Как интерпретировать среднее значение силуэтной меры, близкое к 1?

```{r q19}
question(
  "",
  answer("Объекты плохо соответствуют своим кластерам и находятся на границах."),
  answer("Кластеры хорошо разделены, объекты в целом ближе к своим кластерам, чем к соседним.", correct = TRUE),
  answer("Кластеризация случайна и не имеет структуры."),
  answer("Это всегда признак переобучения модели."),
  incorrect = "Вспомните, как определяется силуэт: большие значения соответствуют чёткой кластерной структуре.",
  allow_retry = TRUE
)
```

### Вопрос 20

Как можно использовать силуэтную меру для выбора числа кластеров $k$ в алгоритме k-means?

```{r q20}
question(
  "",
  answer("Выбрать k, при котором средний силуэт максимален.", correct = TRUE),
  answer("Выбрать k, при котором средний силуэт минимален."),
  answer("Всегда брать k = числу переменных."),
  answer("Силуэтная мера не зависит от k."),
  incorrect = "Обычно сравнивают средний силуэт по ряду значений k и выбирают максимум, но не забывают о содержательном смысле кластеров.",
  allow_retry = TRUE
)
```

## Упражнения; Алгоритм k-means

Алгоритм $k$-means решает задачу разбиения $n$ объектов на $k$ кластеров так, чтобы минимизировать сумму квадратов расстояний объектов до центров кластеров.

**Целевая функция**

Пусть у нас есть какое-то количество объектов $x_n$, которые описываются некоторым количеством характеристик $p$,значимых для будущей типологии. Математически это можно записать так: $x_1,\dots,x_n \in \mathbb{R}^p$, а число кластеров – $k$.  
Обозначим центр кластера $j$ как $\mu_j$, а множество объектов в кластере как $j$ — $C_j$.

При применении метода k-means (k средних), максимизируется функция:
$$
J = \sum_{j=1}^k \sum_{x_i \in C_j} \|x_i - \mu_j\|_2^2.
$$

То есть, простыми словами, мы ищем такое разбиение респондентов или организаций на $k$ групп и такие центры этих групп, при которых объекты внутри своих кластеров в среднем находятся как можно ближе к центру своего кластера.

Еще более по‑социологическ: допустим мы создаем типологию пользователей социального сервиса на основе анализа результатов опроса. Тогда:

- каждый респондент представлен набором количественных признаков (возраст, доход, расходы по категориям и т.п.);

- алгоритм подбирает типологию (кластеры) так, чтобы «средний респондент» внутри каждого типа был как можно более «похож» на остальных в своем типе, а разброс внутри типов должен быть как можно меньшим.


Разберем процесс по шагам:


**Шаг 1. Инициализация**

Выбираются начальные центры кластеров $\mu_1^{(0)},\dots,\mu_k^{(0)}$ (обычно случайно из набора наблюдений).

**Шаг 2. Присвоение объектов кластерам**

Для каждого объекта $x_i$ находится ближайший центр и объект присваивается соответствующему кластеру:
$$
c_i^{(t)} = \arg\min_{j \in \{1,\dots,k\}} \|x_i - \mu_j^{(t)}\|^2.
$$

**Шаг 3. Пересчет центров кластеров (update)**

Для каждого кластера $j$ новый центр вычисляется как среднее всех объектов, отнесенных к данному кластеру:
$$
\mu_j^{(t+1)} = \frac{1}{|C_j^{(t)}|} \sum_{x_i \in C_j^{(t)}} x_i.
$$

**Шаг 4. Критерий остановки**

Шаги 2–3 повторяются, пока не выполнится одно из условий остановки:

- присвоения $c_i$ не перестанут меняться;
- изменение центров $\mu_j^{(t)}$ или значения $J$ не станет меньше заданного порога;
- не будет достигнуто максимальное число итераций.

Результат: разбиение на $k$ кластеров и центры $\mu_1,\dots,\mu_k$, которые локально минимизируют $J$.

Вот как это может выглядеть:

![](images\k-means.gif){width=300px} ![](images\Kmeans2.gif){width=300px}

Разберем пример с набором данных `Wholesale customers`. Этот датасет описывает клиентов оптового дистрибьютора (Португалия) и содежит следующие переменные:

- `Channel` — тип клиента: 1 = Horeca (отели, рестораны, кафе), 2 = Retail (розница).
- `Region` — регион: 1 = Lisbon, 2 = Oporto, 3 = Other.
- `Fresh` — годовые траты на свежие продукты.
- `Milk` — траты на молочную продукцию.
- `Grocery` — бакалея (сухие и долгохранящиеся продукты).
- `Frozen` — замороженные продукты
- `Detergent` — моющие средства.
- `Delicatessen` — деликатесы.

Всего 440 клиентов, 8 переменных, пропусков нет.

Используем числовые переменные набора `wholesale` (траты по товарным категориям) и их стандартизованные версии `wh_scaled`.

### Упражнение 1 — Просмотр и базовая статистика

Посмотрите на первые строки числовых данных `wh_num` и на основные статистики по переменным `summary()`.

```{r exk1, exercise = TRUE}
# выведите первые строки wh_num
# и сводные статистики с помощью summary()
```

```{r exk1-solution, include=FALSE}
head(wh_num)
summary(wh_num)
```

```{r exk1-check}
grade_code()
```


### Упражнение 2 — Запуск k-means с k = 4

Выполните кластеризацию k-means на стандартизованных данных `wh_scaled` с помощью метода `kmeans()`  с числом кластеров `centers = 4` и `nstart = 25`.  
Сохраните результат в объект `km_wh4`.

```{r exk2, exercise = TRUE}
# ваш код
```

```{r exk2-solution, include=FALSE}
km_wh4 <- kmeans(wh_scaled, centers = 4, nstart = 25)
```

```{r exk2-check}
grade_this({
  if (!exists("km_wh4")) fail("Создайте объект km_wh4 с результатами kmeans().")
  if (!"kmeans" %in% class(km_wh4)) fail("Проверьте, что вы использовали функцию kmeans().")
  pass("k-means с k = 4 выполнен.")
})
```


### Упражнение 3 — Размеры и центры кластеров

1. Посмотрите на количество наблюдений в каждом кластере (`km_wh4$size`).  
2. Выведите матрицу центров кластеров (`km_wh4$centers`).

```{r exk3, exercise = TRUE}
# ваш код
```

```{r exk3-solution, include=FALSE}
km_wh4$size
km_wh4$centers
```

```{r exk3-check}
grade_code()
```

```{r exk3-q}
question_numeric(
  "Сколько клиентов оказалось в самом маленьком кластере?",
  answer(as.character(min(km_wh4$size)), correct = TRUE),
  allow_retry = TRUE,
  try_again = "Подсказка: используйте min(km_wh4$size)."
)
```


### Упражнение 4 — Визуализация кластеров

Постройте визуализацию кластеров, используя функцию `fviz_cluster()` из пакета `factoextra`, с данными `wh_scaled` и результатом `km_wh4`.

```{r exk4, exercise = TRUE}
# ваш код
```

```{r exk4-solution, include=FALSE}
fviz_cluster(km_wh4, data = wh_scaled)
```

```{r exk4-check}
grade_code()
```

```{r exk4-q}
question_radio(
  "Что показывает этот график кластеров?",
  answer("Как клиенты группируются по двум главным компонентам, с выделением кластеров k-means.", correct = TRUE),
  answer("Только распределения одной переменной."),
  answer("Только дендрограмму иерархического анализа."),
  allow_retry = TRUE
)
```


- На графике мы видим, что в **первом кластере**  у нас только один клиент с особыми потребностями, чьи значения далеко отстоят от других клиентов. Он тратит аномально много по всем категориям (особенно на деликатесы, заморозку, молочные и свежие продукты) даже относительно других оптовых покупателей. По сути это один «сверхтяжелый» клиент, который настолько выбивается по объемам закупок, что k-means выделяет его в отдельный кластер, чтобы не портить структуру остальных. Он настолько специфичен, что его стоит выделить в отдельный кейс, анализировать отдельно, поскольку стастически - это определенно «выброс».

- Во **втором кластере** центральные значения все слегка отрицательные. Интерпретация: это клиенты с объемом закупок ниже среднего по всем категориям — небольшие/редко покупающие клиенты без выраженной специализации.

В **третьем кластере** по центральным значениям Milk ≈ 1.85, Grocery ≈ 2.22, Detergent ≈ 2.25 (сильно выше среднего), Frozen слегка ниже 0, Fresh немного ниже 0.
Интерпретация: клиенты, делающие крупные закупки в «сухих» категориях — молочные продукты, бакалея, моющие средства, но не особенно выделяющиеся по свежим продуктам и заморозке. Типично похоже на розничные магазины/супермаркеты.

- У **четвертого кластера** Fresh ≈ 1.6 и Frozen ≈ 1.0 (выше среднего), Milk и Grocery чуть ниже или около нуля, Detergent чуть ниже среднего, Delicatessen немного выше нуля. Это клиенты с повышенными закупками свежих и замороженных продуктов, менее акцентированные на бакалее и бытовой химии. 

### Упражнение 5 — Сравнение кластеров по типу клиента (Channel)

Сравните полученные кластеры с переменной `Channel` (1 = Horeca, 2 = Retail):

1. Постройте таблицу сопряженности `tab_km_channel` между `km_wh4$cluster` и `wholesale$Channel`.  
2. Выведите на экран `tab_km_channel` и проинтерпретируйте, какие кластеры ближе к Horeca, а какие — к розничной торговле.

```{r exk5, exercise = TRUE}
# ваш код
```

```{r exk5-solution, include=FALSE}
tab_km_channel <- table(km_wh4$cluster, wholesale$Channel)
tab_km_channel
```

```{r exk5-check}
grade_code()
```

```{r exk5-q}
question_text(
  "Кратко опишите, какие кластеры в большей степени состоят из клиентов Horeca, а какие — из розничных клиентов.",
  answer(" ", correct = TRUE)
)
```


## Упражнения: иерархический кластерный анализ


Иерархический кластерный анализ строит дерево вложенных кластеров (дендрограмму), показывая, как объекты последовательно объединяются (*агломеративный подход*) или разделяются (*дивизимный подход*). В отличие от k-means, число кластеров заранее не фиксируется: его выбирают позже, «обрезая» дендрограмму на нужной высоте. Результаты иерархического кластерного анализа могут разниться и зависят от того, как определяются расстояния между объектами и какой метод выбран для определения границ кластера.

### Расстояния между объектами

Пусть объекты описаны векторами признаков $x_i = (x_{i1},\dots,x_{ip})$, $x_l = (x_{l1},\dots,x_{lp})$. Наиболее распространенными являются следующие метрики:

- **Евклидово расстояние**
  $$
  d_{\text{E}}(x_i,x_l) = \sqrt{\sum_{m=1}^p (x_{im} - x_{lm})^2}.
  $$
  
 Где: 
- $x_i = (x_{i1},\dots,x_{ip})$ — вектор признаков **для объекта номер \(i\)**;  
- $x_l = (x_{l1},\dots,x_{lp})$ — вектор признаков **для объекта номер \(l\)**;  
- индекс $m$ в сумме — номер признака (переменной) от 1 до $p$.


То есть для двух объектов (респондентов, организаций и пр.)
по каждому признаку $m$ считается разность, возводится в квадрат, считается сумма и далее извлекается корень. Индексы нужны лишь для того, чтобы явно различать «номер объекта» и «номер признака».

Данная метрика хорошо работает, когда признаки сравнимы по масштабу (часто после стандартизации), чувствительно к выбросам и сильным различиям по отдельным координатам.

- **Манхеттенское (city-block) расстояние**
  $$
  d_{\text{M}}(x_i,x_l) = \sum_{m=1}^p |x_{im} - x_{lm}|.
  $$
  Менее чувствительно к крупным отклонениям одной переменной, удобно, когда различия по признакам складываются аддитивно (суммарная «разница по всем шкалам»).

- **Расстояние Чебышева**
  $$
  d_{\text{Ch}}(x_i,x_l) = \max_{1 \le m \le p} |x_{im} - x_{lm}|.
  $$

Ориентируется на максимальное различие по одной координате, полезно, когда важно учитывать «наибольшую разницу» между объектами (например, если один критический показатель доминирует в оценке близости).

![](images\dist.png){width=300px}

### Другие расстояния

Помимо евклидова, манхеттенского и расстояния Чебышева, в кластерном анализе часто используют несколько других типов расстояний и мер «несходства».

#### Квадрат евклидова расстояния

Иногда вместо евклидова расстояния используют его квадрат без извлечения корня:

$d_{\text{E2}}(x,y) = \sum_{m=1}^p (x_m - y_m)^2.$

В этом случае, крупные отличия по отдельным признакам получают еще больший вес; удобно в алгоритмах, которые прямо работают с суммой квадратов (метод Уорда, k-means).

#### Косинусное расстояние

Сначала определяют косинусное сходство между объектами:

$\cos(x,y) = \dfrac{\sum_{m=1}^p x_m y_m}{\lVert x \rVert \,\lVert y \rVert},$

где $\lVert x \rVert$ — длина вектора $x$. Тогда косинусное расстояние можно задать как

$d_{\text{cos}}(x,y) = 1 - \cos(x,y).$

Важно, когда интересует **форма профиля** признаков, а не абсолютный масштаб (направление вектора, а не его длина).

#### Расстояния для бинарных (0–1) признаков

Если признаки представлены как «есть/нет» (0–1), используют меры, основанные на совпадениях/расхождениях. Пусть для двух объектов:

- $a$ — число признаков, где оба равны 1;
- $b$ — число признаков, где первый = 1, второй = 0;
- $c$ — число признаков, где первый = 0, второй = 1;
- $d$ — число признаков, где оба равны 0.

**Мера Жаккара (Jaccard)**

Сходство:

$s_{\text{J}} = \dfrac{a}{a + b + c}.$

Расстояние:

$d_{\text{J}} = 1 - s_{\text{J}} = \dfrac{b + c}{a + b + c}.$

Игнорирует совпадения «0–0», акцентирует совместные «единицы».

**Sokal–Michener**

Сходство:

$s_{\text{SM}} = \dfrac{a + d}{a + b + c + d},$

расстояние:

$d_{\text{SM}} = 1 - s_{\text{SM}}.$

Учитывает и совпадения «1–1», и совпадения «0–0».

#### Расстояние Гауэра для смешанных данных

При совместном анализе количественных, порядковых и бинарных переменных используют расстояние Гауэра. Для каждого признака $m$ считается частичное несходство $d_m(x,y)$ (нормированное от 0 до 1), затем берется среднее:

$d_{\text{Gower}}(x,y) = \dfrac{\sum_{m=1}^p w_m d_m(x,y)}{\sum_{m=1}^p w_m},$

где $w_m$ — вес (обычно 1 при отсутствии пропуска и 0 при пропуске).

Примеры частичных расстояний:

- для количественных признаков:  
  $d_m(x,y) = \dfrac{|x_m - y_m|}{R_m},$  
  где $R_m$ — диапазон (max − min) по признаку $m$;
- для бинарных и номинальных: $d_m(x,y) = 0$ при совпадении и $1$ при различии;
- для порядковых: сначала перевод в $[0;1]$, затем как количественный.

#### Шпаргалка по выбору мер сходства/близости

- Только количественные признаки: семейство $L_p$ (манхеттен, евклидово, квадрат евклидова) — после стандартизации.  
- Бинарные признаки: Жаккар (если важны «единицы»), Sokal–Michener (если важны и «0», и «1»).  
- Смешанные данные: расстояние Гауэра, позволяющее совместно учитывать разные типы шкал.

На практике в иерархическом анализе чаще всего используют евклидово или манхеттенское расстояние; выбор метрики должен соответствовать смыслу признаков и способу, которым исследователь определяет «похожесть».

### Агломеративные и дивизимные алгоритмы

1. **Агломеративный иерархический анализ (bottom-up)**  
   - старт: каждый объект образует отдельный кластер;  
   - на каждом шаге находятся **два наиболее близких кластера** и объединяются в новый;  
   - процесс продолжается до тех пор, пока не останется один кластер, содержащий все объекты.  

2. **Дивизимный иерархический анализ (top-down)**  
   - старт: все объекты в одном кластере;  
   - на каждом шаге один кластер разбивается на два более мелких (по некоторому критерию ухудшения/улучшения качества);  
   - процесс продолжается, пока не будут получены отдельные объекты или не выполнен критерий остановки.  

В прикладной практике и в стандартных реализациях в R чаще используется именно **агломеративный** вариант.

### Модели объединения кластеров (linkage)

При агломеративном анализе необходимо определить, как считать расстояние между кластерами $A$ и $B$, если известно расстояние между отдельными объектами этих кластеров.

Пусть $d(x_i,x_l)$ — расстояние между объектами. Тогда:

- **Метод одиночной связи (single linkage)**  
  $$
  D_{\text{single}}(A,B) = \min_{x_i \in A,\; x_l \in B} d(x_i,x_l).
  $$
Ориентируется на наиболее близкую пару объектов. Стремится цеплять кластеры «по цепочке», может приводить к эффекту «цепей» (один длинный вытянутый кластер).
  


- **Метод полной связи (complete linkage)**
  $$
  D_{\text{complete}}(A,B) = \max_{x_i \in A,\; x_l \in B} d(x_i,x_l).
  $$

Ориентируется на наиболее далекую пару объектов. Стремится формировать более **компактные и плотные кластеры**, но может разделять «естественные» группы, если внутри них есть удаленные точки.

- **Метод средней связи (average linkage)**
  $$
  D_{\text{average}}(A,B) = \frac{1}{|A|\;|B|} \sum_{x_i \in A} \sum_{x_l \in B} d(x_i,x_l).
  $$

Берет среднее расстояние между всеми парами объектов в кластерах. Компромисс между одиночной и полной связью, менее чувствителен к выбросам внутри кластеров.

- **Метод Уорда (Ward)**  
  В классической формулировке метод Уорда выбирает для объединения те кластеры, при которых **минимально увеличивается внутрикластерная сумма квадратов**. Если обозначить центры кластеров как $\mu_A$, $\mu_B$, а их размеры как $n_A$, $n_B$, то критерий роста суммарной инерции при объединении можно записать как
  $$
  \Delta(A,B) = \frac{n_A n_B}{n_A + n_B}\,\lVert \mu_A - \mu_B \rVert_2^2.
  $$
  На каждом шаге объединяются кластеры с минимальным $\Delta(A,B)$. По смыслу это близко к идее k-means: минимизация разброса внутри кластеров.

![](images\linkage.png){width=500px}

Главная идея: иерархический анализ не только дает разбиение на кластеры, но и показывает **иерархию объединений**, что удобно для интерпретации типологий в социальных и маркетинговых исследованиях.


Попробуем провести иерархический кластерный анализ на данных по оптовым закупкам.


### Упражнение 6 — Матрица расстояний

Постройте евклидову матрицу расстояний для стандартизованных данных `wh_scaled` и сохраните ее в объект `dist_wh`. Функция для подсчета расстояний - `dist()`, `method="euclidean"`. Выведите матрицу расстояний на экран. 

```{r hex1, exercise = TRUE}
# ваш код
```

```{r hex1-solution, include=FALSE}
dist_wh <- dist(wh_scaled, method = "euclidean")
dist_wh
```

```{r hex1-check}
grade_code()
```


### Упражнение 7 — Агломеративная кластеризация (метод Уорда)

Выполните иерархическую агломеративную кластеризацию на `dist_wh` с помощью `hclust()` и методом `ward.D2`.  
Сохраните результат в объект `hc_wh`.

```{r hex2, exercise = TRUE}
# ваш код
```

```{r hex2-solution, include=FALSE}
hc_wh <- hclust(dist_wh, method = "ward.D2")
```

```{r hex2-check}
grade_this({
  if (!exists("hc_wh")) fail("Создайте объект hc_wh с результатами hclust().")
  if (!"hclust" %in% class(hc_wh)) fail("Проверьте, что вы использовали hclust().")
  pass("Иерархическая кластеризация (Ward.D2) выполнена.")
})
```


### Упражнение 8 — Дендрограмма

Постройте дендрограмму для объекта `hc_wh`.  
Используйте функцию `plot()`. Уберите подписи клиентов (`labels = FALSE`), чтобы график был более читаемым. Сделайте название с помощью параметра `main="  "`, добавьте подписи осей через `xlab` и `ylab'.


```{r hex3, exercise = TRUE}
# ваш код
```

```{r hex3-solution, include=FALSE}
plot(hc_wh, labels = FALSE,
     main = "Иерархическая кластеризация клиентов wholesale (Ward.D2)",
     xlab = "Клиенты", ylab = "Расстояние")
```

```{r hex3-check}
grade_code()
```

```{r hex3-q}
question_radio(
  "Ориентируясь на дендрограмму, сколько кластеров кажется разумным оставить?",
  answer("Примерно 3–5 кластеров (до крупного скачка высоты).", correct = TRUE),
  answer("Всего один кластер."),
  answer("Столько кластеров, сколько наблюдений."),
  allow_retry = TRUE
)
```


### Упражнение 9 — Обрезка дендрограммы и размеры кластеров

1. Обрежьте дендрограмму на 4 кластера (`k = 4`) с помощью `cutree()` и сохраните результат в `hc_clusters4`.  
2. Посчитайте, сколько клиентов попало в каждый кластер с помощью функции `table()`.

```{r hex4, exercise = TRUE}
# ваш код
```

```{r hex4-solution, include=FALSE}
hc_clusters4 <- cutree(hc_wh, k = 4)
table(hc_clusters4)
```

```{r hex4-check}
grade_this({
  if (!exists("hc_clusters4")) fail("Создайте объект hc_clusters4 с помощью cutree().")
  pass("Метки кластеров получены.")
})
```

```{r hex4-q}
question_numeric(
  "Сколько клиентов оказалось в самом многочисленном кластере (по результатам иерархической кластеризации)?",
  answer(as.character(max(table(hc_clusters4))), correct = TRUE),
  allow_retry = TRUE
)
```


### Упражнение 10 — Визуализация кластеров и интерпретация

1. Постройте визуализацию кластеров с помощью функции `fviz_cluster()`, используя список с данными и метками кластеров:  
   `fviz_cluster(list(data = wh_scaled, cluster = hc_clusters4))`.  
2. Сравните полученные кластеры с переменной `wholesale$Channel` (Horeca vs Retail), построив таблицу сопряженности `tab_hc_channel`.

```{r hex5, exercise = TRUE}
# ваш код
```

```{r hex5-solution, include=FALSE}
fviz_cluster(list(data = wh_scaled, cluster = hc_clusters4))

tab_hc_channel <- table(hc_clusters4, wholesale$Channel)
tab_hc_channel
```

```{r hex5-check}
grade_code()
```

```{r hex5-q}
question_text(
  "Кратко опишите, какие кластеры по иерархическому анализу ближе к Horeca, а какие — к розничным клиентам. Совпадает ли это с интерпретацией кластеров в k-means?",
  answer(" ", correct = TRUE)
)
```

## Упражнения: метод k-ближайших соседей (kNN)

Метод k-ближайших соседей — это простой метрический алгоритм **с учителем**: для нового объекта класс определяется по классам его ближайших соседей в обучающей выборке. Предполагается, что объекты с похожими признаками должны иметь одинаковую или близкую метку класса.

![](images\knn.jpg){width=400px}

### Основная идея

Имеется обучающая выборка $(x_i, y_i)$, $i = 1,\dots,n$, где:

- $x_i = (x_{i1},\dots,x_{ip})$ — вектор признаков объекта $i$;
- $y_i$ — известный класс (например, тип клиента: Horeca или Retail).

Пусть $x$ — новый объект, для которого нужно предсказать класс.

1. Выбираем число соседей $k$ и метрику расстояния $d(x, x_i)$ (обычно евклидово или манхеттенское расстояние).
2. Вычисляем расстояния от $x$ до всех обучающих объектов:
   $$
   d_i = d(x, x_i), \quad i = 1,\dots,n.
   $$
3. Находим индексы $k$ ближайших соседей:
   $$
   N_k(x) = \{i_1,\dots,i_k\}, \quad \text{таких что } d_{i_j} \text{ — } k \text{ наименьших расстояний.}
   $$
4. Для **классификации** выбираем класс, который чаще всего встречается среди соседей:
   $$
   \hat{y}(x) = \arg\max_{c} \sum_{i \in N_k(x)} \mathbf{1}(y_i = c),
   $$
   где $\mathbf{1}(\cdot)$ — индикатор (1, если условие верно, и 0 иначе).

Таким образом, kNN не строит явной модели; он опирается на локальное «окружение» объекта в пространстве признаков.


Мы будем классифицировать клиентов оптового дистрибьютора по типу `Channel` (Horeca vs Retail) на основе профиля расходов по категориям. 


### Упражнение 11 — Размеры обучающей и тестовой выборок

Посчитайте, сколько наблюдений попало в обучение и сколько — в тест:

```{r knn_ex1, exercise = TRUE}
idx_train <- sample(1:n, size = round(0.7 * n))
train_x <- wh_scaled[idx_train, ]
train_y <- channel[idx_train]
test_x  <- wh_scaled[-idx_train, ]
test_y  <- channel[-idx_train]
# посмотрите, сколько попало в обучающую выборку через length(idx_train)
# посчитайте объем тестовой выборки n - length(idx_train)
```

```{r knn_ex1-solution, include=FALSE}
length(idx_train)
n - length(idx_train)
```

```{r knn_ex1-check}
grade_code()
```

```{r knn_ex1-q}
question_numeric(
  "Сколько наблюдений в тестовой выборке?",
  answer(as.character(n - length(idx_train)), correct = TRUE),
  allow_retry = TRUE
)
```


### Упражнение 12 — kNN с k = 5

Выполните классификацию клиентов на тестовой выборке с помощью функции `knn()`:

- используйте настройки`train =train_x`, `train_y`, `test = test_x`, `cl = train_y`;
- задайте `k = 5`;
- сохраните предсказанные классы в объект `pred_knn5`.

```{r knn_ex2, exercise = TRUE}
```

```{r knn_ex2-solution, include=FALSE}
pred_knn5 <- knn(train = train_x, test = test_x, cl = train_y, k = 5)
```

```{r knn_ex2-check}
grade_this({
  if (!exists("pred_knn5")) fail("Создайте объект pred_knn5 с результатами knn().")
  pass("kNN с k = 5 выполнен.")
})
```


### Упражнение 13 — Матрица ошибок и точность

1. Постройте таблицу сопряженности между `pred_knn5` и `test_y`.  
2. Оцените долю верно классифицированных клиентов (accuracy):

````markdown
```r
acc_knn5 <- sum(diag(tab_knn5)) / sum(tab_knn5)
acc_knn5
```
````

```{r knn_ex3, exercise = TRUE}
```
```{r knn_ex3-solution, include=FALSE}
tab_knn5 <- table(pred_knn5, test_y)
tab_knn5
acc_knn5 <- sum(diag(tab_knn5)) / sum(tab_knn5)
acc_knn5
```

```{r knn_ex3-check}
grade_code()
```

```{r knn_ex3-q}
true_acc5 <- round(acc_knn5, 2)
question_numeric(
  "Какова точность (accuracy) kNN при k = 5? (округлите до двух знаков после запятой)",
  answer(as.character(true_acc5), correct = TRUE),
  allow_retry = TRUE
)
```


### Упражнение 14 — Влияние параметра k

Сравните точность классификации для нескольких значений k: 1, 3, 5, 9, 15.

```{r knn_ex4, exercise = TRUE}
ks <- c(1, 3, 5, 9, 15)
accs <- sapply(ks, function(k) {
  pred <- knn(train = train_x, test = test_x, cl = train_y, k = k)
  tab  <- table(pred, test_y)
  sum(diag(tab)) / sum(tab)
})
data.frame(k = ks, accuracy = accs)
```

```{r knn_ex4-solution, include=FALSE}
ks <- c(1, 3, 5, 9, 15)
accs <- sapply(ks, function(k) {
  pred <- knn(train = train_x, test = test_x, cl = train_y, k = k)
  tab  <- table(pred, test_y)
  sum(diag(tab)) / sum(tab)
})
data.frame(k = ks, accuracy = accs)
```

```{r knn_ex4-check}
grade_code()
```

```{r knn_ex4-q}
question_text(
  "Как меняется точность при увеличении k? Есть ли признаки переобучения при очень малых k (например, k = 1)?",
  answer(" ", correct = TRUE)
)
```

---

### Упражнение 15 — Интерпретация признаков и расстояния

Сделаем график, визуализирующий разделение между каналами клиентов по исходным и предсказанным значениям. 

```{r knn_ex5, exercise = TRUE}
# Сделаем PCA на стандартизированных данных
pca_wh <- prcomp(wh_scaled, scale. = FALSE) # так как данные уже стандартизированы, второй раз этого делать не нужно

# отберем координаты только тестовых наблюдений, потому что только по ним у нас есть предсказанные значения
scores <- as.data.frame(pca_wh$x[-idx_train, 1:2])
scores$true <- test_y          # истинные классы на тесте
scores$pred <- pred_knn5       # предсказания kNN

library(ggplot2)
ggplot(scores, aes(x = PC1, y = PC2, color = true, shape = pred)) +
  geom_point(alpha = 0.8) +
  labs(
    x = "PC1", y = "PC2",
    color = "Channel (истинный)",
    shape  = "Channel (kNN)"
  ) +
  theme_minimal()
```

```{r knn_ex5-solution, include=FALSE}
# Сделаем PCA на стандартизированных данных
pca_wh <- prcomp(wh_scaled, scale. = FALSE) # так как данные уже стандартизированы, второй раз этого делать не нужно

# отберем координаты только тестовых наблюдений, потому что только по ним у нас есть предсказанные значения
scores <- as.data.frame(pca_wh$x[-idx_train, 1:2])
scores$true <- test_y          # истинные классы на тесте
scores$pred <- pred_knn5       # предсказания kNN

library(ggplot2)
ggplot(scores, aes(x = PC1, y = PC2, color = true, shape = pred)) +
  geom_point(alpha = 0.8) +
  labs(
    x = "PC1", y = "PC2",
    color = "Channel (истинный)",
    shape  = "Channel (kNN)"
  ) +
  theme_minimal()
```

```{r knn_ex5-check}
grade_code()
```

```{r knn_ex5-q}
question_text(
  "В каких случаях просходит непопадание в профиль клиента? Ваши предположения на этот счет",
  answer(" ", correct = TRUE)
)
```
```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
