---
title: "Снижение размерности: анализ соответствий"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(FactoMineR)
library(factoextra)
knitr::opts_chunk$set(echo=FALSE)
data(children)
children<-children
children2<-children[1:14,]
res.ca <- CA(children2, graph = FALSE)
eig.ca <- get_eigenvalue(res.ca)
tutorial_options(
  exercise.reveal_solution = FALSE  # НЕ показывать решения студентам
)
data(hobbies)
hobbies<-hobbies
hobbies.active <- hobbies[, 1:15]
res.mca <- MCA(hobbies.active, graph = FALSE)
res.mca.sup <- MCA(hobbies, quali.sup = 19:22, quanti.sup = 23, graph = FALSE)
hob_famd <- hobbies |> 
  transmute(
    nb.activitees,
    sex = factor(Sex),
    age = factor(Age),
    marital.status = factor(`Marital status`),
    TV = factor(TV)
  )
res.famd <- FAMD(hob_famd, graph = FALSE)
```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Предварительные замечания

Анализ соответствий (CA) и множественный анализ соответствий (MCA) относятся к методам снижения размерности для **категориальных** данных и позволяют визуализировать связь между категориями переменных в пространстве низкой размерности.
CA применяется к двумерным перекрестным таблицам (часто – частотным), тогда как MCA расширяет подход на несколько номинальных/порядковых переменных.

Приведем несколько кейсов, где анализ соответствий может быть использован для решения прикладных задач.

### 1. Маркетинг: образ  и карта воспрития бренда (perceptual map)

- В такого рода исследованиях исходными данными являются результаты опроса  потребителей по шкалам, таким как «Какие из перечисленных брендов вы ассоциируете с атрибутами X, Y, Z» (надёжный, модный, семейный, недорогой и т.п.). 
- Затем строится перекрестная таблица «бренды × имиджевые атрибуты» (доли упоминаний) и проводится анализ соответствий; на карте расположены и бренды, и атрибуты. [Пример исследования](https://journals.sagepub.com/doi/10.1177/1470785318801480)
- Как интерпретировать? Близость бренда к атрибуту показывает, что он ассоциируется с этим образом; близкие бренды конкурируют за одну позицию, удаленные занимают отличающиеся ниши. 
  - Так, анализ соответствий был проведен для рынка шампуней и профилей брендов, результаты представлены следующим образом (пример взят из следующего [источника](https://www.ashokcharan.com/Marketing-Analytics/images/image020.png):

![](images\ca.png){width=500}
  
### 2. Поведение и критерии покупательского выбора

- Еще один пример из области маркетинговых исследований: исходными данными выступают результаты анкетирования  с номинальными/порядковыми ответами о критериях покупки (цена, вкус, бренд, страна происхождения), типах магазинов, частоте потребления и социально‑демографических признаках. [Источник для дополнительного чтения](https://journals.economic-research.pl/ceref/article/download/4054/2567/14123).
- Модель: проводится множественный анализ соответствий по совокупности категориальных переменных. 
- Интерпретация: карта показывает кластеры потребителей (группы уровней переменных), «притянутых» к разным атрибутам. Так, в одной из работ о критериях покупки пива MCA позволил выделить чёткие сегменты и связать атрибуты продукта с предпочтениями групп покупателей. [Пример статьи](https://www.oeconomia.pl/ceref/article/download/4054/2567).

### 3. E‑commerce: критерии онлайн‑покупок по категориям товаров

- Исходные данные: опрос пользователей, где для различных товарных категорий (электроника, одежда, продукты, книги) фиксируются значимые критерии онлайн‑покупки (доставка, цена, бренд, отзывы, безопасность платежа и др.). [Пример исследования](https://www.academic-pub.org/ojs/index.php/ijecs/article/download/1818/374).
- Модель: строится таблица «категория товара × критерий выбора», проводится CA или MCA, иногда с добавлением демографических переменных как дополнительных. 
- Анализ позволяет сегментировать пользователей, их стратегии коммуникации и опыт использования маркетплейса - UX (фильтры, подсказки) под разные товарные группы. 

### 4. Исследование политических предпочтений: карта позиций партий и установок

- Исходные данные: опрос электората, где  измеряются установки по политическим вопросам, самоидентификация, поддержка партий [Хороший пример](https://journals.plos.org/plosone/article/file?id=10.1371%2Fjournal.pone.0287180&type=printable).
- Модель: MCA по множеству вопросов (позиции по миграции, налогам, морали, гендерной политике и т.п.), партии и группы избирателей часто добавляются как категории/дополнительные переменные. 
- Интерпретация: двухмерная карта отражает латентные измерения идеологии (например, «социальный либерализм – консерватизм» и «экономический левый – правый») и показывает, какие ответы и какие партии группируются вблизи друг друга. - Это помогает выявлять идеологические кластеры внутри партий и различия между «ядерными» и «плавающими» избирателями, а также использовать карту как визуальную опору для коммуникационных стратегий. 

![](images\mca.png)

В рамках данного тьюториала мы разберем, как проводить простой и множественный анализ соответствий в R. 

## Тестовые вопросы

### Вопрос 1

Что является входными данными для классического анализа соответствий (CA)?

```{r q1}
question(
  "",
  answer("Матрица ковариаций количественных переменных."),
  answer("Двумерная контингентная таблица (частоты или профили).", correct = TRUE),
  answer("Дизъюнктивная матрица индикаторов 0/1 для множества категориальных переменных."),
  answer("Любая матрица расстояний между объектами."),
  incorrect = "CA выполняется на перекрестных таблицах для двух категориальных переменных.",
  allow_retry = TRUE
)
```

### Вопрос 2

Что представляют собой координаты строк и столбцов в пространстве главных осей при CA?

```{r q2}
question(
  "",
  answer("Оценки регрессионной модели."),
  answer("Координаты профилей строк и столбцов на факторных осях, полученных из сингулярного разложения.", correct = TRUE),
  answer("Случайные точки без статистического смысла."),
  answer("Координаты индивидуумов в PCA."),
  incorrect = "CA раскладывает профили строк и столбцов и отображает их на осях максимальной инерции.",
  allow_retry = TRUE
)
```

### Вопрос 3

Как интерпретируется расстояние между двумя категориями на карте CA?

```{r q3}
question(
  "",
  answer("Если точки близко — категории связаны похожими профилями/частотами.", correct = TRUE),
  answer("Если точки близко — категории статистически независимы."),
  answer("Расстояние не имеет интерпретации."),
  answer("Расстояние отражает линейную корреляцию Пирсона."),
  incorrect = "Близость на карте CA означает сходство профилей категорий, а не отсутствие связи.",
  allow_retry = TRUE
)
```

### Вопрос 4

Какой критерий в CA и MCA чаще всего рассматривают аналогом дисперсии в PCA?

```{r q4}
question(
  "",
  answer("Коэффициент детерминации R^2."),
  answer("Инерция (inertia), связанная с собственными значениями.", correct = TRUE),
  answer("Среднеквадратическая ошибка."),
  answer("Статистика хи-квадрат."),
  incorrect = "Инерция – мера вариации, объясняемой осями, и связана с собственными значениями в CA/MCA.",
  allow_retry = TRUE
)
```

### Вопрос 5

Какой тип матрицы обычно анализируется в MCA (при методе по умолчанию)?

```{r q5}
question(
  "",
  answer("Матрица корреляций Пирсона."),
  answer("Burt-матрица (все пары перекрестных таблиц)."),
  answer("Индикаторная (дизъюнктивная) матрица 0/1 по категориям всех переменных.", correct = TRUE),
  answer("Матрица расстояний Гаусса."),
  incorrect = "По умолчанию MCA в FactoMineR выполняется на индикаторной матрице (method = 'Indicator').",
  allow_retry = TRUE
)
```

### Вопрос 6

Что означает высокая сумма cos2 для категории на первых двух осях MCA?

```{r q6}
question(
  "",
  answer("Категория хорошо представлена на этих осях и лежит близко к окружности корреляции.", correct = TRUE),
  answer("Категория не информативна и должна быть исключена."),
  answer("Категория имеет нулевой вклад в инерцию."),
  answer("Это признак мультиколлинеарности."),
  incorrect = "Большое cos2 означает, что координаты категории хорошо объясняются выбранными осями.",
  allow_retry = TRUE
)
```

### Вопрос 7

Для чего используют функции `fviz_ca_row()` и `fviz_ca_col()` из пакета factoextra?

```{r q7}
question(
  "",
  answer("Для вычисления CA."),
  answer("Для визуализации строк (row) и столбцов (col) таблицы в пространстве факторов.", correct = TRUE),
  answer("Для выполнения кластерного анализа."),
  answer("Для стандартизации данных."),
  incorrect = "CA считают с помощью `CA()` из FactoMineR, а в factoextra — визуализируют результаты.",
  allow_retry = TRUE
)
```

### Вопрос 8

Какой объект создает функция `CA()` из пакета FactoMineR?

```{r q8}
question(
  "",
  answer("Объект класса 'lm'."),
  answer("Объект, содержащий eigenvalues, координаты строк и столбцов, cos2, вклады и др.", correct = TRUE),
  answer("Только матрицу собственных векторов."),
  answer("Графический объект ggplot."),
  incorrect = "CA() возвращает богатый списковый объект с результатами для строк и столбцов.",
  allow_retry = TRUE
)
```

### Вопрос 9

В чем основное отличие MCA от CA?

```{r q9}
question(
  "",
  answer("MCA применим только к количественным переменным."),
  answer("MCA – это просто PCA, переименованный для категориальных данных."),
  answer("MCA обобщает CA для случая, когда переменных категориальных три и более.", correct = TRUE),
  answer("Отличий нет, это синонимы."),
  incorrect = "CA работает с одной перекрестной таблицей двух переменных, MCA – с несколькими категориальными переменными.",
  allow_retry = TRUE
)
```

### Вопрос 10

Как можно добавить количественные переменные в MCA в FactoMineR?

```{r q10}
question(
  "",
  answer("Как активные переменные, предварительно дискретизировав их."),
  answer("Как дополнительные количественные переменные (quanti.sup), не влияющие на построение осей.", correct = TRUE),
  answer("Никак — MCA не допускает количество переменных."),
  answer("Только через предварительный PCA."),
  incorrect = "Quantitative supplementary variables (quanti.sup) могут добавляться для интерпретации осей, не влияя на их построение.",
  allow_retry = TRUE
)
```

### Вопрос 11

Какой тип анализа использовать, если в данных смешаны количественные и категориальные переменные и нужно рассматривать их совместно?

```{r q11}
question(
  "",
  answer("Только PCA."),
  answer("FAMD (factor analysis of mixed data), реализованный в FactoMineR.", correct = TRUE),
  answer("Только MCA."),
  answer("Иерархический кластерный анализ."),
  incorrect = "FAMD специально разработан для совместного анализа количественных и качественных переменных.",
  allow_retry = TRUE
)
```

### Вопрос 12

Какой функцией из пакета factoextra удобно визуализировать результаты MCA (индивиды и категории вместе) из FactoMineR?

```{r q12}
question(
  "",
  answer("fviz_pca_biplot()"),
  answer("fviz_mca_biplot()", correct = TRUE),
  answer("fviz_famd_ind()"),
  answer("fviz_cluster()"),
  incorrect = "Функции семейства fviz_mca_* предназначены для визуализации MCA.",
  allow_retry = TRUE
)
```

### Вопрос 13

Какой аргумент функции `MCA()` позволяет указать дополнительных (неактивных) индивидов?

```{r q13}
question(
  "",
  answer("ind.sup", correct = TRUE),
  answer("quali.sup"),
  answer("row.sup"),
  answer("sup.ind"),
  incorrect = "Аргумент ind.sup — индексы наблюдений, которые добавляются как дополнительные.",
  allow_retry = TRUE
)
```

### Вопрос 14

Как можно получить собственные значения и долю объясненной инерции для объекта MCA/CA из FactoMineR с помощью factoextra?

```{r q14}
question(
  "",
  answer("summary(res.mca)$coef"),
  answer("get_eigenvalue(res.mca)", correct = TRUE),
  answer("res.mca$eig.only"),
  answer("fviz_mca_ind(res.mca)"),
  incorrect = "Функция get_eigenvalue() извлекает eigenvalues и проценты объясненной инерции.",
  allow_retry = TRUE
)
```

### Вопрос 15

На карте MCA категории одной и той же переменной образуют "звезду". Что означает расположение центра этой звезды близко к началу координат?

```{r q15}
question(
  "",
  answer("Переменная в целом мало связана с основными осями и слабо участвует в общей структуре.", correct = TRUE),
  answer("Переменная имеет очень высокую дисперсию."),
  answer("Переменная идеально предсказывает все другие переменные."),
  answer("Это признак ошибки в данных."),
  incorrect = "Если барицентр категорий близко к началу координат, переменная мало участвует в объяснении общей инерции на первых осях.",
  allow_retry = TRUE
)
```


## Упражнения: Анализ соответствий (CA)

Разберем возможности анализа соответствий используем данные `children` из пакета `FactoMineR`.

- Данные взяты из [книги](https://cran.r-project.org/web/packages/FactoMineR/FactoMineR.pdf) *Traitements Statistiques des Enquêtes* (ред. D. Grange, L. Lebart, Dunod, 1993).
- В оригинальном исследовании респондентам задавали вопрос:  
  *«По вашему мнению, какие причины могут заставить женщину или пару колебаться в решении иметь детей?»*. Индивидуальные ответы были агрегированы в таблицу сопряженности (contingency table), включенную в пакет FactoMineR под названием `children`. 

- Набор данных представляет собой дата‑фрейм размером **18 строк × 8 столбцов**. 
- **Строки**: 18 возможных **причин** колебаний относительно рождения детей (например, финансовые трудности, жилищные условия, карьерные соображения и т.п. — конкретные формулировки в документации не перечислены, но трактуются именно как «reasons mentioned»). 
- **Столбцы**: 8 категорий респондентов, заданных комбинациями **образования и возраста**; в описании они указаны как «different categories (education, age) people belong to». 
- **Ячейка** $(i,j)$: количество упоминаний причины $i$ в группе $j$. 

**Столбцы (уровень образования и возраст)**:

- *unqualified* – без квалификации, без диплома  
- *cep* – диплом уровня CEP (аналог базового свидетельства об образовании)  
- *bepc* – диплом уровня BEPC (неполное среднее образование)  
- *high_school_diploma* – аттестат о среднем образовании  
- *university* – высшее образование (университетский диплом)  
- *thirty* – возраст около 30 лет (30 лет / 30–39)  
- *fifty* – возраст около 50 лет (50 лет / 40–49)  
- *more_fifty* – старше 50 лет  

**Строки (причины нежелания иметь детей)**:

- *money* – деньги, финансовые средства  
- *future* – будущее (неуверенность в будущем)  
- unemployment – безработица  
- *circumstances* – обстоятельства (жизненные обстоятельства)  
- *hard* – трудности, тяжело  
- *economic* – экономические условия  
- *egoism* – эгоизм (желание жить для себя)  
- *employment* – занятость, работа (наличие/поиск работы)  
- *finances* – финансовое положение  
- *war* – война  
- *housing* – жильё, жилищные условия  
- *fear* – страх  
- *health* – проблемы со здоровьем  
- *work* – работа (трудовая нагрузка, карьера)  
- *comfort* – комфорт, удобства  
- *disagreement* – разногласия (в семье/с партнёром)  
- *world* – мир (состояние мира, общества, «мир устроен плохо»)  
- *to_live* – жить (качество жизни / «как жить дальше»)

### Упражнение 1 — Смотрим на данные

Выведите на экран данные набора `children'. 

```{r ex1, exercise = TRUE}
children
```


```{r ex1-solution, include=FALSE}
children
```

```{r ex1-check}
grade_code()
```

Для анализа соответствий нам нужна таблица без пропусков. Создайте таблицу `children2`, в которой будут все столбцы и строки с 1 по 14.

```{r ex1a, exercise = TRUE}

```


```{r ex1a-solution, include=FALSE}
children2<-children[1:14,]
```

```{r ex1a-check}
grade_code()
```


### Упражнение 2 — Запуск CA

- Выполните анализ соответствий на  данных `children2` с помощью `CA()` из пакета `FactoMineR` (без графиков).  

- Сохраните результат в объект `res.ca`.

```{r ex2, exercise = TRUE}
# ваш код
```


```{r ex2-solution, include=FALSE}
res.ca <- CA(children2, graph = FALSE)
```

```{r ex2-check}
grade_this({
  if (!exists("res.ca")) fail("Создайте объект res.ca с результатами CA.")
  if (!"CA" %in% class(res.ca)) fail("Похоже, res.ca получен не функцией CA() из FactoMineR.")
  pass("Анализ соответствий успешно выполнен.")
})
```

### Упражнение 3 — Собственные значения и инерция

Поскольку анализ соответствий -- это метод снижения размерности, в результате взамен исходных переменных (строк и столбцов) в нашей таблице мы должны получить новые измерения.

Основная идея напоминает метод главных компонент, однако вместо обычной дисперсии здесь используется понятие **инерции**, связанной с $\chi^2$‑расстояниями. 


#### Как считаются собственные значения в CA

1. **Исходная таблица нормируется до относительных частот**  
   В таблице с частотами (это наша таблица `children`) $N = (n_{ij})$ значение в каждой ячейке делится на сумму всех частот по таблице $n = \sum_{i,j} n_{ij}$.  
   Получается матрица относительных частот:
   $$
   P = (p_{ij}), \quad p_{ij} = \frac{n_{ij}}{n},
   $$
   которую часто называют *correspondence matrix*. 

2. **Подсчитываются маргинальные распределения (массы строк и столбцов)**  
   $$
   r_i = \sum_j p_{ij} \quad (\text{масса строки}), \qquad
   c_j = \sum_i p_{ij} \quad (\text{масса столбца}).
   $$ 

3. **Подсчитываются профили строк и столбцов**  
   - Профиль строки $i$: условные вероятности $p_{ij} / r_i$ (распределение по столбцам при фиксированной строке).  
   - Профиль столбца $j$: $p_{ij} / c_j$ (распределение по строкам). 
   
4. **Строится матрица стандартизированных остатков**


   $$
   S = D_r^{-1/2} (P - r c^\top) D_c^{-1/2},
   $$
   где $D_r, D_c$ — диагональные матрицы масс, в которых на диагонали находятся $\frac{1}{\sqrt{r_i}}$ или $\frac{1}{\sqrt{c_i}}$. 
   
5. **Сингулярное разложение матрицы S**

На $S$ делается сингулярное разложение (SVD), откуда и берутся сингулярные значения $\sigma_k$, их квадраты $\lambda_k = \sigma_k^2$ — это **собственные значения**, а их сумма — общая инерция. 

   $$
   S = D_r^{-1/2}(P - r c^\top) D_c^{-1/2} = U \Sigma V^\top,
   $$
   
 Каждая ось (Dim k) соответствует собственному значению $\lambda_k$, которое показывает, сколько «информации» (вариации профилей в $\chi^2$-метрике) она объясняет.

#### Что такое инерция и как её интерпретировать

В CA термин **инерция** — это аналог дисперсии в PCA, но измеренный в $\chi^2$-метрике. 

- **Общая инерция** двухмерной таблицы:
  $$
  I = \frac{\chi^2_{\text{таблицы}}}{n},
  $$
  то есть общее значение критерия Пирсона $\chi^2$, поделенное на сумму всех частот (объем выборки). 
  
  
- Эта же величина равна сумме всех собственных значений:
  $$
  I = \sum_k \lambda_k.
  $$
  
**Процент инерции, объясненный осью $k$**:
$$
\text{Perc}_k = \frac{\lambda_k}{\sum_j \lambda_j} \times 100\%.
$$


Интерпретация:

- Инерция измеряет, насколько сильно **профили строк и столбцов** отклоняются от ситуации статистической независимости. 

- Если первая ось объясняет, например, 48% инерции, а вторая ещё 40%, то в плоскости (Dim1, Dim2) уже «видно» почти 88% всей структуры ассоциаций между строками и столбцами. 
- Чем **больше доля инерции** на первых осях, тем более информативна двумерная карта: она хорошо восстанавливает структуру зависимостей в исходной таблице (почти весь $\chi^2$). 


**Задание**:

- Извлеките собственные значения и долю объясненной инерции с помощью `get_eigenvalue(res.ca)` и сохраните результат в объект `eig.ca`.
- Выведите на экран `eig.ca`.

```{r ex3, exercise = TRUE}
# ваш код
```

```{r ex3-solution, include=FALSE}
eig.ca <- get_eigenvalue(res.ca)
eig.ca
```

```{r ex3-check}
grade_code()
```
```{r ex3-q}
true_value<-round(eig.ca[2,3], 1)
question_numeric(
  "Какую долю инерции объясняют первые две оси? (округлите до одной десятой)",
   answer(as.character(true_value), correct = T), 
  allow_retry = TRUE,
  try_again   = "Посмотрите второй элемент столбца cumulative.variance.percent в eig.ca."
)
```

### Упражнение 4 — Визуализация строк и столбцов

Постройте карту строк и столбцов для CA с помощью функций `fviz_ca_row(res.ca)` и `fviz_ca_col(res.ca)`.

```{r ex4, exercise = TRUE}
# ваш код
```


```{r ex4-solution, include=FALSE}
fviz_ca_row(res.ca)
fviz_ca_col(res.ca)
```

```{r ex4-check}
grade_code()
```

```{r ex4-q}
question_radio(
  "Как интерпретировать близко расположенные строки на графике строк?",
  answer("Они имеют похожие профили распределения по столбцам.", correct = TRUE),
  answer("Они независимы друг от друга."),
  answer("Это артефакт построения графика."),
  allow_retry = TRUE,
  try_again = "Помните, CA отражает сходство профилей строк."
)
```


### Упражнение 5 — Вклады категорий

Извлеките результаты для столбцов (категорий переменной столбцов) с помощью `res.ca$col`.  
Сравните вклады категорий (`$contrib`) на первую ось и найдите категорию с наибольшим вкладом.

```{r ex5, exercise = TRUE}
# ваш код
```

```{r ex5-hint, eval=FALSE}
# contrib_col <- res.ca$col$contrib
# sort(contrib_col[,1], decreasing = TRUE)
```

```{r ex5-solution, include=FALSE}
contrib_col <- res.ca$col$contrib
sort(contrib_col[, 1], decreasing = TRUE)
```

```{r ex5-check}
grade_code()
```

```{r ex5-q}
question(
  "Какая категория вносит наибольший вклад в общую вариацию профилей столбцов?",
  answer("more_fifty", correct = TRUE),
  answer("thirty", correct = TRUE),
  answer("high_school_diploma"),
  answer("bepc"),
  answer("unqualified"),
  answer("university"),
  answer("cep"),
  answer("fifty"),
  incorrect = "Обратите внимание на самые большие значения инерции: они соответствуют переменным, которые сильнее всего отклоняются от модели независимости.",
  allow_retry = TRUE
)
```


### Упражнение 6 — Интерпретация cos2

Рассчитайте cos2 для строк: `row_cos2 <- res.ca$row$cos2`.  
Выберите строки с наибольшими значениями cos2 по первой оси `head(sort(row_cos2[, 1], decreasing = TRUE))`.

```{r ex6, exercise = TRUE}
# ваш код
```

```{r ex6-solution, include=FALSE}
row_cos2 <- res.ca$row$cos2
head(sort(row_cos2[, 1], decreasing = TRUE))
```

```{r ex6-check}
grade_code()
```

```{r ex6-q}
question_radio(
  "Что означает высокий cos2 для строки на первой оси?",
  answer("Строка хорошо представлена на этой оси, её положение надежно интерпретировать.", correct = TRUE),
  answer("Строка не влияет на анализ."),
  answer("Строка обязательно является выбросом."),
  allow_retry = TRUE
)
```


### Упражнение 7 — Объединённая карта (biplot)

Постройте объединенный график строк и столбцов при помощи `fviz_ca_biplot(res.ca)` и интерпретируйте близость некоторых категорий строк и столбцов.

```{r ex7, exercise = TRUE}
# ваш код
```

```{r ex7-hint, eval=FALSE}
# fviz_ca_biplot(res.ca, repel = TRUE)
```

```{r ex7-solution, include=FALSE}
fviz_ca_biplot(res.ca, repel = TRUE)
```

```{r ex7-check}
grade_code()
```

## Упражнения: Множественный анализ соответствий

Используем набор `hobbies`для проведения множественного анализа соответствий. Массив получен из реального опроса 8403 человек о формах досуга (данные INSEE/CSA).

Переменные:

Первые 18 столбцов — активные категориальные переменные
Это бинарные (0/1) вопросы о занятиях:

- чтение (Reading)
- прослушивание музыки (Listening music)
- кино (Cinema)
- театр/шоу (Show)
- выставки (Exhibition)
- компьютер (Computer)
- спорт (Sport)
- прогулки (Walking)
- путешествия (Travelling)
- игра на музыкальных инструментах (Playing music)
- садоводство (Gardening)
- вязание (Knitting)
- кулинария (Cooking)
- рыбалка (Fishing)
- телевизор (TV).

Следующие 4 столбца — дополнительные (supplementary) качественные переменные:

- sex — пол (male, female)
- age — возрастные группы (15–25, 26–35, …, 86–100)
- marital.status — семейное положение (single, married, widowed, divorced, remarried)
- profession — профессия/социальный статус (manual labourer, unskilled worker, technician, foreman, senior management, employee, other).

23‑й столбец — дополнительная количественная переменная

- nb.activitees — количество хобби, которыми человек занимается (из 18 возможных).
 
### Упражнение 8 — Подготовка данных для MCA

Создайте подмножество данных `hobbies.active`, включив только активные категориальные переменные (без дополнительных).  
Например, возьмите первые 15 столбцов `hobbies` как активные.

```{r ex8, exercise = TRUE}
hobbies.active <-
```


```{r ex8-solution, include=FALSE}
hobbies.active <- hobbies[, 1:15]
```

```{r ex8-check}
grade_code()
```


### Упражнение 9 — Запуск MCA

Выполните MCA на `hobbies.active` с помощью `MCA()`, без графиков.  
Сохраните результат в объект `res.mca`.

```{r ex9, exercise = TRUE}
# ваш код
```

```{r ex9-solution, include=FALSE}
res.mca <- MCA(hobbies.active, graph = FALSE)
```

```{r ex9-check}
grade_this({
  if (!exists("res.mca")) fail("Создайте объект res.mca.")
  if (!"MCA" %in% class(res.mca)) fail("Проверьте, что вы использовали функцию MCA() из FactoMineR.")
  pass("MCA успешно выполнен.")
})
```


### Упражнение 10 — Собственные значения и scree plot MCA

Извлеките собственные значения с помощью `get_eigenvalue(res.mca)` и сохраните их в `eig.mca`.  
Постройте график каменистой осыпи с помощью `fviz_eig(res.mca)`.

```{r ex10, exercise = TRUE}
# ваш код
```

```{r ex10-hint, eval=FALSE}
# eig.mca <- get_eigenvalue(res.mca)
# fviz_eig(res.mca)
```

```{r ex10-solution, include=FALSE}
eig.mca <- get_eigenvalue(res.mca)
fviz_eig(res.mca)
```

```{r ex10-check}
grade_code()
```

```{r ex10-q}
question_radio(
  "Сколько осей вы бы оставили для интерпретации на основе графика и доли инерции?",
  answer("Те оси, которые до накопленной инерции примерно 60–70%.", correct = TRUE),
  answer("Только одну ось."),
  answer("Все оси без исключения."),
  allow_retry = TRUE
)
```


### Упражнение 11 — Карта категорий MCA

Постройте:

- график категорий: `fviz_mca_var(res.mca, repel = TRUE)`

```{r ex11, exercise = TRUE}
# ваш код
```


```{r ex11-solution, include=FALSE}
fviz_mca_var(res.mca, repel = TRUE)
```

```{r ex11-check}
grade_code()
```


### Упражнение 12 — Вклады категорий в оси MCA

Извлеките результаты для категорий: `var_mca <- res.mca$var`.  
Посмотрите на вклад категорий (`var_mca$contrib`) в первую ось и выведите 10 категорий с наибольшим вкладом.

```{r ex12, exercise = TRUE}
# ваш код
```

```{r ex12-hint, eval=FALSE}
# var_mca <- res.mca$var
# head(sort(var_mca$contrib[,1], decreasing = TRUE), 10)
```

```{r ex12-solution, include=FALSE}
var_mca <- res.mca$var
head(sort(var_mca$contrib[, 1], decreasing = TRUE), 10)
```

```{r ex12-check}
grade_code()
```


```{r ex12-q}
question_text("Напишите краткий комментарий по поводу того, что Вы видите",
answer(" ", correct = TRUE))
```


### Упражнение 13 — MCA с дополнительными переменными

Добавьте в MCA из набора `hobbies` несколько количественных переменных как дополнительные (quanti.sup).

Создайте объект `res.mca.sup`, указав, какие столбцы являются дополнительными.

```{r ex13, exercise = TRUE}
# ваш код
```

```{r ex13-hint, eval=FALSE}
# Пример:
# res.mca.sup <- MCA(hobbies, quali.sup = 19:22, quanti.sup = 23, graph = FALSE)
```

```{r ex13-solution, include=FALSE}
res.mca.sup <- MCA(hobbies, quali.sup = 19:22, quanti.sup = 23, graph = FALSE)
```

```{r ex13-check}
grade_this({
  if (!exists("res.mca.sup")) fail("Создайте объект res.mca.sup.")
  pass("Дополнительные переменные добавлены корректно.")
})
```


### Упражнение 14 — График с новыми переменными

Сделайте график по переменным по результатам дополнительного множественного анализа соответствий `res.mca.sup`.

```{r ex14, exercise = TRUE}
# ваш код
```

```{r ex14-hint, eval=FALSE}
fviz_mca_var(res.mca.sup)
```

```{r ex14-solution, include=FALSE}
fviz_mca_var(res.mca.sup)
```

```{r ex14-check}
grade_code()
```

```{r ex14-q}
question_text("Какие хобби имеют респонденты, указавшие в качестве профессии - управление (management)",
answer(" ", correct = TRUE))
```


### Упражнение 15 — Смешанный анализ FAMD

Используем данные `hobbies`, где есть и количественная, и категориальные переменные.

**Задание:** подготовьте подмножество данных и выполните FAMD:

- Создайте новый датафрейм `hob_famd`, включив:
  - количественную переменную `nb.activitees`;  
  - категориальные переменные: `sex`, `age`, `marital.status`, `TV` (частота просмотра TV как фактор);  
- Выполните FAMD: `res.famd <- FAMD(hob_famd, graph = FALSE)`.


```{r ex15, exercise = TRUE}
# ваш код
```

```{r ex15-hint, eval=FALSE}
# hob_famd <- hobbies %>%
#   transmute(
#     nb.activitees,
#     sex = factor(Sex),
#     age = factor(Age),
#     marital.status = factor(`Marital status`),
#     TV = factor(TV)  # частота просмотра TV как категориальная
#   )
# res.famd <- FAMD(hob_famd, graph = FALSE)
```

```{r ex15-solution, include=FALSE}
hob_famd <- hobbies |> 
  transmute(
    nb.activitees,
    sex = factor(Sex),
    age = factor(Age),
    marital.status = factor(`Marital status`),
    TV = factor(TV)
  )
res.famd <- FAMD(hob_famd, graph = FALSE)
```

```{r ex15-check}
grade_this({
  if (!exists("res.famd")) fail("Создайте объект res.famd с результатами FAMD.")
  pass("FAMD на реальных данных выполнен корректно.")
})
```

### Упражнение 16 — Интерпретация FAMD

Теперь визуализируем результаты FAMD и интерпретируем их.

**Задание:**

1. Постройте:
   - карту переменных: `fviz_famd_var(res.famd, repel = TRUE)`.  
2. Обратитесь к карте переменных и ответьте на вопрос:  
   - какие переменные/категории ближе всего к первой оси;  
   - какие — ко второй;  
   - как в этом пространстве ведет себя `nb.activitees` (число хобби).

```{r ex16, exercise = TRUE}
# ваш код
```

```{r ex16-hint, eval=FALSE}
fviz_famd_var(res.famd, repel = TRUE)
```

```{r ex16-solution, include=FALSE}
fviz_famd_var(res.famd, repel = TRUE)
```

```{r ex16-check}
grade_code()
```

```{r ex16-q}
question_radio(
  "Что даёт FAMD по сравнению с раздельным PCA и MCA на этих данных?",
  answer("Позволяет одновременно учитывать nb.activitees (число хобби) и категориальные переменные (пол, возраст, семейное положение, TV) в едином факторном пространстве.", correct = TRUE),
  answer("Никаких преимуществ, это просто другое название MCA."),
  answer("Применим только для временных рядов."),
  allow_retry = TRUE
)
```


```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
