---
title: "Снижение размерности: метод главных компонент и факторный анализ"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(caret)
library(dplyr)
library(olsrr)
library(psych)
data(bfi)

library(FactoMineR)
library(factoextra)
df_pca <- bfi[, 1:25] |>
  na.omit()

res.pca <- PCA(df_pca, scale.unit = TRUE, ncp = 10, graph = FALSE)

df_fa <- bfi |>
  select(paste0("E", 1:5), paste0("N", 1:5)) |> 
           na.omit()
res.fa <- fa(df_fa, nfactors = 2, rotate = "varimax", fm="ml")
knitr::opts_chunk$set(echo=FALSE, warning=FALSE, message=FALSE)
```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## Предварительные замечания

**Методы снижения размерности** -- это совокупность различных по своей природе методов, направленных на сокращение пространства признаков, представленных в наборе данных. 

В результат снижения размерности формируются обобщенные переменные, которые в дальнейшем могут быть использованы например, в качестве интегральных показателей при конструировании различных рейтингов, в качестве предикторов в регрессионном анализе или признаков, используемых для кластеризации объектов и других аналитических задач.

![](https://upload.wikimedia.org/wikipedia/commons/thumb/9/9c/PCA_Projection_Illustration.gif/1280px-PCA_Projection_Illustration.gif){width=400px}

Эти методы:

- помогают выявить скрытые структуры и обнаруживать латентные факторы /свойства, стоящие за ответами респондентов
- упрощают интерпретацию: вместо десятков разрозненных вопросов исследователь работает с несколькими обобщенными показателями (факторами или компонентами)
- готовят данные для дальнейшего анализа: уменьшение количества переменных помогает избежать "проклятия размерности" (это фундаментальная проблема анализа данных: при росте числа переменных (признаков) объем пространства растет экспоненциально, а данные становятся разреженными и "пустыми") и повышает надежность статистических выводов.


> Допустим, для проведения комплексной оценки социального благополучия студентов мы отобрали 150 показателей, описывающих все грани разнообразной и насыщенной студенческой жизни - их материальное положение, академические успехи, профессиональное развитие, участие в творческой деятельности, социальную активность и пр. Анализировать все 150 показателей сразу невозможно, особенно, если объем выборки невелик. Необходимо их укрупнение до нескольких базовых «осей», по которым будет удобно сравнивать как самих студентов, так и образовательные учреждения. Это может стать хорошим инструментом для развития образовательной среды в отдельном университете или регионе. 

Снижение размерности как раз и помогает перейти от данных, представленных в высокоразмерном пространстве (150 показателей) в пространство меньшей размерности (например, 10 основных "доменов" или "сфер"). 


Как это сделать?

Выбор метода снижения размерности в первую очередь зависит от типа данных. В социальных исследованиях до 70-90% переменных имеют нечисловую природу (номинальные или порядковые шкалы) . Это ключевой момент, разделяющий все методы на две большие группы.

**1. Для количественных (непрерывных) данных**:
<style>
.red { color: red; }
.green { color: green; }
</style>

- [Метод главных компонент (Principal Component Analysis, PCA)]{.red}. Преобразует исходные переменные в новые, некоррелированные оси — главные компоненты, которые являются линейными комбинациями исходных данных. Первая компонента объясняет максимум дисперсии, вторая — максимум оставшейся, и так далее. Позволяет сократить данные и устраняет избыточность, часто используется как первый шаг перед другими методами.
- [Факторный анализ (Factor Analysis, FA)]{.red}. В отличие от PCA, который просто преобразует данные, FA ищет скрытые (латентные) причины, порождающие корреляции между наблюдаемыми переменными. Предполагается, что несколько ненаблюдаемых факторов влияют на ответы респондентов. Факторный анализ позволяет найти и проинтерпретировать латентные переменные (например, "религиозность", или "удовлетворенность жизнью") . Это один из самых важных методов для проверки теоретических конструктов. Существует множество методов выделения факторов (метод главных осей, максимального правдоподобия и др.), и выбор между ними — отдельная исследовательская задача.
- [Кластерный анализ переменных / пунктов опросника (Item Cluster Analysis)]{.red}. Хотя кластерный анализ чаще всего применяется для выделения гомогенных групп людей или неживых объектов (товаров и услуг, организаций и пр.), он может применяться и с переменными (например, результатами тестирования). Так, иерархический кластерный анализ образует кластеры, вложенные в другие кластеры, что удобно представить в виде дендрограммы (диаграммы, напоминающей дерево).В зависимости от того, какая мера применяется для вычисления расстояний, будут зависеть результаты, они могут сильно напоминать результаты факторного анализа (например, если в качестве меры сходства используется коэффициент корреляции Пирсона), или быть отличными от него.  

**2. Для категориальных (качественных) данных**:

Эта группа методов специально разработана для анализа данных, представленных в номинальных  (пол, профессия) и порядковых шкалах (образование, шкалы согласия).

- [Анализ соответствий (Correspondence Analysis, CA)]{.red}. Позволяет визуализировать связи между категориями двух или более номинальных переменных в пространстве меньшей размерности. На выходе получается карта, где близость категорий говорит о силе их взаимосвязи. 
- [Множественный анализ соответствий (Multiple Correspondence Analysis, MCA)]{.red} — расширение метода для анализа связей между тремя и более категориальными переменными .
- [Категориальный метод главных компонент (Categorical PCA, CatPCA)]{.red}. Это адаптация классического PCA для работы с порядковыми или номинальными данными. Он оптимальным образом приписывает числовые значения категориям (квантифицирует их) так, чтобы максимизировать объясненную дисперсию.
- [Анализ латентных классов (Latent Class Analysis, LCA)]{.red}. Предполагает, что за наблюдаемыми категориальными переменными стоит одна ненаблюдаемая (латентная) категориальная переменная. Другими словами, популяция неоднородна и состоит из нескольких скрытых классов (латентных групп), которые по-разному отвечают на вопросы.

В данном тьюториале рассматривается первая группа методов (метод главных компонент и факторный анализ), а в "Assignment12" -- вторая группа, методам кластерного анализа посвящен отдельный тренажер --  "Assignment13".  Мы будем следовать материалу, изложенному в разделе 12 [электронного пособия](https://domelia.github.io/R-book/PCA_CA.html), поэтому лучше всего сначала изучить все материалы, а потом уже переходить к практическим занятиям.

Мы выполним тестовые задания и рассмотрим кейсы, иллюстрирующие возможности применения метода главных компонент 

## Тестовые вопросы

### Вопрос 1

Какая из следующих задач НЕ является главной задачей метода главных компонент (PCA)?

```{r q1}
question("", answer("Выявление скрытой структуры в наборе данных"), answer("Снижение размерности данных путем удаления избыточных переменных", correct = TRUE), answer("Идентификация взаимосвязанных переменных"), answer("Визуализация многомерных данных в пространстве меньшей размерности"), incorrect = "Вспомните три главные задачи PCA. Один из вариантов — это распространенное заблуждение, но не основная задача метода.", allow_retry = TRUE )
```

### Вопрос 2

Что представляют собой главные компоненты в PCA?

```{r q2}
question("", answer("Исходные переменные, отсортированные по важности."), answer("Собственные значения матрицы корреляций."), answer("Новые переменные, являющиеся линейной комбинацией исходных переменных.", correct = TRUE), answer("Координаты наблюдений в исходном пространстве."), incorrect = "Вспомните определение метода главных компонент.", allow_retry = TRUE )
```

### Вопрос 3

Для чего обычно выполняется стандартизация переменных перед проведением PCA?

```{r q3}
question("", answer("Чтобы все переменные имели среднее значение 0 и дисперсию 1.", correct = TRUE), answer("Чтобы уменьшить влияние выбросов."), answer("Чтобы преобразовать категориальные переменные в количественные."), answer("Чтобы увеличить объясненную дисперсию первых компонент."), incorrect = "Стандартизация делает переменные сопоставимыми, особенно если они измеряются в разных единицах", allow_retry = TRUE )
```

### Вопрос 4

Как называется график, на котором одновременно отображаются и наблюдения (индивиды), и переменные в пространстве главных компонент?

```{r q4}
question("", answer("Scree plot (график каменистой осыпи)"), answer("Корреляционный круг"), answer("Биплот (biplot)", correct = TRUE), answer("Гистограмма вкладов"), incorrect = "Этот график объединяет информацию и о переменных, и о наблюдениях.", allow_retry = TRUE )
```

### Вопрос 5

Что показывает значение cos2 (квадрат косинуса) для переменной на графике PCA?

```{r q5}
question("", answer("Вклад переменной в формирование главной компоненты (в процентах)."), answer("Корреляцию между переменной и главной компонентой."), answer("Качество представления (репрезентации) переменной на плоскости главных компонент.", correct = TRUE), answer("Значение переменной для конкретного наблюдения."), incorrect = "Качество представленности переменных в факторном пространстве называется cos2 (квадрат косинуса, квадрат координаты)'. Высокое значение cos2 означает, что переменная хорошо описывается данными компонентами.", allow_retry = TRUE )
```

### Вопрос 6

Вы провели PCA и получили следующие собственные значения для первых трех компонент: 4.12, 1.84, 1.24. Общее количество переменных — 10. Какой процент дисперсии объясняют первые две компоненты вместе? (Ответ округлите до целого числа)

```{r q6}
question("", answer("41%"), answer("60%", correct = TRUE), answer("72%"), answer("33%"), incorrect = "Процент дисперсии для компоненты рассчитывается как (ее собственное значение / общее число переменных) * 100%. Для двух компонент проценты нужно сложить. Общее число переменных = сумме всех собственных значений (10).", allow_retry = TRUE )
``` 

### Вопрос 7

Согласно правилу Кайзера, сколько главных компонент следует оставить для интерпретации, если анализ проводился для 8 переменных и получены следующие собственные значения: 3.5, 2.1, 1.2, 0.9, 0.8, 0.3, 0.1, 0.1?

```{r q7}
question("", answer("1"), answer("2"), answer("3", correct = TRUE), answer("4"), incorrect = "Правило Кайзера гласит, что оставляют компоненты с собственным значением > 1. В данном списке это 3.5, 2.1 и 1.2.", allow_retry = TRUE )
```

### Вопрос 8

Что означает тот факт, что на корреляционном круге (графике переменных) переменные расположены близко друг к другу и далеко от центра?

```{r q8}
question("", answer("Эти переменные имеют низкое качество представления (cos2)."), answer("Между этими переменными существует сильная положительная корреляция.", correct = TRUE), answer("Эти переменные имеют маленькую дисперсию."), answer("Эти переменные отрицательно коррелируют с главными компонентами."), incorrect = "Интерпретация корреляционного круга: переменные с положительной корреляцией группируются вместе. Положение далеко от центра говорит о высоком качестве представления (cos2).", allow_retry = TRUE )
```


### Вопрос 9
Для проведения самого анализа методом главных компонент в R, согласно материалу, можно использовать пакет FactoMineR. А для чего в первую очередь предназначен пакет factoextra?

```{r q9}
question("", answer("Для проведения кластерного анализа."), answer("Для расчета собственных векторов и значений."), answer("Для извлечения результатов и визуализации многомерного анализа.", correct = TRUE), answer("Для импорта и экспорта данных."), incorrect = "Сам анализ проводится с помощью других библиотек, библиотека factoextra используется для извлечения результатов и визуализации'.", allow_retry = TRUE )
```

### Вопрос 10

Какая функция из пакета factoextra используется для создания графика, показывающего долю дисперсии, объясняемую каждой главной компонентой (scree plot)?

```{r q10}
question("", answer("fviz_pca_var()"), answer("fviz_pca_biplot()"), answer("fviz_contrib()"), answer("fviz_eig()", correct = TRUE), incorrect = "fviz_eig(res.pca): визуализация собственных значений. Эта функция и строит график 'каменистой осыпи'.", allow_retry = TRUE )
```


### Вопрос 11

Какой из перечисленных объектов НЕ возвращается функцией get_pca_var() после выполнения PCA?

```{r q11}
question("", answer("$coord — координаты переменных"), answer("$cos2 — качество представления переменных"), answer("$contrib — вклад переменных в компоненты"), answer("$dist — расстояние до центра наблюдений", correct = TRUE), incorrect = "Функция get_pca_var()извлекает результаты для *переменных*. Расстояния до центра ($dist) возвращаются для наблюдений (индивидов) функциейget_pca_ind().", allow_retry = TRUE )
```

### Вопрос 12

У вас есть объект res.pca, созданный функцией PCA() из пакета FactoMineR. Какой командой можно получить таблицу с собственными значениями и процентом объясненной дисперсии?

```{r q12}
question("", answer("res.pca$var$coord"), answer("summary(res.pca)"), answer("get_eigenvalue(res.pca)", correct = TRUE), answer("fviz_pca_ind(res.pca)"), incorrect = "get_eigenvalue(res.pca): - извлекает собственные значения / дисперсию главных компонент'. Это наиболее прямая команда из пакета factoextra.",
allow_retry = TRUE
)
```


### Вопрос 13
Посмотрите на график вкладов переменных в первые две компоненты. 

![](https://domelia.github.io/R-book/PCA_CA_files/figure-html/unnamed-chunk-20-1.png){width=600px}

Основываясь на информации из текста, какие переменные вносят наибольший вклад в формирование пространства главных компонент (имеют вклад выше среднего уровня, обозначенного красной линией)?

```{r q13}
question("Это вопрос с множественным выбором.", answer("X100m (бег на 100 м)", correct = TRUE), answer("Shot.put (толкание ядра)"), answer("Long.jump (прыжки в длину)", correct = TRUE), answer("High.jump (прыжки в высоту)"), answer("Pole.vault (прыжки с шестом)", correct = TRUE), allow_retry = TRUE )
```

### Вопрос 14

Анализ главных компонент на данных decathlon2 показал, что переменная X100m (бег на 100 м) имеет отрицательную координату по первой компоненте, а переменная Long.jump (прыжки в длину) — положительную. Что это означает?

```{r q14}
question("", answer("Спортсмены, хорошо бегущие 100 м, как правило, плохо прыгают в длину, и наоборот.", correct = TRUE), answer("Между бегом на 100 м и прыжками в длину нет никакой связи."), answer("Переменная X100m плохо представлена на графике."), answer("Обе переменные вносят малый вклад в первую компоненту."), incorrect = "Переменные с положительной корреляцией сгруппированы вместе; переменные, имеющие отрицательную корреляцию, находятся на разных сторонах графика.", allow_retry = TRUE )
```

### Вопрос 15

В анализе decathlon2 используются "дополнительные наблюдения" (ind.sup). Для чего они нужны?

```{r q15}
question("", answer("Чтобы увеличить размер выборки и повысить точность анализа."), answer("Чтобы их координаты были предсказаны на основе модели, построенной по активным данным, и не влияли на определение главных компонент.", correct = TRUE), answer("Чтобы использовать их для стандартизации данных."), answer("Чтобы определить количество главных компонент."), incorrect = "Дополнительные элементы не используются в основном анализе, а их координаты предсказываются исходя из информации, полученной в результате анализа, полученного на активных данных.", allow_retry = TRUE )
```


## Упражнения: Анализ главных компонент

Разберем возможности метода главных компонент на примере данных `bfi` из библиотеки `psych`, основанных на психологической модели **Big Five (Большой пятерки)** — одной из наиболее признанных теорий личности.


🔬 Происхождение модели

- Сформирована в рамках **лексического подхода**, согласно которому важные личностные различия отражаются в языке.

👉 В результате было выделено **5 устойчивых факторов личности**:

- Доброжелательность  
- Добросовестность  
- Экстраверсия  
- Нейротизм  
- Открытость опыту  

🧪 Опросник BFI:

- Разработан в 1990-е годы  
- Представляет собой **сокращенную версию измерения Big Five**  
- Содержит **25 утверждений** (порядковые шкалы)  
- демографические данные
- пропуски (реалистичные данные)


📦 Структура переменных

- A — Доброжелательность (Agreeableness)


**Переменные:** `A1–A5`

Характеризует:

- доверие к людям  
- альтруизм  
- кооперативность  

⚠️ Обратные шкалы:
- A1  

-  C — Добросовестность (Conscientiousness)


**Переменные:** `C1–C5`

Характеризует: 

- организованность  
- дисциплину  
- ответственность  

⚠️ Обратные:
- C4  

- E — Экстраверсия (Extraversion)


**Переменные:** `E1–E5`

Характеризует:

- общительность  
- энергичность  
- активность  

-  N — Нейротизм (Neuroticism)


**Переменные:** `N1–N5`

Характеризует:

- тревожность  
- эмоциональную нестабильность  
- склонность к стрессу  

- O — Открытость опыту (Openness)


**Переменные:** `O1–O5`

Характеризует:

- креативность  
- любознательность  
- интерес к новому  

- 📌 Дополнительные переменные

  - `age` — возраст  
  - `gender` — пол  
  - `education` — образование  

### Упражнение 1

Как всегда, начнем с просмотра данных. Загрузите набор `bfi` и выведите первые строки.


```{r ex1, exercise = TRUE, eval=TRUE}

```

```{r ex1-hint, eval = FALSE}
head()
```


```{r ex1-solution, include = FALSE, echo=FALSE}
head(bfi)
```

```{r ex1-code-check}
grade_code()
```

```{r ex1-q}
question_radio(
  "Что представляет собой датасет bfi: результаты эксперимента, опросные данные или административные данные?",
  answer("опросные данные", correct = TRUE),
  answer("административные данные"),
  answer("лабораторный эксперимент"),
  allow_retry = TRUE,
  try_again = "Посмотри на значения переменных и подумай, что именно измерялось."
)
```

### Упражнение 2

Вам необходимо:

- Создать объект`df_pca` на основе `bfi`.
- Оставить только первые 25 переменных (без социально-демографических параметров).
- Удалить пропуски с помощью `na.omit()`(в противном случае программа их заменит на средние значения).


```{r ex2, exercise = TRUE, eval=TRUE}

```

```{r ex2-hint, eval = FALSE}
df_pca <- bfi[, ] |>
  ()
```


```{r ex2-solution, include = FALSE}
df_pca <- bfi[, 1:25] |>
  na.omit()
```

```{r ex2-code-check}
grade_code()
```

```{r ex2-q}
question_radio(
  "Почему перед выполнением PCA в этом упражнении нужно удалить пропущенные значения?",
  answer("Потому что PCA строится по числовой матрице без пропусков", correct = TRUE),
  answer("Потому что пропуски автоматически увеличивают число компонент"),
  answer("Потому что пропуски превращают переменные в факторы"),
  allow_retry = TRUE,
  try_again = "Подумай о том, что использует алгоритм PCA на входе."
)
```


### Упражнение 3

Выполните анализ главных компонент для 25 пунктов опросника Big Five.

- Используйте функцию `PCA()` из пакета `FactoMineR`.
- Работайте с объектом `df_pca`, созданным ранее.
- Установите `scale.unit = TRUE`, чтобы стандартизовать переменные.
- Сохраните результат в объект `res.pca`.
- Выведите `res.pca` на экран.

```{r ex3, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex3-solution, include = FALSE}
res.pca <- PCA(df_pca, scale.unit = TRUE, ncp = 10, graph = FALSE)
```

```{r ex3-code-check}
grade_code()
```

```{r ex3-q}
question_radio(
  "Зачем в этом упражнении используется аргумент scale.unit = TRUE?",
  answer("Чтобы все переменные имели одинаковую дисперсию и вклад не зависел от единиц измерения.", correct = TRUE),
  answer("Чтобы увеличить количество главных компонент."),
  answer("Чтобы удалить пропущенные значения."),
  allow_retry = TRUE,
  try_again = "Вспомните, что происходит с переменными при стандартизации."
)
```


### Упражнение 4

Оцените, сколько компонент имеет смысл оставить для интерпретации.

- Извлеките собственные значения с помощью `get_eigenvalue(res.pca)` и сохраните в `eig.val`.
- Посмотрите на долю объясненной дисперсии по компонентам.
- Постройте график каменистой осыпи c помощью функции `fviz_eig()` из пакета `factoextra`.

```{r ex4, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex4-solution, include = FALSE}
eig.val <- get_eigenvalue(res.pca)
eig.val
fviz_eig(res.pca)
```

```{r ex4-code-check}
grade_code()
```

```{r ex4-q}
question(
  "Сколько главных компонент вы бы оставили для интерпретации на основе графика и таблицы собственных значений?",
  answer("1–2 компоненты"),
  answer("3–5 компонент", correct = TRUE),
  answer("10 и более компонент"),
  incorrect = "Ориентируйтесь на правило Кайзера (λ > 1) и на то, какая доля дисперсии суммарно объясняется первыми компонентами.",
  allow_retry = TRUE
)
```


### Упражнение 5

Рассмотрите качество представления переменных и их вклад в первые две компоненты.

- Используйте `get_pca_var(res.pca)` и сохраните результат в `var`.
- Выведите таблицы `var$cos2` и `var$contrib`.

👉  Напомним, `cos2` показывает качество представленности переменной в компоненте, его можно интерпретировать как долю дисперсии переменной, объясненную данной компонентой, `contrib` - это вклад переменной в определение фактора (компоненты), измеряемый в процентах как `var.cos2 * 100) / (total cos2)`.



- Постройте график переменных с помощью `fviz_pca_var(res.pca)`.

👉  Что означает длина стрелки? 

- Координаты переменной по осям = ее нагрузки (корреляции с компонентами).

- Длина стрелки = корень из суммы квадратов этих координат → то есть связана с `cos2`.

- Чем длиннее стрелка (ближе к единичной окружности), тем выше `cos2`, тем лучше переменная представлена выбранными компонентами и тем она «значимее» для интерпретации этого пространства.

```{r ex5, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex5-solution, include = FALSE}
var <- get_pca_var(res.pca)
var$cos2
var$contrib
fviz_pca_var(res.pca, repel = TRUE)
```

```{r ex5-code-check}
grade_code()
```

```{r ex5-q}
question(
  "Какие группы пунктов (A, C, E, N, O) лучше всего представлены первыми двумя компонентами (имеют высокие cos2 и вклад)?",
  answer("Преимущественно пункты, относящиеся к одному–двум факторам личности (например, Экстраверсии и Нейротизму).", correct = TRUE),
  answer("Все пункты представлены примерно одинаково хорошо."),
  answer("Лучше всего представлены демографические переменные (age, gender, education)."),
  incorrect = "Посмотрите, какие переменные находятся далеко от центра на корреляционном круге и имеют высокие cos2 и contrib.",
  allow_retry = TRUE
)
```

---

### Упражнение 6

Постройте биплот, чтобы одновременно увидеть индивидов и переменные в пространстве первых двух компонент.

- Используйте функцию `fviz_pca_biplot(res.pca, repel = TRUE)`.

```{r ex6, exercise = TRUE, eval=TRUE}
# ваш код
```

```{r ex6-solution, include = FALSE}
fviz_pca_biplot(res.pca, repel = TRUE)
```

```{r ex6-code-check}
grade_code()
```

```{r ex6-q}
question_radio(
  "Что показывает направление и длина стрелок на биплоте?",
  answer("Средние значения переменных."),
  answer("Корреляцию переменных с главными компонентами и их вклад в пространство компонент.", correct = TRUE),
  answer("Размер выборки."),
  allow_retry = TRUE,
  try_again = "Стрелки на биплоте интерпретируются так же, как и на корреляционном круге."
)
```


### Упражнение 7 - Координаты переменных и индивидов - типичные паттерны ответов и профилей по переменным


Что такое координаты?


👉  В PCA «координаты» – это  числа, показывающие, **где именно расположен объект (индивид) или переменная на осях главных компонент**.


**1. Координаты индивидов (scores)**

- При PCA каждое наблюдение проецируется на новые оси $PC_1, PC_2, \dots$. Получающиеся значения называются **scores** или **координаты индивидов**: в **FactoMineR** это `ind$coord`. 
- Геометрически: это «адрес» точки на плоскости/в пространстве компонент, то есть те числа, которые мы видим на графике индивидов.

Формально: координаты = исходные стандартизированные данные, умноженные на матрицу собственных векторов. 

**2. Координаты переменных**

- Аналогично можно спроецировать **переменные** на те же оси. 
- В FactoMineR/factoextra `var$coord` = **координаты переменных** по компонентам. 
- На корреляционном круге именно эти координаты являются концами стрелок (по оси X – Dim.1, по оси Y – Dim.2 и т.д.). 

Связь с нагрузками:

- В случае PCA по корреляционной матрице координаты переменных по оси $PC_k$ совпадают с **нагрузками** с этой компонентой 
- Отсюда: `var$coord`, `cos2` и `loadings` – это один и тот же показатель, представляемый по-разному.

Если попытаться придать этим числам смысл: мы получим, что у индивидов координата означает, насколько выражен паттерн (поведения, отношения к жизни и пр.), представленный обобщенной переменной (компонентой).

Большая положительная координата по $PC_1$  → у индивида высоки те переменные, которые имеют большие положительные нагрузки на $PC_1$ и низки – по переменным с отрицательными нагрузками. Это не среднее и не «общий балл» по всем вопросам, а взвешенная комбинация переменных, формирующих фактор.

Координата переменной по компоненте = «насколько эта переменная воплощает/отражает смысл этой компоненте - и, одновременно, сколько в самой переменной содержится ее обобщенного смысла».


**Задание**:

Создайте новую переменную на основе первой главной компоненты.

- Извлеките координаты индивидов с помощью `get_pca_ind(res.pca)` и сохраните их в `ind`.
- Сохраните координаты по первой компоненте в новый вектор `pc1`.
- Добавьте `pc1` в исходный датафрейм `df_pca`.
- Выведите на экран первые значения `pc1`.

```{r ex7, exercise = TRUE, eval=TRUE}
# ваш код
```

```{r ex7-hint, eval = FALSE}
ind <- get_pca_ind(res.pca)
pc1 <- ind$coord[, 1]
df_pca$pc1<-pc1
head(pc1)
```

```{r ex7-solution, include = FALSE}
ind <- get_pca_ind(res.pca)
pc1 <- ind$coord[, 1]
df_pca_pc1 <- df_pca |>
  mutate(pc1 = pc1)
head(df_pca_pc1$pc1)
```

```{r ex7-code-check}
grade_code()
```

```{r ex7-q}
question_radio(
  "Как можно использовать переменную pc1 в дальнейшем анализе?",
  answer("Как интегральный показатель, отражающий сочетание нескольких пунктов опросника.", correct = TRUE),
  answer("Только для проверки нормальности распределения."),
  answer("Переменная pc1 не имеет практического смысла."),
  allow_retry = TRUE,
  try_again = "Вспомните, что главные компоненты могут использоваться как новые предикторы или индексы."
)
```

## Упражнения: Факторный анализ

В этом блоке упражнений рассмотрим особенности применения факторного анализа на тех же данных - `bfi`. Для простоты отберем переменные только по двум блокам - экстраверсии и нейротизму.

### Упраждение 8

Создайте датафрейм для факторного анализа.

- Оставьте только переменные `E1:E5` (Экстраверсия) и `N1:N5` (Нейротизм).
- Удалите пропуски с помощью `na.omit()`.
- Сохраните результат в объект `df_fa`.

```{r ex8, exercise = TRUE, eval=TRUE}
# ваш код
```

```{r ex8-hint, eval = FALSE}
df_fa <- bfi |>
  select(paste0("E", 1:5), paste0("N", 1:5)) |> 
           na.omit()
```

```{r ex8-solution, include = FALSE}
df_fa <- bfi |>
  select(paste0("E", 1:5), paste0("N", 1:5)) |> 
           na.omit()
```

```{r ex8-code-check}
grade_this({
  # проверяем, что объект существует
  if (!exists("df_fa")) {
    fail("Объект df_fa не найден.")
  }

  # проверяем, что это data.frame / tibble
  if (!inherits(df_fa, c("data.frame", "tbl_df"))) {
    fail("df_fa должен быть таблицей (data.frame / tibble).")
  }

  # правильные имена колонок
  needed <- c(paste0("E", 1:5), paste0("N", 1:5))

  if (!setequal(colnames(df_fa), needed)) {
    fail(paste0(
      "В df_fa должны быть только переменные: ",
      paste(needed, collapse = ", "), "."
    ))
  }

  pass("Отлично! df_fa сформирован корректно.")
})
```

```{r ex8-q}
question_radio(
  "Почему для факторного анализа удобно сначала сузить набор переменных до теоретически связанной группы?",
  answer("Это упрощает интерпретацию факторов и позволяет проверить теоретическую модель.", correct = TRUE),
  answer("Потому что факторный анализ не работает с большим числом переменных."),
  answer("Потому что факторный анализ можно проводить только на 10 переменных."),
  allow_retry = TRUE,
  try_again = "Подумайте про связь между теоретическими конструктами и латентными факторами."
)
```


### Упражнение 9 -- Проверяем пригодность данных для факторного анализа

Зачем нам проверять, пригодны ли данные для анализа, если мы и так уже отобрали нужные переменные?

 Факторный анализ предполагает, что **у переменных есть общая часть дисперсии**, которая объясняется факторами, а часть – своя собственная, **специфическая** дисперсия + ошибка.  
 
 - Проверка пригодности нужна, чтобы убедиться, что между переменными достаточно общих связей, которые имеет смысл сворачивать в общие факторы, и что матрица корреляций существенно отличается от единичной - то есть такой, где по диагонали -- единицы, а в остальных ячейках -- нули.

- Если переменные почти не коррелируют или корреляции не образуют связной структуры, то выделять факторы бессмысленно: модель _общих факторов_ не подходит.


#### KMO (Measure of Sampling Adequacy)

KMO оценивает, насколько хорошо переменные подходят для факторного анализа с точки зрения **соотношения обычных корреляций и частных корреляций**. 

Пусть:

- $r_{ij}$ – обычная корреляция между переменными $i$ и $j$  
- $p_{ij}$ – частная корреляция между $i$ и $j$ при контроле всех остальных переменных

Тогда общая мера KMO:

$$ KMO = \dfrac{\sum_{i \ne j} r_{ij}^{2}}{\sum_{i \ne j} r_{ij}^{2} + \sum_{i \ne j} p_{ij}^{2}} $$

Интерпретация:

- $KMO \to 1$ → обычные корреляции велики, частные малы → структура хорошо описывается общими факторами, данные **подходят** для факторного анализа. 
- $KMO \to 0$ → обычные корреляции малы, частные велики → связи скорее уникальны/случайны, факторный анализ **нежелателен**.  
- Эмпирические «правила большого пальца»:  
  - $KMO \ge 0.8$ – очень хорошо  
  - $0.7–0.8$ – хорошо  
  - $0.6–0.7$ – приемлемо  
  - $< 0.6$ – лучше поискать другую модель / переменные. 
  
Функция `KMO()` в пакете `psych` возвращает как общую меру, так и KMO по отдельным переменным.


#### Тест сферичности Бартлетта

Тест Бартлетта проверяет гипотезу о том, что **матрица корреляций равна единичной**, т.е. все переменные некоррелированы в популяции. 

- Нулевая гипотеза:  
  $H_0: \mathbf{R} = \mathbf{I}$ (нет общей факторной структуры).  
- Альтернативная:  
  $H_1: \mathbf{R} \ne \mathbf{I}$ (есть существенные корреляции, матрица факторизуема).

Тестовая статистика (для выборки объема $n$, $p$ переменных):

$\chi^2 = -\left(n - 1 - \dfrac{2p + 5}{6}\right) \ln|\mathbf{R}|$

где $|\mathbf{R}|$ – определитель матрицы корреляций.  

- При справедливости $H_0$ эта статистика асимптотически имеет распределение $\chi^2$ с $\dfrac{p(p-1)}{2}$ степенями свободы.
- Малое значение $|\mathbf{R}|$ (корреляции значимы) → большое значение $\chi^2$ → маленький $p$ → **отклоняем $H_0$** и делаем вывод, что матрица пригодна для факторного анализа.  


Если KMO ≥ 0.6–0.7 и тест Бартлетта значим (p < 0.05), можно считать, что данные **подходят** для факторного анализа и переходить к выбору числа факторов и их интерпретации. 


**Задание:**

Проверьте пригодность данных для факторного анализа.

- Рассчитайте KMO‑критерий с помощью `KMO()` из пакета `psych`.
- Рассчитайте тест сферичности Бартлетта с помощью `cortest.bartlett()`.

```{r ex9, exercise = TRUE, eval=TRUE}
# ваш код
```

```{r ex9-hint, eval = FALSE}
KMO(df_fa)
cortest.bartlett(cor(df_fa), n = nrow(df_fa))
```

```{r ex9-solution, include = FALSE}
KMO(df_fa)
cortest.bartlett(cor(df_fa), n = nrow(df_fa))
```

```{r ex9-code-check}
grade_code()
```

```{r ex9-q}
question(
  "Как интерпретировать результаты KMO и теста Бартлетта в вашем случае?",
  answer("KMO достаточно высок, а тест Бартлетта значим, данные подходят для факторного анализа.", correct = TRUE),
  answer("KMO очень низкий, а тест Бартлетта незначим, факторный анализ неприменим."),
  answer("Эти критерии не имеют отношения к факторному анализу."),
  incorrect = "Напомнить: KMO > 0.6 обычно считается приемлемым, а значимый тест Бартлетта говорит о наличии корреляций.",
  allow_retry = TRUE
)
```


### Упражнение 10

Выполните факторный анализ с двумя факторами.

- Используйте функцию `fa()` из пакета `psych`.
- Задайте `nfactors = 2`, `rotate = "varimax", fm="ml"`. Поскольку мы заранее знаем, что наши переменные должны образовывать два фактора - экстраверсии и нейротизма, логично было бы это проверить. Вращение "Варимакс" необходимо для получения простой факторной структуры (когда переменные имеют большие нагрузки по одним факторам, и маленькие -- по другим), а выбор метод `ml` -- максимального правдоподобия может быть объяснен тем, что он позволяет адекватно и более строго оценивать метрики качества, что не совсем позволяет метод по умолчанию -- `minres` - метод нелинейных наименьших квадратов.   


<details><summary><strong>▶ Подробнее о методах извлечения факторов в psych::fa()</strong></summary>

**Методы извлечения факторов в `fa()`**

Функция `fa()` поддерживает несколько методов (аргумент `fm`):

- `"minres"` – минимизация остатков (minimum residuals).
- `"pa"` – principal axis factoring.
- `"ml"` – maximum likelihood (максимальное правдоподобие).
- плюс другие (`"wls"`, `"gls"` и др.).

**`"minres"` (minimum residuals)**

- **Идея**: найти такие нагрузки, чтобы сумма квадратов **остатков корреляций** (разницы между наблюдаемой матрицей корреляций и смоделированной) была минимальной.  
- **Интуиция**: подгоняем модель так, чтобы максимально хорошо воспроизвести взаимные корреляции между переменными.  
- **Плюсы**:
  - довольно устойчив к нарушениям нормальности и «грязным» данным;
  - часто сходится там, где ML более капризен;
  - удобен как базовый метод для разведочного и учебного анализа.
- **Минусы**:
  - не является ML–оцениванием, поэтому классические χ²‑тесты подгонки и SE параметров интерпретируются более ограниченно;
  - fit‑индексы носят более эвристический характер.


**`"pa"` (principal axis factoring)**

- **Идея**: классический метод **главных осей** (principal axis factoring), ориентирован на моделирование **общей дисперсии**.  
- **Как работает**:
  - диагональ корреляционной матрицы заменяется оценками коммунальностей;
  - далее итеративно извлекаются факторы, объясняющие общую часть дисперсии, а не полную, как в PCA.
- **Отличие от PCA**:
  - PCA использует полную матрицу корреляций/ковариаций;
  - PA работает с «очищенной» матрицей, где на диагонали стоят коммунальности.
- **Когда полезен**:
  - когда нужен «чистый» факторный подход без ML;
  - как классический метод EFA в психометрической традиции.


**`"ml"` (maximum likelihood)**

- **Идея**: подобрать нагрузки и уникальные дисперсии так, чтобы максимизировать **функцию правдоподобия** при предположении многомерной нормальности.  
- **Преимущества**:
  - можно формально тестировать, **достаточно ли заданного числа факторов** (χ²‑тест подгонки модели);
  - доступны **стандартные ошибки** нагрузок и других параметров;
  - можно использовать **AIC, BIC** для сравнения конкурирующих моделей;
  - метод естественно встраивается в рамку SEM (структурное моделирование).
- **Ограничения**:
  - чувствителен к нарушениям многомерной нормальности;
  - может быть менее устойчив при маленьких выборках или сильно «кривых» распределениях.


**`"wls"`, `"gls"` и другие**

- `"wls"` (weighted least squares) и `"gls"` (generalized least squares) – варианты взвешенного/обобщённого МНК:

  - учитывают неодинаковую надёжность элементов матрицы корреляций;
  - полезны при работе с особыми типами корреляций (полихорические и т.п.) и сложными структурами данных.

</details>

- Сохраните результат в объект `res.fa`.

Почему мы используем именно этот пакет? Коротко: потому что psych::fa() – это специализированный инструмент именно для общего факторного анализа в психометрии, а FactoMineR в основном «заточен» под методы главных компонент и «родственных» им описательных методов, а не под классический FA.

```{r ex10, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex10-solution, include = FALSE}
res.fa <- fa(df_fa, nfactors = 2, rotate = "varimax", fm="ml")
res.fa
```

```{r ex10-code-check}
grade_code()
```

```{r ex10-q}
question_radio(
  "Чем факторный анализ отличается от PCA в контексте этой задачи?",
  answer("Факторный анализ предполагает наличие латентных факторов, порождающих корреляции между переменными.", correct = TRUE),
  answer("Факторный анализ просто пересортировывает переменные по дисперсии."),
  answer("Отличий нет, это одно и то же."),
  allow_retry = TRUE,
  try_again = "Сравните цель PCA (снижение размерности) и факторного анализа (модель латентных факторов)."
)
```

В выводе содержится довольно много информации, о чем она?

#### Гипотеза «2 фактора достаточно»

- Сначала оценивается **нулевая модель** (все переменные некоррелированы):  
  - df null model = 45, Chi Square = 8217.3 – очень плохое соответствие, как и ожидается;  
- Для нашей **2‑факторной модели**:  
  - df model = 26, objective function = 0.22 – намного лучше, чем нулевая модель.

Потом считаются несколько индексов подгонки. В блоке:

> The harmonic n.obs … empirical chi square 428.7, prob < 1.7e-74  
> Likelihood Chi Square = 586.99, prob < 3.1e-107  

- Очень маленькие p‑значения ⇒ модель **формально не идеально согласуется** с данными (что при больших выборках почти всегда). 

#### Остатки: RMSR

> RMSR = 0.04, df-corrected RMSR = 0.06  
> Fit based upon off diagonal values = 0.98  

- RMSR (root mean square of residuals) – среднеквадратичная ошибка аппроксимации корреляций. 

- 0.04–0.06 – довольно **малые остатки**, а fit по вне‑диагональным значениям 0.98 говорит, что модель хорошо воспроизводит корреляции между переменными (off-diagonal). 


####  Индексы глобальной подгонки

> Tucker Lewis Index (TLI) = 0.881  
> RMSEA = 0.091 (90% CI: 0.085–0.097)  
> BIC = 382.37  

- **TLI ≈ 0.88** – чаще всего порог «хорошо» ставят 0.90–0.95, то есть модель **чуть не дотягивает до «хорошей», но лучше, чем совсем плохая**. - **RMSEA ≈ 0.09** (CI 0.085–0.097):  
  - < 0.05 – хорошее соответствие,  
  - 0.05–0.08 – приемлемое,  
  - 0.08–0.10 – «пограничное»,  
  - > 0.10 – плохое. 
  
  Здесь: **скорее «так себе»: лучше, чем совсем плохая модель, но заметно хуже желаемого < 0.08**.  
- **BIC** полезен для сравнения моделей: чем меньше, тем лучше; его смысл проявится, если сравнивать решения с 1, 2, 3 факторами. 

Итого по глобальным индексам: 2‑факторная модель **объясняет корреляции, но не идеально**, возможно, данные «просят» больше факторов или имеют сложную структуру.

#### Адекватность факторных баллов

Блок:

> Correlation of scores with factors: 0.92, 0.88  
> Multiple R² of scores with factors: 0.85, 0.78  
> Minimum correlation of possible factor scores: 0.71, 0.56  

Интерпретация: 

- Корреляция регрессионных оценок факторных баллов с истинными факторами:  
  - Фактор 1: 0.92, Фактор 2: 0.88 – **очень хорошие оценки**, мало искажений.  
- $Multiple\space R^2$  – доля дисперсии «истинного» фактора, объясненная факторными баллами:  
  - 0.85 и 0.78 – большая часть вариации фактора восстановлена.  
- Минимальная корреляция возможных факторных баллов (по Брауну/Томсону):  
  - 0.71, 0.56 – даже «худшие» возможные оценки факторов (при другом методе расчета баллов) все равно достаточно хорошо с ними связаны, особенно для первого фактора.  

Вывод: **факторные баллы достаточно надежны для дальнейшего анализа**, особенно по первому фактору.

Таким образом, двухфакторная модель **лучше нулевой**, неплохо воспроизводит корреляции (RMSR низкий, fit 0.98). Факторные оценки по обоим факторам **адекватны и пригодны** для использования в регрессиях, групповых сравнениях и т.д.


### Упражнение 11

Проинтерпретируйте факторные нагрузки.

- Посмотрите на матрицу нагрузок `res.fa$loadings`.
- Определите, какие пункты загружаются на первый фактор, а какие – на второй.
- Попробуйте дать содержательные названия факторам.

```{r ex11, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex11-solution, include = FALSE}
res.fa$loadings
```

```{r ex11-code-check}
grade_code()
```

```{r ex11-q}
question(
  "Как вы бы назвали первый и второй факторы, исходя из того, какие пункты на них загружаются?",
  answer("Первый фактор — Нейротизм, второй — Экстраверсия.", correct = TRUE),
  answer("Оба фактора отражают один и тот же конструкт."),
  answer("Названия факторов не имеют смысла и не используются."),
  incorrect = "Сопоставьте содержание утверждений (E1–E5, N1–N5) с тем, на какие факторы они загружаются.",
  allow_retry = TRUE
)
```


### Упражнение 12

Получите факторные баллы и сравните их между группами.

- Сохраните факторные баллы в объект `scores_fa` (`res.fa$scores` или с помощью `factor.scores()`).
- Объедините факторные баллы с исходными демографическими переменными (`gender` или `age`) в `fa_df`.
- Постройте простой график boxplot для одного фактора по группам пола.

```{r ex12, exercise = TRUE, eval=TRUE}
# ваш код
```


```{r ex12-solution, include = FALSE}
scores_fa <- res.fa$scores
fa_df <- cbind(scores_fa,
               bfi[rownames(df_fa), c("gender","age")])
boxplot(ML1 ~ gender, data = fa_df)
```

```{r ex12-code-check}
grade_code()
```

```{r ex12-q}
question_radio(
  "Как можно интерпретировать различия в факторных баллах между группами пола?",
  answer("Как различия в уровне латентной характеристики, измеряемой фактором, между мужчинами и женщинами.", correct = TRUE),
  answer("Как различия только в уровне шума измерения."),
  answer("Факторные баллы нельзя сравнивать между группами."),
  allow_retry = TRUE,
  try_again = "Помните, что факторные баллы — это оценки латентных переменных на уровне индивидов."
)
```



```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
