---
title: "Регрессионный анализ"
output: 
  learnr::tutorial:
    language: ru.json
runtime: shiny_prerendered
---

```{r setup-locale, include=FALSE}
Sys.setenv(LANGUAGE = "ru_RU.UTF-8")
Sys.setlocale("LC_ALL", "Russian")
```

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(tutorial.helpers)
library(tidyverse)
library(caret)
library(dplyr)
library(olsrr)
knitr::opts_chunk$set(echo=FALSE, warning=FALSE, message=FALSE)
life<-read.csv("data/life.csv")
mod_full <- life |>
  select(-c(Country, Status, Year)) |>
  lm(Life.expectancy ~ ., data = _)
model_backward <- olsrr::ols_step_backward_p(mod_full)
model_forward <- olsrr::ols_step_forward_p(mod_full)
mod_back    <- model_backward$model
mod_forward <- model_forward$model
tutorial_options(
  exercise.reveal_solution = FALSE
)


```


```{r info-section, child = "child_documents/info_section.Rmd"}
```

## План действий

Тема регрессионного анализа в курсе разбита на две части. Первая часть - данный тьюториал, посвящена введению в регрессионный анализ и разбору случая множественной линейной регрессии. Вторая часть содержится в отдельном тьюториале "Assignment 10" и посвящена бинарной логистической регресии. 

В теоретическом плане тренажеры следуют материалу, изложенному в разделе 11 [электронного пособия](https://domelia.github.io/R-book/Regression.html), поэтому лучше всего сначала изучить все материалы, а потом уже переходить к практическим занятиям.


## Предварительные вопросы

```{r regression-quiz, echo=FALSE}
quiz(
  caption = "Тест: Регрессионный анализ",

  question_radio(
    "Что оценивает коэффициент детерминации $R^2$?",
    answer("Среднюю ошибку предсказания"),
    answer("Долю дисперсии зависимой переменной, объяснённую моделью", correct = TRUE),
    answer("Статистическую значимость предикторов"),
    answer("Силу корреляции между предикторами"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какое допущение линейной регрессии требует, чтобы остатки имели постоянную дисперсию?",
    answer("Нормальность остатков"),
    answer("Независимость наблюдений"),
    answer("Гомоскедастичность", correct = TRUE),
    answer("Линейность"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Что означает p-значение < 0.05 для коэффициента регрессии?",
    answer("Коэффициент практически значим"),
    answer("Модель объясняет более 95% дисперсии"),
    answer("Коэффициент статистически значимо отличается от нуля", correct = TRUE),
    answer("Между переменными есть причинно-следственная связь"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какой метод чаще всего используется для оценки коэффициентов линейной регрессии?",
    answer("Метод максимального правдоподобия"),
    answer("Метод наименьших квадратов (МНК)", correct = TRUE),
    answer("Байесовский метод"),
    answer("Метод моментов"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Что такое мультиколлинеарность в регрессионном анализе?",
    answer("Нелинейная связь между Y и X"),
    answer("Высокая корреляция между независимыми переменными", correct = TRUE),
    answer("Гетероскедастичность остатков"),
    answer("Автокорреляция остатков"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какой показатель используют для диагностики мультиколлинеарности?",
    answer("Критерий Дарбина–Уотсона"),
    answer("Тест Бройша–Пагана"),
    answer("Variance Inflation Factor (VIF)", correct = TRUE),
    answer("Тест Шапиро–Уилка"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Как интерпретируется коэффициент $\\beta = 3.5$ в модели $y = \\beta_0 + \\beta_1x$?",
    answer("При X = 0 значение Y равно 3.5"),
    answer("При увеличении X на 1 единицу Y увеличивается в среднем на 3.5", correct = TRUE),
    answer("Корреляция между X и Y равна 3.5"),
    answer("Стандартное отклонение Y равно 3.5"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Что тестирует F-статистика в регрессионной модели?",
    answer("Нормальность остатков"),
    answer("Значимость каждого коэффициента по отдельности"),
    answer("Значимость модели в целом — все ли коэффициенты одновременно равны нулю", correct = TRUE),
    answer("Гомоскедастичность"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Что происходит со значением $R^2$ при добавлении новых предикторов в модель?",
    answer("Всегда уменьшается"),
    answer("Не изменяется"),
    answer("Всегда увеличивается или остается прежним", correct = TRUE),
    answer("Изменяется непредсказуемо"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какой показатель, в отличие от $R^2$, штрафует за добавление незначимых предикторов?",
    answer("Стандартная ошибка регрессии"),
    answer("Скорректированный $R^2$ (Adjusted $R^2$)", correct = TRUE),
    answer("Коэффициент корреляции Пирсона"),
    answer("AIC (информационный критерий Акаике)"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Как называется ситуация, когда дисперсия остатков зависит от значений предиктора?",
    answer("Мультиколлинеарность"),
    answer("Автокорреляция"),
    answer("Гетероскедастичность", correct = TRUE),
    answer("Эндогенность"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какой тест используется для проверки нормальности остатков?",
    answer("Тест Левена"),
    answer("Тест Дарбина–Уотсона"),
    answer("Тест Шапиро–Уилка", correct = TRUE),
    answer("Тест Бройша–Пагана"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Что такое «влиятельное наблюдение» (influential observation) в регрессии?",
    answer("Наблюдение с наибольшим значением Y"),
    answer("Наблюдение, удаление которого существенно меняет коэффициенты модели", correct = TRUE),
    answer("Наблюдение с наибольшим остатком"),
    answer("Наблюдение с пропущенными значениями"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),

  question_radio(
    "Какая метрика измеряет влияние наблюдения на коэффициенты регрессии?",
    answer("Стандартизованный остаток"),
    answer("Расстояние Кука (Cook's distance)", correct = TRUE),
    answer("VIF"),
    answer("Значение плеча (leverage)"),
    allow_retry = TRUE,
    random_answer_order = TRUE
  ),
  question(
  "Что такое мультиколлинеарность в модели множественной линейной регрессии и чем она опасна?",
  answer("Ситуация, когда зависимая переменная линейно зависит от ошибки, что увеличивает смещение оценок", correct = FALSE),
  answer("Ситуация, когда объясняющие переменные сильно коррелируют между собой, что приводит к нестабильности оценок коэффициентов и росту их дисперсий", correct = TRUE),
  answer("Ситуация, когда остатки модели не имеют нормального распределения", correct = FALSE),
  answer("Ситуация, когда в модели присутствует один предиктор вместо нескольких", correct = FALSE),
  allow_retry = TRUE
)
)
```

## Практические задания

Чтобы отработать навыки проведения регрессионного анализа, используем набор данных **Life Expectancy (WHO)** — это таблица показателей здоровья и развития разных стран, собранная на основе статистики Всемирной организации здравоохранения и ООН. Он содержит сведения по 193 странам мира за 2000–2015 гг. по 22 показателям и позволяет изучать, какие факторы связаны с продолжительностью жизни.

![](images\who.jpg){width=400px}

- Единица анализа: страна–год (одна строка ≈ одна страна в конкретном году).  
- Целевая переменная: **life expectancy** — ожидаемая продолжительность жизни при рождении (в годах).  
- Географический охват: 193 страны.  
- Временной охват: последовательные годы (2000–2015).  

Переменные:

- **Country** — страна.  
- **Year** — год наблюдения.  
- **Status** — статус  — развитая или развивающаяся страна.
- **Life_expectancy** — ожидаемая продолжительность жизни при рождении.  
- **Adult_Mortality** — смертность взрослых (обычно вероятность умереть между 15 и 60 годами на 1000 человек).  
- **Infant_deaths** — смертность младенцев (обычно число смертей до 1 года на 1000 или на 1000 рождений).  
- **Alcohol** — потребление алкоголя на душу населения среди лиц 15+ лет, литров чистого алкоголя в год. Характеризует уровень употребления алкоголя в стране как фактор риска для здоровья.
- **percentage expenditure** — доля расходов на здравоохранение в процентах от валового внутреннего продукта. Отражает приоритетность здравоохранения в структуре экономических затрат.
- **Hepatitis B**  - охват вакцинацией против гепатита B среди детей (обычно в процентах соответствующей возрастной группы). Используется как индикатор доступности и эффективности программ иммунизации.
- **Measles**  - показатель заболеваемости корью (часто — число зарегистрированных случаев на 1000 населения или абсолютное число случаев за год). Характеризует эпидемиологическую ситуацию и качество вакцинации.
- **BMI**  - средний индекс массы тела населения (Body Mass Index), безразмерная величина. Отражает распространенность избыточного веса и ожирения как факторов риска хронических заболеваний.
- **under-five deaths**  — смертность детей до 5 лет (количество случаев смерти в этой возрастной группе, обычно на 1000 живорождений или за год). Ключевой индикатор здоровья детского населения и эффективности системы здравоохранения.
- **Polio**  — охват вакцинацией против полиомиелита (в процентах детей целевой возрастной группы). Является показателем выполнения программ иммунизации и контроля над инфекционными заболеваниями.
- **Total expenditure**  — совокупные расходы на здравоохранение (как правило, в процентах от государственных расходов или ВВП). Характеризует общий ресурс, выделяемый государством и/или обществом на систему здравоохранения.
- **Diphtheria** - охват вакцинацией против дифтерии (часто — комбинированная вакцина DTP: дифтерия, коклюш, столбняк), в процентах детей. Индикатор качества программ детской вакцинации.
- **HIV/AIDS**  - показатель смертности или заболеваемости, связанной с ВИЧ/СПИДом (как правило, число смертей на 1000 или 100 000 населения). Отражает бремя эпидемии ВИЧ в стране.
- **GDP**  - валовой внутренний продукт на душу населения, как правило в текущих долларах США. Используется как макроэкономический индикатор уровня экономического развития страны.
- **Population** - численность населения страны (обычно в абсолютных величинах). Позволяет учитывать масштаб страны и использовать нормированные показатели (на 1000 или 100 000 человек).
- **thinness 1-19 years**  - доля населения в возрасте 1–19 лет с дефицитом массы тела (thinness), как правило в процентах. Индикатор нутритивного статуса детей и подростков, связанный с недоеданием.
- **thinness 5-9 years**  - доля детей 5–9 лет с дефицитом массы тела, в процентах. Более узкий возрастной индикатор нутритивной недостаточности.
- **Income composition of resources**  — индекс композиции дохода (основан на показателе человеческого развития, HDI), отражающий распределение доходов и доступность ресурсов для населения. Интерпретируется как обобщенный социально‑экономический индикатор благосостояния.
- **Schooling** — ожидаемая продолжительность обучения, среднее количество лет обучения, которое человек может получить в данной стране. Характеризует уровень развития системы образования и человеческого капитала.

### Упражнение 1 — Загрузка данных

Давайте посмотрим на наши данные. 
Загрузите данные, указав имя набора - `life`
```{r data, exercise = TRUE, eval=TRUE}

```

```{r data-hint-1, eval = FALSE}
life
```


```{r data-solution, include = FALSE}
life
```

```{r data-test, include = FALSE}
life
```
```{r data-code-check}
grade_code()
```

### Упражнение 2 — Анализ взаимосвязей между переменными

Создадим график взаимосвязей между всеми предикторами (кроме года) и зависимой переменной. 

```{r echo=FALSE}
life<-life |> 
  filter(Year==2015)

```

Вам необходимо отобрать все переменные, кроме `Country, Status, Year, Life.expectancy` в объект `x`, а целевую переменную `ife$Life.expectancy` в объект `y`. Затем используйте функцию `featurePlot()`из пакета `caret` для того, чтобы сделать визуализацию.

Учитывайте, что пакет `caret` уже загружен, и дополнительно его загружать не нужно. 
```{r ex2, exercise = TRUE, eval=TRUE}
x <- life |> select()
y <- life$

featurePlot(x = , y = )
```

```{r ex2-hint-1, eval = FALSE}
x <- life |> select(-c(..., ..., ..., ...))
y <- life$...

featurePlot(x = ..., y = ...)
```


```{r ex2-solution, include = FALSE}
x <- life |> select(-c(Country, Status, Year, Life.expectancy))
y <- life$Life.expectancy

featurePlot(x = x, y = y)
```

```{r ex2-test, include = FALSE}
x <- life |> select(-c(Country, Status, Year, Life.expectancy))
y <- life$Life.expectancy

featurePlot(x = x, y = y)
```
```{r ex2-code-check}
grade_code()
```


```{r ex2-q}
question_text(
    "Какие переменные Вам кажутся лучшими кандидатами в предикторы?",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 3 - Анализ общей модели

Следующий шаг – анализ общей модели множественной линейной регрессии, куда в качестве зависимой переменной будет включена  `Life.expectancy`, а в независимые -- все остальные переменные, кроме `Country, Status, Year`.

Создайте пайп, в котором на первой "ступеньке" будут загружаться данные, на второй -- будет произведен отбор переменных, на третьей -- построена модель, на четвертой -- создастся таблица с результатами. Вместо классического `summary()` воспользуйтесь `sjPlot::tab_model()`.

Вам нужно дополнить код в упражнении. Обратите внимание, что `data=_` используется, чтобы не упоминать дважды источник данных. Оператор "_" – это такая "заглушка" для объекта, который проходит по пайпу.

```{r ex3, exercise = TRUE, eval=TRUE}
life |> 
  select() |> 
  lm(, data=_) |> 
  sjPlot::tab_model()
```

```{r ex3-hint-1, eval = FALSE}
life |> 
  select(-c(..., ..., ...)) |> 
  lm(Life.expectancy ~ ., data=_) |> 
  sjPlot::tab_model()
```


```{r ex3-solution, include = FALSE}
life |> 
  select(-c(Country, Status, Year)) |> 
  lm(Life.expectancy ~ ., data=_) |> 
  sjPlot::tab_model()
```

```{r ex3-test, include = FALSE}
life |> 
  select(-c(Country, Status, Year)) |> 
  lm(Life.expectancy ~ ., data=_) |> 
  sjPlot::tab_model()
```
```{r ex3-code-check}
grade_code()
```


```{r ex3-q}
question_checkbox(
  "Какие предикторы оказались статистически значимыми в модели (p < 0.05)?",
  answer("Adult.Mortality",                correct = TRUE),
  answer("infant.deaths",                 correct = TRUE),
  answer("Alcohol",                       correct = FALSE),
  answer("percentage.expenditure",        correct = TRUE),
  answer("Hepatitis.B",                   correct = FALSE),
  answer("Measles",                       correct = FALSE),
  answer("BMI",                           correct = TRUE),
  answer("under.five.deaths",             correct = TRUE),
  answer("Polio",                         correct = FALSE),
  answer("Total.expenditure",             correct = FALSE),  # p = 0.050 — можно обсудить порог
  answer("Diphtheria",                    correct = TRUE),
  answer("HIV.AIDS",                      correct = TRUE),
  answer("GDP",                           correct = FALSE),
  answer("Population",                    correct = FALSE),
  answer("thinness..1.19.years",          correct = FALSE),
  answer("thinness.5.9.years",            correct = FALSE),
  answer("Income.composition.of.resources", correct = TRUE),
  answer("Schooling",                     correct = TRUE),
  allow_retry = TRUE,
  random_answer_order = TRUE,
  try_again = "Посмотрите на p-value в таблице результатов: выбирайте только те предикторы, у которых p < 0.05."
)
```

### Упражнение 4 - Прямой отбор предикторов

В этом упражнении мы попробуем подобрать значимые предикторы с помощью пошагового отбора и функции `ols_step_forward_p()` из пакета `olsrr`.

Метод **прямого отбора** начинается с модели без предикторов и постепенно добавляет каждую новую переменную, проверяя ее статистическую значимость

Создайте полную модель `mod_full`, а затем примените процедуру подбора предикторов. Сохраните результаты в `model_forward`.


```{r ex4-setup, include=FALSE}
mod_full <- life |>
  select(-c(Country, Status, Year)) |>
  lm(Life.expectancy ~ ., data = _)

model_forward <- olsrr::ols_step_forward_p(mod_full)
model_forward
```


```{r ex4, exercise = TRUE, eval=TRUE}
mod_full <- life |> 

model_forward <- olsrr::ols_step_forward_p(mod_full)
model_forward

```

```{r ex4-hint-1, eval = FALSE}
mod_full <- life |> 
  select() |> 
  lm(Life.expectancy ~ ., ... = _)

model_selected <- olsrr::ols_step_forward_p(mod_full)
model_selected
```


```{r ex4-solution, include = FALSE}
mod_full <- life |>
  select(-c(Country, Status, Year)) |>
  lm(Life.expectancy ~ ., data = _)

model_forward <- olsrr::ols_step_forward_p(mod_full)
model_forward
```


```{r ex4-code-check}
grade_code()
```

### Упражнение 5 - Графическая диагностика отбора

Теперь посмотрим на результаты `ols_step_forward_p()` графически.
Функция `plot()` для объекта `model_forward` позволяет визуально сравнить модели по различным метрикам.

```{r ex5, exercise = TRUE, eval=TRUE, exercise.setup="ex4-setup"}

```

```{r ex5-hint-1, eval = FALSE}
plot(model_forward)
```

```{r ex5-solution, include = FALSE}
plot(model_forward)
```

```{r ex5-code-check}
grade_code()
```

На графиках хорошо заметно, как увеличивается $R^2$ по мере добавления новых предикторов в модель, а информационный критерий Акаике AIC (само по себе значение этого критерия ничего не значит, но его снижение означает улучшение качества модели) и RMSE (Root mean squared error) -- снижаются.

Напомним, что RMSE —- это среднеквадратичная ошибка прогноза, мера «среднего» отклонения предсказанных значений от фактических в тех же единицах, что и зависимая переменная.

Для регрессии с наблюдениями $y_i$ и предсказаниями $\hat{y_i}$:

$$
\text{RMSE} = \sqrt{\frac{1}{n} \sum_{i=1}^{n} (y_i - \hat y_i)^2}
$$

RMSE показывает, насколько в среднем **величина ошибки** прогноза велика; это по сути стандартное отклонение остатков. Чем меньше RMSE, тем модель ближе к данным; RMSE = 0 означал бы идеальное совпадение (в реальных данных почти не встречается). Сравнивать RMSE имеет смысл для моделей на **одной и той же выборке** и в единицах измерения вашей целевой переменной (например, годы жизни, рубли, баллы шкалы). 


### Упражнение 6 - Метод обратного исключения

Метод обратного исключения начинается с полной модели и затем по очереди удаляет наименее статистически значимые переменные. 

Проведите аналогичный анализ, но только с помощью функции `ols_step_backward_p`. Выведите результаты в `model_backward`. 
Учитывайте, что полная модель уже создана и Вам нужно только отбирать предикторы.

```{r ex6-setup, include = FALSE}
mod_full <- life |>
  select(-c(Country, Status, Year)) |>
  lm(Life.expectancy ~ ., data = _)
```

```{r ex6, exercise = TRUE, eval=TRUE}
model_backward <- olsrr::ols_step_backward_p(mod_full, prem = 0.05)
model_backward

```

```{r ex6-hint-1, eval = FALSE}
model_backward <- olsrr::ols_step_backward_p(mod_full)
model_backward

```


```{r ex6-solution, include = FALSE}
model_backward <- olsrr::ols_step_backward_p(mod_full)
model_backward
```


```{r ex6-code-check}
grade_code()
```

Обратный отбор почти ничего не «выкинул» по сути: AIC и $R^2$ практически не меняются, значит все удалённые предикторы были слабо информативны и модель уже была довольно компактной. 

AIC последовательно чуть уменьшается (с 8921.24 до 8914.26), то есть формально каждый шаг улучшает баланс «качество + простота модели».

$R^2$ и $Adjusted R^2$ меняются только в третьем знаке (0.83474 → 0.83464, 0.83291 → 0.83322), то есть объясненная дисперсия почти не ухудшается, а скорректированный $R^2$ даже слегка растет из‑за уменьшения числа параметров.

Удаленные переменные (thinness 1–19 years, Population, GDP, Measles) давали минимальный вклад в объяснение, поэтому их исключение почти не сказалось на качестве предсказания, но чуть упростило модель. 



### Упражнение 7 — Сравнение моделей


Логично будет сравнить модели, полученные в ходе прямого отбора и обратного исключения.

Давайте сделаем сравнение по следующим показателям:

- набору отобранных предикторов - какие являются общими в двух моделях, какие отличаются
- основным метрикам - $Adj. R^2$, $AIC$, $BIC$

Шаг 1. Просто запустите код в ячейке:
```{r ex7-setup, include = FALSE}
mod_full <- life |>
  select(-c(Country, Status, Year)) |>
  lm(Life.expectancy ~ ., data = _)
model_backward <- olsrr::ols_step_backward_p(mod_full)
model_forward <- olsrr::ols_step_forward_p(mod_full)
```

```{r ex7, exercise = TRUE, eval=FALSE}
mod_back    <- model_backward$model
mod_forward <- model_forward$model

# Вытаскиваем из модели имена предикторов (без константы)
pred_back    <- setdiff(names(coef(mod_back)), "(Intercept)")
pred_forward <- setdiff(names(coef(mod_forward)), "(Intercept)")

# Отбираем уникальные предикторы из обеих моделей
all_preds <- sort(unique(c(pred_back, pred_forward)))

# Сводим все в датафрейм
pred_table <- data.frame(
  predictor = all_preds,
  in_backward = all_preds %in% pred_back,
  in_forward  = all_preds %in% pred_forward
)
# Проверим, входит ли предиктор в прямую или обратную модель
pred_table$in_backward <- ifelse(pred_table$in_backward, "Да", "Нет")
pred_table$in_forward  <- ifelse(pred_table$in_forward,  "Да", "Нет")

pred_table
```

```{r ex7-q}
question_radio(
    "У моделей одинаковый набор предикторов?",
    answer("Да", correct = TRUE),
    answer("Нет", correct = FALSE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```


Шаг 2. Теперь сравним по метрикам:

```{r ex7_2, exercise=TRUE, eval=FALSE}
compare_models <- tibble::tibble(
  model   = c("backward", "forward"),
  adj_R2  = c(summary(mod_back)$adj.r.squared,
              summary(mod_forward)$adj.r.squared),
  AIC     = c(AIC(mod_back), AIC(mod_forward)),
  BIC     = c(BIC(mod_back), BIC(mod_forward)),
  n_param = c(length(coef(mod_back)),
              length(coef(mod_forward)))
)
compare_models
```

Видим, что модели имеют идентичные значения метрик, что в целом указывает на стабильность модели (разные методы приводят к одинаковым результатам).

Еще один вариант отбора - комбинированный, сочетающий оба алгоритма - пошагового включения и исключения переменных. 

`ols_step_both_p()` в пакете `olsrr` реализует классический `stepwise`‑отбор предикторов по p‑значениям: на каждом шаге переменные могут как добавляться, так и удаляться из модели, пока не останется ни одной, которую можно ввести или исключить по заданным порогам.

На вход принимает полную модель `lm`, в которой уже перечислены все предикторы. Пошагово строит цепочку моделей, на каждом шаге:

- рассматривает кандидатов на включение (еще не вошедшие переменные);
- рассматривает кандидатов на исключение (уже включенные переменные);
- выбирает действие (добавить или удалить ту переменную, у которой p‑value наилучшим образом удовлетворяет порогам).

Останавливается, когда нет переменных, которые можно еще ввести (p_enter) или удалить (p_remove).

Давайте создадим финальную модель путем такого смешанного отбора. Запустите код и дайте свой комментарий:

```{r ex7_p_both, exercise=TRUE, eval=FALSE}
mod_final <- olsrr::ols_step_both_p(mod_full)
mod_final
```

```{r ex7_2-q}
question_text("Что Вы можете сказать по поводу финальной модели (например, какой предиктор имеет максимальное стандартизированное значение  Std. Beta?",
    answer(" ", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```

### Упражнение 8 — Диагностика модели

Итак, модель у нас уже есть, необходимо теперь проверить, соответствует ли она всем требованиям.

```{r ex8-setup}
mod_final <- lm(
  Life.expectancy ~ Schooling +
    HIV.AIDS +
    Adult.Mortality +
    Income.composition.of.resources +
    percentage.expenditure +
    BMI +
    Diphtheria +
    under.five.deaths +
    infant.deaths +
    thinness.5.9.years +
    Total.expenditure +
    Alcohol +
    Polio +
    Hepatitis.B,
  data = life
)

```


Постройте 4 диагностических графика для финальной модели множественной регрессии.

```{r ex8, exercise=TRUE, exercise.setup="ex8-setup"}
mod_final <- lm(
  Life.expectancy ~ Schooling +
    HIV.AIDS +
    Adult.Mortality +
    Income.composition.of.resources +
    percentage.expenditure +
    BMI +
    Diphtheria +
    under.five.deaths +
    infant.deaths +
    thinness.5.9.years +
    Total.expenditure +
    Alcohol +
    Polio +
    Hepatitis.B,
  data = life
)
par(mfrow = c(2, 2))
plot()
```

```{r ex8-hint-1, eval = FALSE}
mod_final <- lm(
  Life.expectancy ~ Schooling +
    HIV.AIDS +
    Adult.Mortality +
    Income.composition.of.resources +
    percentage.expenditure +
    BMI +
    Diphtheria +
    under.five.deaths +
    infant.deaths +
    thinness.5.9.years +
    Total.expenditure +
    Alcohol +
    Polio +
    Hepatitis.B,
  data = life
)
par(mfrow = c(2, 2))
plot(...)
```

```{r ex8-solution}
mod_final <- lm(
  Life.expectancy ~ Schooling +
    HIV.AIDS +
    Adult.Mortality +
    Income.composition.of.resources +
    percentage.expenditure +
    BMI +
    Diphtheria +
    under.five.deaths +
    infant.deaths +
    thinness.5.9.years +
    Total.expenditure +
    Alcohol +
    Polio +
    Hepatitis.B,
  data = life
)
par(mfrow = c(2, 2))
plot(mod_final)
```
```{r ex8-code-check}
grade_code()
```

- Residuals vs Fitted: остатки в среднем вокруг 0, но есть легкая нелинейность и разная ширина «ленты» по оси X — возможная слабая нелинейность и гетероскедастичность.

- Q–Q Residuals: хвосты заметно отклоняются от прямой, нормальность остатков нарушена, особенно в нижнем и верхнем хвостах.

- Scale–Location: разброс стандартизированных остатков слегка меняется с fitted, что поддерживает идею умеренной гетероскедастичности.

- Residuals vs Leverage: есть несколько наблюдений с повышенным leverage и Cook’s distance близко к порогу — отдельные страны/годы сильно влияют на оценки и их стоит отдельно проверить/задокументировать.

По-хорошему, надо отсмотреть страны с большим расстоянием Кука и левериджа, возможно, убрать какие-то периоды, и все перепроверить еще раз.

### Упражнение 9 — Мультиколлинеарность (VIF)

Проверьте VIF для предикторов модели.

```{r ex-vif, exercise=TRUE, exercise.setup="ex8-setup"}
library(car)
vif(___)
```

```{r ex-vif-solution}
library(car)
vif(mod_final)
```

---

```{r q-vif, echo=FALSE}
question_radio(
  "При каком значении VIF мультиколлинеарность считается проблематичной?",
  answer("VIF > 1"),
  answer("VIF > 5 (умеренная) или VIF > 10 (серьёзная)", correct = TRUE),
  answer("VIF > 0.9"),
  answer("VIF > 100"),
  allow_retry = TRUE
)
```


VIF показывает очень сильную мультиколлинеарность у переменных детской смертности:

- `under.five.deaths`: VIF ≈ 182.5  
- `infant.deaths`: VIF ≈ 181.9  

Это означает, что эти две переменные почти линейно зависят друг от друга (почти одно и то же измерение смертности детей), и их одновременное включение делает оценки коэффициентов нестабильными.


Кроме того, пусть это прозвучит немного цинично, если человек не умер в раннем возрасте, то у него есть больше шансов дожить до старости, это вещь очевидная. На самом деле и у высокой продолжительности жизни и у низкой детской смертности общие детерминанты, и включение этих переменных в модель дает только очередное подтверждение их взаимосвязи.


### Упражнение 10 — Тест на нормальность остатков

Проверьте нормальность остатков с помощью теста Шапиро–Уилка.

```{r ex10, exercise=TRUE, exercise.setup="ex8-setup"}
# Постройте гистограмму остатков финальной модели
hist(resid(mod_final),
     breaks = 30,
     main = "Гистограмма остатков",
     xlab = "Остатки")

```
```{r ex10-hint-1}
hist(resid(mod_final),
     breaks = 30,
     main = "...",
     xlab = "...")


```
```{r ex10-solution}
hist(resid(mod_final),
     breaks = 30,
     main = "Гистограмма остатков",
     xlab = "Остатки")

```
```{r ex10-q}
question_text(
  "Выглядт ли распределение остатков приблизительно симметричным и колоколовидным, или есть явная асимметрия и толстые хвосты?",
  answer(" ", correct = TRUE),
  allow_retry = TRUE,
  try_again = "Сравните форму гистограммы с классической нормальной кривой.",
  incorrect = NULL
)
```

### Упражнение 11 — Предсказание для новой страны

Используя `mod_final`, предскажите ожидаемую продолжительность жизни для страны с характеристиками:  

- Schooling = 12
- Adult.Mortality = 100
- HIV.AIDS = 0.5
- Income.composition.of.resources = 0.7
- percentage.expenditure = 200
- BMI = 25
- Diphtheria = 90
- under.five.deaths = 20
- infant.deaths = 15
- thinness.5.9.years = 5
- Total.expenditure = 6
- Alcohol = 5
- Polio = 90
- Hepatitis.B = 90

```{r ex-predict, exercise=TRUE, exercise.setup="ex8-setup"}
new_country <- data.frame(
  Schooling                      = 12,
  HIV.AIDS                       = 0.5,
  Adult.Mortality                = 100,
  Income.composition.of.resources = 0.7,
  percentage.expenditure         = 200,
  BMI                            = 25,
  Diphtheria                     = 90,
  under.five.deaths              = 20,
  infant.deaths                  = 15,
  thinness.5.9.years             = 5,
  Total.expenditure              = 6,
  Alcohol                        = 5,
  Polio                          = 90,
  Hepatitis.B                    = 90
)

predict(___, newdata = ___, interval = "confidence")
```

```{r ex-predict-solution}
new_country <- data.frame(
  Schooling                      = 12,
  HIV.AIDS                       = 0.5,
  Adult.Mortality                = 100,
  Income.composition.of.resources = 0.7,
  percentage.expenditure         = 200,
  BMI                            = 25,
  Diphtheria                     = 90,
  under.five.deaths              = 20,
  infant.deaths                  = 15,
  thinness.5.9.years             = 5,
  Total.expenditure              = 6,
  Alcohol                        = 5,
  Polio                          = 90,
  Hepatitis.B                    = 90
)

predict(mod_final, newdata = new_country, interval = "confidence")
```

```{r ex11-q}
question_numeric(
    "Предсказанное значение ожидаемой продолжительности жизни составляет _______лет (округлите до одной десятой",
    answer("71.5", correct = TRUE),
    allow_retry = TRUE,
    try_again = "Edit Answer",
    incorrect = NULL)
```


```{r download-answers, child ="child_documents/download_answers.Rmd"}
```
